i: [1,17918,0,67,23,1,0,0] t6_zf_lang
a: [3,18315,0,2,1,17918]
s: 18315
i: [1,637,0,2,1,1,0,0] t6_zf_lang
a: [3,30401,0,2,1,637]
s: 30401
i: [1,638,0,3,1,1,0,0] t6_zf_lang
a: [3,30396,0,3,1,638]
s: 30396
i: [1,17919,0,2,1,0,0,255] fc7_card_1
a: [3,30406,0,2,1,17919]
s: 30406
i: [1,17920,0,2,1,0,0,1] fc1_zf_lang
a: [3,30411,0,2,1,17920]
s: 30411
i: [1,17921,0,2,1,0,0,255] spc5_boole
a: [3,30416,0,2,1,17921]
s: 30416
i: [1,17922,0,2,1,0,0,4] spc3_boole
a: [3,30421,0,2,1,17922]
s: 30421
i: [1,17923,0,2,1,0,0,3] spc4_boole
a: [3,30426,0,2,1,17923]
s: 30426
i: [1,17924,0,2,1,0,0,4] spc2_boole
a: [3,30431,0,2,1,17924]
s: 30431
i: [1,17925,0,2,1,0,0,4] spc1_boole
a: [3,30436,0,2,1,17925]
s: 30436
i: [1,17926,0,2,1,0,0,255] fc4_card_1
a: [3,30441,0,2,1,17926]
s: 30441
i: [1,17927,0,2,1,0,0,255] fc1_xboole_0
a: [3,30446,0,2,1,17927]
s: 30446
i: [1,17980,0,2,1,0,0,255] fc6_ordinal1
a: [3,30702,0,2,1,17980]
s: 30702
i: [1,17979,0,2,1,0,0,255] fc6_ordinal1
a: [3,30707,0,2,1,17979]
s: 30707
i: [1,17981,0,2,1,0,0,255] spc5_numerals
a: [3,30727,0,2,1,17981]
s: 30727
i: [1,17985,0,2,1,0,0,255] spc4_numerals
a: [3,30747,0,2,1,17985]
s: 30747
i: [1,17989,0,2,1,0,0,255] spc3_numerals
a: [3,30767,0,2,1,17989]
s: 30767
i: [1,17993,0,2,1,0,0,255] spc2_numerals
a: [3,30787,0,2,1,17993]
s: 30787
i: [1,17997,0,2,1,0,0,255] spc1_numerals
a: [3,30807,0,2,1,17997]
s: 30807
i: [1,18162,0,2,1,0,0,255] rc9_finseq_1
a: [3,30812,0,2,1,18162]
s: 30812
i: [1,18161,0,2,1,0,0,255] rc9_finseq_1
a: [3,30817,0,2,1,18161]
s: 30817
i: [1,18160,0,2,1,0,0,255] rc9_finseq_1
a: [3,30822,0,2,1,18160]
s: 30822
i: [1,18159,0,2,1,0,0,255] rc9_finseq_1
a: [3,30827,0,2,1,18159]
s: 30827
i: [1,18158,0,2,1,0,0,255] rc9_finseq_1
a: [3,30832,0,2,1,18158]
s: 30832
i: [1,18164,0,2,1,0,0,255] rc3_ordinal1
a: [3,30837,0,2,1,18164]
s: 30837
i: [1,18163,0,2,1,0,0,255] rc3_ordinal1
a: [3,30842,0,2,1,18163]
s: 30842
i: [1,18198,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,30846,0,2,1,18198]
s: 30846
i: [1,18257,0,2,1,0,0,255] rc3_card_1
a: [3,30854,0,2,1,18257]
s: 30854
i: [1,18258,0,2,1,0,0,2] rc2_xboole_0
a: [3,30859,0,2,1,18258]
s: 30859
i: [1,18259,0,2,1,0,0,255] rc1_card_1
a: [3,30864,0,2,1,18259]
s: 30864
i: [1,18260,0,2,1,0,0,255] rc4_ordinal1
a: [3,30869,0,2,1,18260]
s: 30869
i: [1,18261,0,2,1,0,0,5] rc1_xreal_0
a: [3,30874,0,2,1,18261]
s: 30874
i: [1,18262,0,2,1,0,0,255] rc1_ordinal1
a: [3,30879,0,2,1,18262]
s: 30879
i: [1,18263,0,2,1,0,0,2] rc1_xboole_0
a: [3,30884,0,2,1,18263]
s: 30884
i: [1,18267,0,2,1,0,0,255] rc2_ordinal1
a: [3,30889,0,2,1,18267]
s: 30889
i: [1,18266,0,2,1,0,0,255] rc2_ordinal1
a: [3,30894,0,2,1,18266]
s: 30894
i: [1,18265,0,2,1,0,0,255] rc2_ordinal1
a: [3,30899,0,2,1,18265]
s: 30899
i: [1,18264,0,2,1,0,0,255] rc2_ordinal1
a: [3,30904,0,2,1,18264]
s: 30904
i: [1,18269,0,2,1,0,0,255] rc8_finseq_1
a: [3,30909,0,2,1,18269]
s: 30909
i: [1,18268,0,2,1,0,0,255] rc8_finseq_1
a: [3,30914,0,2,1,18268]
s: 30914
i: [1,18271,0,2,1,0,0,255] rc5_ordinal1
a: [3,30919,0,2,1,18271]
s: 30919
i: [1,18270,0,2,1,0,0,255] rc5_ordinal1
a: [3,30924,0,2,1,18270]
s: 30924
i: [1,18276,0,2,1,0,0,255] rc2_card_1
a: [3,30929,0,2,1,18276]
s: 30929
i: [1,18275,0,2,1,0,0,255] rc2_card_1
a: [3,30934,0,2,1,18275]
s: 30934
i: [1,18274,0,2,1,0,0,255] rc2_card_1
a: [3,30939,0,2,1,18274]
s: 30939
i: [1,18273,0,2,1,0,0,255] rc2_card_1
a: [3,30944,0,2,1,18273]
s: 30944
i: [1,18272,0,2,1,0,0,255] rc2_card_1
a: [3,30949,0,2,1,18272]
s: 30949
i: [1,18280,0,2,1,0,0,255] rc3_xreal_0
a: [3,30954,0,2,1,18280]
s: 30954
i: [1,18279,0,2,1,0,0,255] rc3_xreal_0
a: [3,30959,0,2,1,18279]
s: 30959
i: [1,18278,0,2,1,0,0,255] rc3_xreal_0
a: [3,30964,0,2,1,18278]
s: 30964
d: [2,31091,0,0,0,2,21,30964,30846]
e: 31091
r: 30846
s: 30964
i: [1,18277,0,2,1,0,0,255] rc3_xreal_0
a: [3,30969,0,2,1,18277]
s: 30969
i: [1,18284,0,2,1,0,0,255] rc2_xreal_0
a: [3,30974,0,2,1,18284]
s: 30974
i: [1,18283,0,2,1,0,0,255] rc2_xreal_0
a: [3,30979,0,2,1,18283]
s: 30979
i: [1,18282,0,2,1,0,0,255] rc2_xreal_0
a: [3,30984,0,2,1,18282]
s: 30984
i: [1,18281,0,2,1,0,0,255] rc2_xreal_0
a: [3,30989,0,2,1,18281]
s: 30989
i: [1,18287,0,2,1,0,0,255] rc3_finseq_1
a: [3,30994,0,2,1,18287]
s: 30994
i: [1,18286,0,2,1,0,0,255] rc3_finseq_1
a: [3,30999,0,2,1,18286]
s: 30999
i: [1,18285,0,2,1,0,0,255] rc3_finseq_1
a: [3,31004,0,2,1,18285]
s: 31004
i: [1,18293,0,2,1,0,0,255] rc4_finseq_1
a: [3,31009,0,2,1,18293]
s: 31009
i: [1,18292,0,2,1,0,0,255] rc4_finseq_1
a: [3,31014,0,2,1,18292]
s: 31014
i: [1,18291,0,2,1,0,0,255] rc4_finseq_1
a: [3,31019,0,2,1,18291]
s: 31019
i: [1,18290,0,2,1,0,0,255] rc4_finseq_1
a: [3,31024,0,2,1,18290]
s: 31024
i: [1,18289,0,2,1,0,0,255] rc4_finseq_1
a: [3,31029,0,2,1,18289]
s: 31029
i: [1,18288,0,2,1,0,0,255] rc4_finseq_1
a: [3,31034,0,2,1,18288]
s: 31034
i: [1,18297,0,2,1,0,0,255] rc4_xreal_0
a: [3,31039,0,2,1,18297]
s: 31039
i: [1,18296,0,2,1,0,0,255] rc4_xreal_0
a: [3,31044,0,2,1,18296]
s: 31044
i: [1,18295,0,2,1,0,0,255] rc4_xreal_0
a: [3,31049,0,2,1,18295]
s: 31049
i: [1,18294,0,2,1,0,0,255] rc4_xreal_0
a: [3,31054,0,2,1,18294]
s: 31054
i: [1,18304,0,2,1,0,0,255] rc1_zf_lang
a: [3,31059,0,2,1,18304]
s: 31059
i: [1,18303,0,2,1,0,0,255] rc1_zf_lang
a: [3,31064,0,2,1,18303]
s: 31064
i: [1,18302,0,2,1,0,0,255] rc1_zf_lang
a: [3,31069,0,2,1,18302]
s: 31069
i: [1,18301,0,2,1,0,0,255] rc1_zf_lang
a: [3,31074,0,2,1,18301]
s: 31074
i: [1,18300,0,2,1,0,0,255] rc1_zf_lang
a: [3,31079,0,2,1,18300]
s: 31079
i: [1,18299,0,2,1,0,0,255] rc1_zf_lang
a: [3,31084,0,2,1,18299]
s: 31084
i: [1,18001,0,3,1,0,0,4] fc16_finseq_1
s: 18001
i: [1,18002,0,3,1,0,0,2] fc1_subset_1
s: 18002
i: [1,18003,0,3,1,0,0,4] fc7_finseq_1
s: 18003
i: [1,18005,0,3,1,0,0,4] fc6_finseq_1
s: 18005
i: [1,18006,0,3,1,0,0,4] fc6_finseq_1
s: 18006
i: [1,18166,0,3,1,0,0,255] rc2_finseq_1
s: 18166
i: [1,18167,0,3,1,0,0,255] rc2_finseq_1
s: 18167
i: [1,18168,0,3,1,0,0,255] rc2_finseq_1
s: 18168
i: [1,18169,0,3,1,0,0,255] rc2_finseq_1
s: 18169
i: [1,18170,0,3,1,0,0,255] rc2_finseq_1
s: 18170
i: [1,18172,0,3,1,0,0,2] rc2_subset_1
s: 18172
i: [1,18176,0,3,1,0,0,255] rc1_finseq_1
s: 18176
i: [1,18177,0,3,1,0,0,255] rc1_finseq_1
s: 18177
i: [1,18178,0,3,1,0,0,255] rc1_finseq_1
s: 18178
i: [1,18179,0,3,1,0,0,255] reflexivity_r1_tarski
s: 18179
i: [1,17928,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r5_r3
a: [3,30451,0,3,1,17928]
s: 30451
i: [1,17929,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r5_r4
a: [3,30456,0,3,1,17929]
s: 30456
i: [1,17930,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r5_r2
a: [3,30461,0,3,1,17930]
s: 30461
i: [1,17931,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r5_r1
a: [3,30466,0,3,1,17931]
s: 30466
i: [1,17932,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r4_r3
a: [3,30471,0,3,1,17932]
s: 30471
i: [1,17933,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r4_r2
a: [3,30476,0,3,1,17933]
s: 30476
i: [1,17934,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r3_r2
a: [3,30481,0,3,1,17934]
s: 30481
i: [1,17935,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r3_r1
a: [3,30486,0,3,1,17935]
s: 30486
i: [1,17936,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r4_r1
a: [3,30491,0,3,1,17936]
s: 30491
i: [1,17937,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r2_r1
a: [3,30496,0,3,1,17937]
s: 30496
i: [1,17938,0,3,1,0,0,255] d1_zf_lang
a: [3,30501,0,3,1,17938]
s: 30501
i: [1,17939,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,30506,0,3,1,17939]
s: 30506
r: 30707
r: 30702
r: 30441
r: 30406
d: [2,31093,0,2,1,2,23,30707,30506]
a: [3,31100,0,2,1,31093]
s: 31100
d: [2,31094,0,2,1,2,23,30702,30506]
a: [3,31105,0,2,1,31094]
s: 31105
d: [2,31095,0,2,1,2,23,30441,30506]
a: [3,31110,0,2,1,31095]
s: 31110
d: [2,31096,0,2,1,2,23,30406,30506]
a: [3,31115,0,2,1,31096]
s: 31115
i: [1,17940,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r5_r5
a: [3,30511,0,3,1,17940]
s: 30511
i: [1,17941,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r4_r5
a: [3,30516,0,3,1,17941]
s: 30516
i: [1,17942,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r3_r5
a: [3,30521,0,3,1,17942]
s: 30521
i: [1,17943,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r3_r3
a: [3,30526,0,3,1,17943]
s: 30526
i: [1,17944,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r3_r4
a: [3,30531,0,3,1,17944]
s: 30531
i: [1,17945,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r4_r4
a: [3,30536,0,3,1,17945]
s: 30536
i: [1,17946,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r2_r5
a: [3,30541,0,3,1,17946]
s: 30541
i: [1,17947,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r2_r4
a: [3,30546,0,3,1,17947]
s: 30546
i: [1,17948,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r2_r3
a: [3,30551,0,3,1,17948]
s: 30551
i: [1,17949,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r2_r2
a: [3,30556,0,3,1,17949]
s: 30556
i: [1,17950,0,3,1,0,0,255] rqLessOrEqual__r1_xxreal_0__r1_r5
a: [3,30561,0,3,1,17950]
s: 30561
i: [1,17951,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r1_r3
a: [3,30566,0,3,1,17951]
s: 30566
i: [1,17952,0,3,1,0,0,3] rqLessOrEqual__r1_xxreal_0__r1_r4
a: [3,30571,0,3,1,17952]
s: 30571
i: [1,17953,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r1_r2
a: [3,30576,0,3,1,17953]
s: 30576
i: [1,17954,0,3,1,0,0,4] rqLessOrEqual__r1_xxreal_0__r1_r1
a: [3,30581,0,3,1,17954]
s: 30581
i: [1,17955,0,3,1,0,0,255] t1_numerals
a: [3,30586,0,3,1,17955]
s: 30586
i: [1,17984,0,3,1,0,0,255] spc5_numerals
a: [3,30712,0,3,1,17984]
s: 30712
i: [1,17983,0,3,1,0,0,255] spc5_numerals
a: [3,30717,0,3,1,17983]
s: 30717
i: [1,17988,0,3,1,0,0,255] spc4_numerals
a: [3,30732,0,3,1,17988]
s: 30732
i: [1,17987,0,3,1,0,0,255] spc4_numerals
a: [3,30737,0,3,1,17987]
s: 30737
i: [1,17992,0,3,1,0,0,255] spc3_numerals
a: [3,30752,0,3,1,17992]
s: 30752
i: [1,17991,0,3,1,0,0,255] spc3_numerals
a: [3,30757,0,3,1,17991]
s: 30757
i: [1,17996,0,3,1,0,0,255] spc2_numerals
a: [3,30772,0,3,1,17996]
s: 30772
i: [1,17995,0,3,1,0,0,255] spc2_numerals
a: [3,30777,0,3,1,17995]
s: 30777
i: [1,18000,0,3,1,0,0,255] spc1_numerals
a: [3,30792,0,3,1,18000]
s: 30792
i: [1,17999,0,3,1,0,0,255] spc1_numerals
a: [3,30797,0,3,1,17999]
s: 30797
i: [1,18298,0,3,1,0,0,255] rc1_zf_lang
a: [3,31089,0,3,1,18298]
s: 31089
d: [2,31117,0,3,1,2,22,30586,30506]
a: [3,31121,0,3,1,31117]
s: 31121
a: [3,18333,0,14,4,17918]
s: 18333
a: [3,18312,0,15,5,17918]
s: 18312
i: [1,18004,0,4,1,0,0,4] fc24_finseq_1
s: 18004
i: [1,18035,0,4,2,0,0,255] cc1_card_1
s: 18035
d: [2,31124,1,2,1,1,40,18035,30864]
a: [3,31129,1,2,1,31124]
s: 31129
i: [1,18040,0,4,2,0,0,255] cc3_card_1
s: 18040
i: [1,18041,0,4,2,0,0,255] cc2_xreal_0
s: 18041
d: [2,31131,1,2,1,1,40,18040,30869]
a: [3,31136,1,2,1,31131]
s: 31136
i: [1,18042,0,4,2,0,0,255] cc6_ordinal1
s: 18042
d: [2,31132,1,2,1,1,40,18040,30919]
a: [3,31141,1,2,1,31132]
s: 31141
i: [1,18043,0,4,2,0,0,255] cc5_card_1
s: 18043
i: [1,18050,0,4,2,0,0,255] cc4_xreal_0
s: 18050
d: [2,31143,1,2,1,1,40,18041,30869]
a: [3,31148,1,2,1,31143]
s: 31148
i: [1,18051,0,4,2,0,0,255] cc3_xreal_0
s: 18051
d: [2,31144,1,2,1,1,40,18041,30919]
a: [3,31153,1,2,1,31144]
s: 31153
i: [1,18060,0,4,2,0,0,255] cc1_ordinal1
s: 18060
i: [1,18061,0,4,2,0,0,255] cc1_ordinal1
s: 18061
d: [2,31162,1,2,1,1,40,18042,30919]
a: [3,31167,1,2,1,31162]
s: 31167
i: [1,18063,0,4,2,0,0,255] cc9_ordinal1
s: 18063
d: [2,31171,1,2,1,1,40,18043,30869]
a: [3,31176,1,2,1,31171]
s: 31176
i: [1,18064,0,4,2,0,0,255] cc8_finseq_1
s: 18064
i: [1,18065,0,4,2,0,0,255] cc2_card_1
s: 18065
d: [2,31172,1,2,1,1,40,18043,30919]
a: [3,31181,1,2,1,31172]
s: 31181
i: [1,18066,0,4,2,0,0,255] cc7_ordinal1
s: 18066
d: [2,31183,1,2,1,1,40,18050,30874]
a: [3,31190,1,2,1,31183]
s: 31190
i: [1,18067,0,4,2,0,0,255] cc3_ordinal1
s: 18067
i: [1,18155,0,4,1,0,0,2] existence_m1_finseq_1
s: 18155
d: [2,31199,1,2,1,1,40,18051,30874]
a: [3,31211,1,2,1,31199]
s: 31211
i: [1,18156,0,4,1,0,0,1] existence_m2_finseq_1
s: 18156
d: [2,31225,1,2,1,1,40,18060,31105]
a: [3,31233,1,2,1,31225]
s: 31233
i: [1,18157,0,4,1,0,0,2] existence_m1_subset_1
s: 18157
i: [1,18165,0,4,1,0,0,255] rc2_finseq_1
s: 18165
d: [2,31227,1,2,1,1,40,18060,30879]
a: [3,31243,1,2,1,31227]
s: 31243
i: [1,18174,0,4,1,0,0,255] rc3_subset_1
s: 18174
d: [2,31245,1,2,1,1,40,18061,31105]
a: [3,31253,1,2,1,31245]
s: 31253
i: [1,18246,0,4,1,0,0,255] rc1_relset_1
s: 18246
i: [1,18247,0,4,1,0,0,255] rc1_relset_1
s: 18247
d: [2,31247,1,2,1,1,40,18061,30879]
a: [3,31263,1,2,1,31247]
s: 31263
i: [1,17956,0,4,1,0,0,1] dt_k1_zf_lang
a: [3,30591,0,4,1,17956]
s: 30591
d: [2,31277,1,2,1,1,40,18063,30446]
a: [3,31285,1,2,1,31277]
s: 31285
i: [1,17958,0,4,1,0,0,255] dt_k5_numbers
a: [3,30597,0,4,1,17958]
s: 30597
i: [1,17982,0,4,1,0,0,255] spc5_numerals
a: [3,30722,0,4,1,17982]
s: 30722
d: [2,31280,1,2,1,1,40,18063,30884]
a: [3,31290,1,2,1,31280]
s: 31290
i: [1,17986,0,4,1,0,0,255] spc4_numerals
a: [3,30742,0,4,1,17986]
s: 30742
d: [2,31281,1,2,1,1,40,18063,31054]
a: [3,31295,1,2,1,31281]
s: 31295
i: [1,17990,0,4,1,0,0,255] spc3_numerals
a: [3,30762,0,4,1,17990]
s: 30762
i: [1,17994,0,4,1,0,0,255] spc2_numerals
a: [3,30782,0,4,1,17994]
s: 30782
d: [2,31297,1,2,1,1,40,18064,30446]
a: [3,31305,1,2,1,31297]
s: 31305
i: [1,17998,0,4,1,0,0,255] spc1_numerals
a: [3,30802,0,4,1,17998]
s: 30802
d: [2,31300,1,2,1,1,40,18064,30884]
a: [3,31310,1,2,1,31300]
s: 31310
i: [1,18008,0,5,2,0,0,255] rc6_subset_1
s: 18008
i: [1,18009,0,5,2,0,0,255] fc9_card_1
s: 18009
d: [2,31301,1,2,1,1,40,18064,31054]
a: [3,31315,1,2,1,31301]
s: 31315
i: [1,18011,0,5,2,0,0,255] rc4_card_1
s: 18011
d: [2,31317,1,2,1,1,40,18065,30446]
a: [3,31325,1,2,1,31317]
s: 31325
i: [1,18015,0,5,2,0,0,255] rc6_finseq_1
s: 18015
i: [1,18016,0,5,2,0,0,255] rc6_finseq_1
s: 18016
d: [2,31320,1,2,1,1,40,18065,30884]
a: [3,31330,1,2,1,31320]
s: 31330
i: [1,18017,0,5,2,0,0,255] rc6_finseq_1
s: 18017
d: [2,31321,1,2,1,1,40,18065,31054]
a: [3,31335,1,2,1,31321]
s: 31335
i: [1,18018,0,5,2,0,0,255] rc6_finseq_1
s: 18018
i: [1,18019,0,5,2,0,0,255] rc6_finseq_1
s: 18019
d: [2,31337,1,2,1,1,40,18066,30446]
a: [3,31345,1,2,1,31337]
s: 31345
i: [1,18020,0,5,2,0,0,255] rc6_finseq_1
s: 18020
d: [2,31340,1,2,1,1,40,18066,30884]
a: [3,31350,1,2,1,31340]
s: 31350
i: [1,18022,0,5,2,0,0,2] rc1_subset_1
s: 18022
i: [1,18024,0,5,2,0,0,255] rc5_subset_1
s: 18024
d: [2,31341,1,2,1,1,40,18066,31054]
a: [3,31355,1,2,1,31341]
s: 31355
i: [1,18025,0,5,2,0,0,255] rc5_subset_1
s: 18025
d: [2,31357,1,2,1,1,40,18067,30446]
a: [3,31365,1,2,1,31357]
s: 31365
i: [1,18037,0,5,2,0,0,255] rc6_card_1
s: 18037
i: [1,18038,0,5,2,0,0,255] rc6_card_1
s: 18038
d: [2,31360,1,2,1,1,40,18067,30884]
a: [3,31370,1,2,1,31360]
s: 31370
i: [1,18044,0,5,2,0,0,255] rc7_finseq_1
s: 18044
d: [2,31361,1,2,1,1,40,18067,31054]
a: [3,31375,1,2,1,31361]
s: 31375
i: [1,18045,0,5,2,0,0,255] rc7_finseq_1
s: 18045
i: [1,18046,0,5,2,0,0,255] rc7_finseq_1
s: 18046
d: [2,31155,2,2,1,1,40,31136,18035]
a: [3,31159,2,2,1,31155]
s: 31159
i: [1,18048,0,5,2,0,0,255] rc7_finseq_1
s: 18048
d: [2,31192,2,2,1,1,40,31148,18050]
a: [3,31196,2,2,1,31192]
s: 31196
i: [1,18049,0,5,2,0,0,255] rc7_finseq_1
s: 18049
i: [1,18068,0,5,2,0,0,255] t6_boole
s: 18068
d: [2,31198,2,2,1,1,40,18051,31148]
a: [3,31206,2,2,1,31198]
s: 31206
i: [1,18069,0,5,2,0,0,255] cc8_card_1
s: 18069
d: [2,31213,2,2,1,1,40,31153,18051]
a: [3,31218,2,2,1,31213]
s: 31218
i: [1,18074,0,5,2,0,0,255] cc7_card_1
s: 18074
i: [1,18075,0,5,2,0,0,255] cc9_card_1
s: 18075
d: [2,31214,2,2,1,1,40,31153,18050]
a: [3,31223,2,2,1,31214]
s: 31223
i: [1,18076,0,5,2,0,0,255] cc9_card_1
s: 18076
d: [2,31226,2,2,1,1,40,18060,31129]
a: [3,31238,2,2,1,31226]
s: 31238
i: [1,18079,0,5,2,0,0,255] cc8_ordinal1
s: 18079
i: [1,18080,0,5,2,0,0,255] cc4_card_1
s: 18080
d: [2,31246,2,2,1,1,40,18061,31129]
a: [3,31258,2,2,1,31246]
s: 31258
i: [1,18081,0,5,2,0,0,255] cc1_xreal_0
s: 18081
d: [2,31491,1,2,1,1,40,18081,30792]
a: [3,31500,1,2,1,31491]
s: 31500
i: [1,18139,0,5,2,0,0,5] dt_m1_finseq_1
s: 18139
i: [1,18140,0,5,2,0,0,5] dt_m1_finseq_1
s: 18140
d: [2,31492,1,2,1,1,40,18081,30772]
a: [3,31505,1,2,1,31492]
s: 31505
i: [1,18141,0,5,2,0,0,5] dt_m1_finseq_1
s: 18141
d: [2,31493,1,2,1,1,40,18081,30752]
a: [3,31510,1,2,1,31493]
s: 31510
i: [1,18142,0,5,2,0,0,255] dt_m2_finseq_1
s: 18142
i: [1,18143,0,5,2,0,0,255] dt_m2_finseq_1
s: 18143
d: [2,31572,1,2,1,1,40,18142,30396]
a: [3,31577,1,2,1,31572]
s: 31577
i: [1,18150,0,5,2,0,0,255] t7_boole
s: 18150
d: [2,31579,1,2,1,1,40,18143,30396]
a: [3,31584,1,2,1,31579]
s: 31584
i: [1,18171,0,5,1,0,0,2] rc2_subset_1
s: 18171
i: [1,18173,0,5,1,0,0,255] rc3_subset_1
s: 18173
d: [2,31494,1,2,1,1,40,18081,30732]
a: [3,31515,1,2,1,31494]
s: 31515
i: [1,17959,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r5_r1_r5
a: [3,30602,0,5,1,17959]
s: 30602
d: [2,31495,1,2,1,1,40,18081,30712]
a: [3,31520,1,2,1,31495]
s: 31520
i: [1,17960,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r5_r5
a: [3,30607,0,5,1,17960]
s: 30607
i: [1,17961,0,5,1,0,0,255] rqRealAdd__k2_xcmplx_0__r3_r2_r5
a: [3,30612,0,5,1,17961]
s: 30612
d: [2,31265,2,2,1,1,40,31167,18061]
a: [3,31270,2,2,1,31265]
s: 31270
i: [1,17962,0,5,1,0,0,255] rqRealAdd__k2_xcmplx_0__r4_r1_r5
a: [3,30617,0,5,1,17962]
s: 30617
d: [2,31266,2,2,1,1,40,31167,18060]
a: [3,31275,2,2,1,31266]
s: 31275
i: [1,17963,0,5,1,0,0,255] rqRealAdd__k2_xcmplx_0__r2_r3_r5
a: [3,30622,0,5,1,17963]
s: 30622
i: [1,17964,0,5,1,0,0,255] rqRealAdd__k2_xcmplx_0__r1_r4_r5
a: [3,30627,0,5,1,17964]
s: 30627
d: [2,31388,2,2,1,1,40,31345,18043]
a: [3,31395,2,2,1,31388]
s: 31395
i: [1,17965,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r3_r1_r3
a: [3,30632,0,5,1,17965]
s: 30632
d: [2,31390,2,2,1,1,40,31345,18041]
a: [3,31401,2,2,1,31390]
s: 31401
i: [1,17966,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r4_r1_r4
a: [3,30637,0,5,1,17966]
s: 30637
i: [1,17967,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r2_r2_r4
a: [3,30642,0,5,1,17967]
s: 30642
d: [2,31403,2,2,1,1,40,31350,18043]
a: [3,31410,2,2,1,31403]
s: 31410
i: [1,17968,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r4_r4
a: [3,30647,0,5,1,17968]
s: 30647
d: [2,31405,2,2,1,1,40,31350,18041]
a: [3,31416,2,2,1,31405]
s: 31416
i: [1,17969,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r3_r3
a: [3,30652,0,5,1,17969]
s: 30652
i: [1,17970,0,5,1,0,0,3] rqRealAdd__k2_xcmplx_0__r3_r1_r4
a: [3,30657,0,5,1,17970]
s: 30657
d: [2,31418,2,2,1,1,40,31355,18043]
a: [3,31425,2,2,1,31418]
s: 31425
i: [1,17971,0,5,1,0,0,3] rqRealAdd__k2_xcmplx_0__r2_r2_r4
a: [3,30662,0,5,1,17971]
s: 30662
d: [2,31428,2,2,1,1,40,31365,18061]
a: [3,31433,2,2,1,31428]
s: 31433
i: [1,17972,0,5,1,0,0,4] rqRealAdd__k2_xcmplx_0__r2_r1_r3
a: [3,30667,0,5,1,17972]
s: 30667
i: [1,17973,0,5,1,0,0,3] rqRealAdd__k2_xcmplx_0__r1_r3_r4
a: [3,30672,0,5,1,17973]
s: 30672
d: [2,31429,2,2,1,1,40,31365,18060]
a: [3,31438,2,2,1,31429]
s: 31438
i: [1,17974,0,5,1,0,0,4] rqRealAdd__k2_xcmplx_0__r1_r2_r3
a: [3,30677,0,5,1,17974]
s: 30677
d: [2,31440,2,2,1,1,40,31370,18061]
a: [3,31445,2,2,1,31440]
s: 31445
i: [1,17975,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r2_r1_r2
a: [3,30682,0,5,1,17975]
s: 30682
i: [1,17976,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r2_r2
a: [3,30687,0,5,1,17976]
s: 30687
d: [2,31441,2,2,1,1,40,31370,18060]
a: [3,31450,2,2,1,31441]
s: 31450
i: [1,17977,0,5,1,0,0,4] rqRealAdd__k2_xcmplx_0__r1_r1_r2
a: [3,30692,0,5,1,17977]
s: 30692
d: [2,31452,2,2,1,1,40,31375,18061]
a: [3,31457,2,2,1,31452]
s: 31457
i: [1,17978,0,5,1,0,0,255] rqRealMult__k3_xcmplx_0__r1_r1_r1
a: [3,30697,0,5,1,17978]
s: 30697
a: [3,30849,0,5,2,18198]
s: 30849
d: [2,31453,2,2,1,1,40,31375,18060]
a: [3,31462,2,2,1,31453]
s: 31462
d: [2,31489,0,5,2,1,22,18079,30506]
s: 31489
d: [2,31638,1,2,1,2,40,31489,30797]
a: [3,31648,1,2,1,31638]
s: 31648
d: [2,31490,0,5,2,1,22,18080,30506]
s: 31490
i: [1,18014,0,6,2,0,0,255] rc6_finseq_1
s: 18014
d: [2,31640,1,2,1,2,40,31489,30777]
a: [3,31653,1,2,1,31640]
s: 31653
i: [1,18027,0,6,2,0,0,5] rc7_card_1
s: 18027
d: [2,31641,1,2,1,2,40,31489,30757]
a: [3,31658,1,2,1,31641]
s: 31658
i: [1,18029,0,6,2,0,0,255] rc4_subset_1
s: 18029
i: [1,18036,0,6,2,0,0,255] rc5_card_1
s: 18036
d: [2,31642,1,2,1,2,40,31489,30737]
a: [3,31663,1,2,1,31642]
s: 31663
i: [1,18039,0,6,2,0,0,255] rc6_card_1
s: 18039
d: [2,31643,1,2,1,2,40,31489,30717]
a: [3,31668,1,2,1,31643]
s: 31668
i: [1,18047,0,6,2,0,0,255] rc7_finseq_1
s: 18047
i: [1,18089,0,6,3,0,0,255] cc2_ordinal1
s: 18089
d: [2,31696,1,2,1,2,40,31490,30777]
a: [3,31705,1,2,1,31696]
s: 31705
i: [1,18090,0,6,3,0,0,255] cc6_card_1
s: 18090
d: [2,31697,1,2,1,2,40,31490,30757]
a: [3,31710,1,2,1,31697]
s: 31710
i: [1,18131,0,6,3,0,0,5] cc4_finseq_1
s: 18131
i: [1,18133,0,6,3,0,0,5] cc1_finseq_1
s: 18133
d: [2,31698,1,2,1,2,40,31490,30737]
a: [3,31715,1,2,1,31698]
s: 31715
i: [1,18138,0,6,2,0,0,255] t1_subset
s: 18138
d: [2,31699,1,2,1,2,40,31490,30717]
a: [3,31720,1,2,1,31699]
s: 31720
i: [1,18146,0,6,2,0,0,2] redefinition_m2_finseq_1
s: 18146
i: [1,18147,0,6,2,0,0,2] redefinition_m2_finseq_1
s: 18147
d: [2,31943,1,3,1,1,40,18146,30396]
a: [3,31948,1,3,1,31943]
s: 31948
i: [1,18182,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 18182
d: [2,31962,2,2,1,1,40,31948,18139]
a: [3,31966,2,2,1,31962]
s: 31966
i: [1,18007,0,7,2,0,0,255] rc6_subset_1
s: 18007
i: [1,18010,0,7,2,0,0,255] rc4_card_1
s: 18010
d: [2,31864,1,4,2,1,40,18090,30822]
a: [3,31876,1,2,1,31864]
s: 31876
i: [1,18021,0,7,2,0,0,2] rc1_subset_1
s: 18021
d: [2,31869,1,4,2,1,40,18090,30934]
d: [2,31882,1,2,1,2,21,31869,30939]
a: [3,31886,1,2,1,31882]
s: 31886
i: [1,18023,0,7,2,0,0,255] rc5_subset_1
s: 18023
i: [1,18026,0,7,2,0,0,5] rc7_card_1
s: 18026
d: [2,31870,1,4,2,1,40,18090,31019]
a: [3,31891,1,2,1,31870]
s: 31891
i: [1,18028,0,7,2,0,0,255] rc4_subset_1
s: 18028
d: [2,31872,1,4,2,1,40,18090,31074]
a: [3,31900,1,2,1,31872]
s: 31900
i: [1,18031,0,7,3,0,0,255] cc10_finseq_1
s: 18031
i: [1,18034,0,7,2,0,0,255] t3_arithm
s: 18034
d: [2,31907,1,4,2,1,40,18131,30446]
a: [3,31916,1,2,1,31907]
s: 31916
i: [1,18062,0,7,3,0,0,255] cc5_ordinal1
s: 18062
d: [2,31911,1,4,2,1,40,18131,30884]
a: [3,31926,1,2,1,31911]
s: 31926
i: [1,18072,0,7,2,0,0,2] dt_k6_zf_lang
s: 18072
i: [1,18073,0,7,2,0,0,2] dt_k6_zf_lang
s: 18073
d: [2,31912,1,4,2,1,40,18131,31054]
a: [3,31935,1,2,1,31912]
s: 31935
i: [1,18082,0,7,2,0,0,3] dt_k3_zf_lang
s: 18082
d: [2,31979,1,4,2,1,40,18031,30792]
a: [3,32020,1,2,1,31979]
s: 32020
i: [1,18083,0,7,2,0,0,3] dt_k3_zf_lang
s: 18083
i: [1,18088,0,7,3,0,0,255] cc10_card_1
s: 18088
d: [2,31980,1,4,2,1,40,18031,30797]
a: [3,32029,1,2,1,31980]
s: 32029
i: [1,18134,0,7,2,0,0,255] fraenkel_a_0_0_zf_lang
s: 18134
d: [2,31527,2,2,1,1,40,31500,18051]
a: [3,31532,2,2,1,31527]
s: 31532
i: [1,18135,0,7,2,0,0,255] fraenkel_a_0_0_zf_lang
s: 18135
i: [1,18136,0,7,2,0,0,255] fraenkel_a_0_0_zf_lang
s: 18136
d: [2,31528,2,2,1,1,40,31500,18050]
a: [3,31537,2,2,1,31528]
s: 31537
i: [1,18148,0,7,2,0,0,255] t3_subset
s: 18148
d: [2,31545,2,2,1,1,40,31505,18051]
a: [3,31550,2,2,1,31545]
s: 31550
i: [1,18149,0,7,3,0,0,1] t8_boole
s: 18149
i: [1,18151,0,7,2,0,0,255] t2_tarski
s: 18151
d: [2,31546,2,2,1,1,40,31505,18050]
a: [3,31555,2,2,1,31546]
s: 31555
i: [1,18175,0,7,1,0,0,255] rc1_finseq_1
s: 18175
d: [2,31560,2,2,1,1,40,31510,18051]
a: [3,31565,2,2,1,31560]
s: 31565
d: [2,32115,0,7,2,1,22,18134,30501]
s: 32115
d: [2,32116,0,7,2,1,22,18135,30501]
s: 32116
d: [2,31561,2,2,1,1,40,31510,18050]
a: [3,31570,2,2,1,31561]
s: 31570
d: [2,32117,0,7,2,1,22,18136,30501]
s: 32117
d: [2,31586,2,2,1,1,40,31515,18051]
a: [3,31591,2,2,1,31586]
s: 31591
i: [1,18030,0,8,3,0,0,255] cc10_ordinal1
s: 18030
i: [1,18032,0,8,3,0,0,255] cc11_finseq_1
s: 18032
d: [2,32135,1,4,2,1,40,18030,30597]
a: [3,32150,1,2,1,32135]
s: 32150
i: [1,18033,0,8,3,0,0,255] cc5_subset_1
s: 18033
a: [3,32153,1,2,1,32135]
s: 32153
i: [1,18070,0,8,3,0,0,2] cc1_subset_1
s: 18070
i: [1,18084,0,8,2,0,0,3] redefinition_k3_zf_lang
s: 18084
d: [2,32175,1,4,2,1,40,18033,30597]
a: [3,32190,1,2,1,32175]
s: 32190
i: [1,18085,0,8,2,0,0,3] redefinition_k3_zf_lang
s: 18085
a: [3,32193,1,2,1,32175]
s: 32193
i: [1,18091,0,8,3,0,0,2] fc4_zf_lang
s: 18091
i: [1,18123,0,8,4,0,0,255] cc6_finseq_1
s: 18123
d: [2,32268,2,6,3,1,40,18123,31584]
d: [2,32278,2,4,2,2,21,32268,31577]
d: [2,32279,2,2,1,3,21,32278,31966]
a: [3,32283,2,2,1,32279]
s: 32283
i: [1,18126,0,8,4,0,0,255] cc2_finseq_1
s: 18126
d: [2,32290,2,6,3,1,40,18126,31584]
d: [2,32300,2,4,2,2,21,32290,31577]
d: [2,32301,2,2,1,3,21,32300,31966]
a: [3,32305,2,2,1,32301]
s: 32305
i: [1,18145,0,8,3,0,0,255] t2_subset
s: 18145
i: [1,18183,0,8,3,0,0,255] fc11_card_1
s: 18183
d: [2,32309,3,4,2,1,40,32305,18090]
a: [3,32313,3,2,1,32309]
s: 32313
i: [1,18184,0,8,3,0,0,255] fc10_card_1
s: 18184
d: [2,32222,2,5,2,1,40,18091,31948]
d: [2,32227,2,3,1,2,21,32222,30401]
a: [3,32231,2,3,1,32227]
s: 32231
i: [1,18191,0,8,3,0,0,255] fc10_subset_1
s: 18191
i: [1,18205,0,8,3,0,0,255] fc6_xreal_0
s: 18205
d: [2,32202,1,4,2,1,40,18070,30597]
d: [2,32214,1,2,1,2,21,32202,31100]
a: [3,32218,1,2,1,32214]
s: 32218
i: [1,18206,0,8,3,0,0,4] fc5_xreal_0
s: 18206
d: [2,31587,2,2,1,1,40,31515,18050]
a: [3,31596,2,2,1,31587]
s: 31596
i: [1,18244,0,8,2,0,0,255] cc1_relset_1
s: 18244
i: [1,18245,0,8,1,0,0,255] rc1_relset_1
s: 18245
d: [2,31598,2,2,1,1,40,31520,18051]
a: [3,31603,2,2,1,31598]
s: 31603
i: [1,18071,0,9,3,0,0,255] cc4_subset_1
s: 18071
d: [2,31599,2,2,1,1,40,31520,18050]
a: [3,31608,2,2,1,31599]
s: 31608
i: [1,18144,0,9,2,0,0,255] dt_m2_finseq_1
s: 18144
i: [1,18154,0,9,3,0,0,255] t5_subset
s: 18154
d: [2,31675,2,2,1,1,40,31648,18043]
a: [3,31682,2,2,1,31675]
s: 31682
i: [1,18309,0,9,3,0,0,255] fraenkel_a_0_0_zf_lang
s: 18309
d: [2,31676,2,2,1,1,40,31648,18042]
a: [3,31687,2,2,1,31676]
s: 31687
d: [2,32517,0,9,3,1,22,18309,30501]
s: 32517
i: [1,18092,0,10,4,0,0,1] d20_zf_lang
s: 18092
r: 18315
a: [3,18316,0,2,1,17918]
s: 18316
i: [1,7550,0,67,23,1,0,0] t6_zf_lang
a: [3,25064,0,2,1,7550]
s: 25064
i: [1,11006,0,68,23,1,0,0] t6_zf_lang
a: [3,22398,0,3,1,11006]
s: 22398
r: 22398
a: [3,18322,0,2,1,17918]
s: 18322
a: [3,22397,0,3,1,11006]
s: 22397
i: [1,4094,0,69,23,1,0,0] t6_zf_lang
a: [3,27730,0,4,1,4094]
s: 27730
i: [1,14462,0,70,23,1,0,0] t6_zf_lang
a: [3,19732,0,5,1,14462]
s: 19732
d: [2,31678,2,2,1,1,40,31648,18040]
a: [3,31692,2,2,1,31678]
s: 31692
a: [3,18319,0,15,5,17918]
s: 18319
a: [3,18326,0,15,5,17918]
s: 18326
d: [2,31731,2,2,1,1,40,31653,18042]
a: [3,31738,2,2,1,31731]
s: 31738
i: [1,18094,0,10,4,0,0,1] d20_zf_lang
s: 18094
d: [2,31733,2,2,1,1,40,31653,18040]
a: [3,31743,2,2,1,31733]
s: 31743
i: [1,18098,0,10,4,0,0,1] d22_zf_lang
s: 18098
a: [3,18330,0,2,1,17918]
s: 18330
d: [2,31748,2,2,1,1,40,31658,18042]
a: [3,31755,2,2,1,31748]
s: 31755
i: [1,18100,0,10,4,0,0,1] d22_zf_lang
s: 18100
d: [2,31750,2,2,1,1,40,31658,18040]
a: [3,31760,2,2,1,31750]
s: 31760
i: [1,18104,0,10,4,0,0,1] d21_zf_lang
s: 18104
r: 18322
a: [3,18323,0,2,1,17918]
s: 18323
a: [3,18336,0,2,1,17918]
s: 18336
i: [1,18106,0,10,4,0,0,1] d21_zf_lang
s: 18106
d: [2,31782,2,2,1,1,40,31663,18042]
a: [3,31789,2,2,1,31782]
s: 31789
i: [1,18111,0,10,4,0,0,1] d23_zf_lang
s: 18111
r: 18336
a: [3,18337,0,2,1,17918]
s: 18337
i: [1,17912,0,67,23,1,0,0] t6_zf_lang
a: [3,18354,0,2,1,17912]
s: 18354
i: [1,17908,0,67,23,1,0,0] t6_zf_lang
a: [3,18366,0,2,1,17908]
s: 18366
i: [1,17914,0,68,23,1,0,0] t6_zf_lang
a: [3,18348,0,3,1,17914]
s: 18348
r: 18348
i: [1,17846,0,67,23,1,0,0] t6_zf_lang
a: [3,18424,0,2,1,17846]
s: 18424
i: [1,17798,0,67,23,1,0,0] t6_zf_lang
a: [3,18476,0,2,1,17798]
s: 18476
a: [3,18347,0,3,1,17914]
s: 18347
i: [1,17910,0,68,23,1,0,0] t6_zf_lang
a: [3,18360,0,3,1,17910]
s: 18360
r: 18360
a: [3,18359,0,3,1,17910]
s: 18359
i: [1,17870,0,68,23,1,0,0] t6_zf_lang
a: [3,18398,0,3,1,17870]
s: 18398
r: 18316
r: 18398
r: 22397
r: 18333
s: 18316
d: [2,32569,1,7,3,1,40,22398,18092]
d: [2,32573,1,5,2,2,21,32569,25064]
a: [3,32581,1,2,1,32573]
s: 32581
a: [3,18397,0,3,1,17870]
s: 18397
d: [2,32571,1,2,1,1,40,22398,18143]
a: [3,32591,1,2,1,32571]
s: 32591
i: [1,17822,0,68,23,1,0,0] t6_zf_lang
a: [3,18450,0,3,1,17822]
s: 18450
r: 18450
a: [3,18449,0,3,1,17822]
s: 18449
d: [2,32572,1,2,1,1,40,22398,18142]
a: [3,32596,1,2,1,32572]
s: 32596
s: 22398
d: [2,33088,1,7,3,1,40,22398,18098]
d: [2,33096,1,5,2,2,21,33088,25064]
a: [3,33104,1,2,1,33096]
s: 33104
i: [1,17916,0,70,23,1,0,0] t6_zf_lang
a: [3,18342,0,5,1,17916]
s: 18342
i: [1,17894,0,70,23,1,0,0] t6_zf_lang
a: [3,18372,0,5,1,17894]
s: 18372
d: [2,33090,1,7,3,1,40,22398,18104]
d: [2,33112,1,5,2,2,21,33090,25064]
a: [3,33120,1,2,1,33112]
s: 33120
i: [1,18152,0,10,3,0,0,255] t2_tarski
s: 18152
d: [2,33092,1,7,3,1,40,22398,18111]
d: [2,33128,1,5,2,2,21,33092,25064]
a: [3,33136,1,2,1,33128]
s: 33136
i: [1,18153,0,10,3,0,0,255] t4_subset
s: 18153
i: [1,18180,0,10,3,0,0,255] cc4_relset_1
s: 18180
d: [2,32987,2,6,3,1,40,32591,18126]
a: [3,32992,2,2,1,32987]
s: 32992
i: [1,18181,0,10,3,0,0,255] cc3_relset_1
s: 18181
d: [2,32545,1,7,3,1,40,18092,30396]
d: [2,32547,1,5,2,2,21,32545,18315]
d: [2,32548,1,3,1,3,21,32547,30401]
a: [3,32552,1,3,1,32548]
s: 32552
i: [1,18196,0,10,3,0,0,3] dt_k12_finseq_1
s: 18196
i: [1,18199,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 18199
d: [2,32570,1,3,1,1,40,22398,18146]
a: [3,32586,1,3,1,32570]
s: 32586
d: [2,33220,3,0,0,2,52,32992,32586,18139]
e: 33220
d: [2,33225,2,2,1,1,40,32586,18139]
d: [2,33228,2,0,0,2,21,33225,32992]
e: 33228
r: 22398
r: 32586
i: [1,17054,0,67,23,1,0,0] t6_zf_lang
a: [3,18994,0,2,1,17054]
s: 18994
i: [1,16478,0,67,23,1,0,0] t6_zf_lang
a: [3,19486,0,2,1,16478]
s: 19486
i: [1,17342,0,68,23,1,0,0] t6_zf_lang
a: [3,18748,0,3,1,17342]
s: 18748
r: 18330
r: 18748
r: 18323
r: 18337
s: 18323
s: 18330
s: 18336
a: [3,18747,0,3,1,17342]
s: 18747
i: [1,16766,0,68,23,1,0,0] t6_zf_lang
a: [3,19240,0,3,1,16766]
s: 19240
r: 19240
a: [3,19239,0,3,1,16766]
s: 19239
s: 22397
i: [1,17630,0,70,23,1,0,0] t6_zf_lang
a: [3,18502,0,5,1,17630]
s: 18502
d: [2,32625,1,7,3,1,40,18098,30396]
d: [2,32627,1,5,2,2,21,32625,30401]
a: [3,32631,1,3,1,32627]
s: 32631
s: 18333
i: [1,18012,0,11,4,0,0,255] cc3_subset_1
s: 18012
d: [2,32666,1,7,3,1,40,18104,30396]
d: [2,32668,1,5,2,2,21,32666,18322]
d: [2,32669,1,3,1,3,21,32668,30401]
a: [3,32673,1,3,1,32669]
s: 32673
i: [1,18013,0,11,4,0,0,255] cc2_subset_1
s: 18013
d: [2,32709,1,7,3,1,40,18111,30396]
d: [2,32711,1,5,2,2,21,32709,18336]
d: [2,32712,1,3,1,3,21,32711,30401]
a: [3,32716,1,3,1,32712]
s: 32716
i: [1,18054,0,11,5,0,0,255] t2_real
s: 18054
i: [1,18055,0,11,5,0,0,255] t1_real
s: 18055
d: [2,33438,1,8,4,1,40,18054,30576]
d: [2,33460,1,6,3,2,21,33438,31505]
d: [2,33461,1,4,2,3,21,33460,31500]
a: [3,33465,1,2,1,33461]
s: 33465
i: [1,18056,0,11,5,0,0,255] t8_real
s: 18056
a: [3,33468,1,2,1,33461]
s: 33468
i: [1,18057,0,11,5,0,0,255] t7_real
s: 18057
i: [1,18058,0,11,5,0,0,255] t3_real
s: 18058
d: [2,33440,1,8,4,1,40,18054,30551]
d: [2,33471,1,6,3,2,21,33440,31510]
d: [2,33472,1,4,2,3,21,33471,31505]
a: [3,33475,1,2,1,33472]
s: 33475
i: [1,18059,0,11,5,0,0,255] t4_real
s: 18059
d: [2,33442,1,8,4,1,40,18054,30571]
d: [2,33481,1,6,3,2,21,33442,31515]
d: [2,33482,1,4,2,3,21,33481,31500]
a: [3,33485,1,2,1,33482]
s: 33485
i: [1,18093,0,11,4,0,0,1] d20_zf_lang
s: 18093
i: [1,18095,0,11,4,0,0,1] d20_zf_lang
s: 18095
d: [2,33446,1,8,4,1,40,18054,30561]
d: [2,33494,1,6,3,2,21,33446,31520]
d: [2,33495,1,4,2,3,21,33494,31500]
a: [3,33498,1,2,1,33495]
s: 33498
i: [1,18099,0,11,4,0,0,1] d22_zf_lang
s: 18099
d: [2,31784,2,2,1,1,40,31663,18040]
a: [3,31794,2,2,1,31784]
s: 31794
i: [1,18101,0,11,4,0,0,1] d22_zf_lang
s: 18101
i: [1,18105,0,11,4,0,0,1] d21_zf_lang
s: 18105
d: [2,31816,2,2,1,1,40,31668,18042]
a: [3,31823,2,2,1,31816]
s: 31823
i: [1,18107,0,11,4,0,0,1] d21_zf_lang
s: 18107
d: [2,31818,2,2,1,1,40,31668,18040]
a: [3,31828,2,2,1,31818]
s: 31828
i: [1,18112,0,11,4,0,0,1] d23_zf_lang
s: 18112
i: [1,18197,0,11,3,0,0,3] redefinition_k12_finseq_1
s: 18197
d: [2,31970,2,2,1,1,40,31886,18041]
a: [3,31975,2,2,1,31970]
s: 31975
i: [1,18200,0,11,3,0,0,255] commutativity_k3_xcmplx_0
s: 18200
d: [2,31464,3,2,1,1,40,31159,18061]
a: [3,31469,3,2,1,31464]
s: 31469
i: [1,18201,0,11,3,0,0,255] commutativity_k2_xcmplx_0
s: 18201
i: [1,18202,0,11,3,0,0,255] dt_k1_nat_1
s: 18202
d: [2,31465,3,2,1,1,40,31159,18060]
a: [3,31474,3,2,1,31465]
s: 31474
i: [1,18235,0,11,3,0,0,2] dt_k7_zf_lang
s: 18235
d: [2,31611,3,2,1,1,40,31401,18051]
a: [3,31617,3,2,1,31611]
s: 31617
i: [1,18236,0,11,3,0,0,2] dt_k7_zf_lang
s: 18236
i: [1,18242,0,11,3,0,0,2] dt_k8_zf_lang
s: 18242
d: [2,31612,3,2,1,1,40,31401,18050]
a: [3,31622,3,2,1,31612]
s: 31622
i: [1,18243,0,11,3,0,0,2] dt_k8_zf_lang
s: 18243
d: [2,31624,3,2,1,1,40,31416,18051]
a: [3,31629,3,2,1,31624]
s: 31629
i: [1,18077,0,12,2,0,0,3] d5_zf_lang
s: 18077
i: [1,18110,0,12,4,0,0,1] d23_zf_lang
s: 18110
d: [2,31625,3,2,1,1,40,31416,18050]
a: [3,31634,3,2,1,31625]
s: 31634
i: [1,18185,0,12,5,0,0,255] fc26_xreal_0
s: 18185
d: [2,32518,3,2,1,1,40,31687,18061]
a: [3,32523,3,2,1,32518]
s: 32523
i: [1,18186,0,12,5,0,0,255] fc9_xreal_0
s: 18186
i: [1,18187,0,12,5,0,0,255] fc23_xreal_0
s: 18187
d: [2,32519,3,2,1,1,40,31687,18060]
a: [3,32528,3,2,1,32519]
s: 32528
i: [1,18188,0,12,5,0,0,255] fc24_xreal_0
s: 18188
d: [2,32601,3,2,1,1,40,31738,18061]
a: [3,32606,3,2,1,32601]
s: 32606
i: [1,18189,0,12,5,0,0,255] fc25_xreal_0
s: 18189
i: [1,18190,0,12,5,0,0,255] fc10_xreal_0
s: 18190
d: [2,32602,3,2,1,1,40,31738,18060]
a: [3,32611,3,2,1,32602]
s: 32611
i: [1,18203,0,12,3,0,0,255] commutativity_k1_nat_1
s: 18203
d: [2,32642,3,2,1,1,40,31755,18061]
a: [3,32647,3,2,1,32642]
s: 32647
i: [1,18204,0,12,3,0,0,255] redefinition_k1_nat_1
s: 18204
i: [1,18207,0,12,5,0,0,255] fc13_xreal_0
s: 18207
d: [2,32643,3,2,1,1,40,31755,18060]
a: [3,32652,3,2,1,32643]
s: 32652
i: [1,18208,0,12,5,0,0,255] fc14_xreal_0
s: 18208
d: [2,32695,3,2,1,1,40,31789,18061]
a: [3,32700,3,2,1,32695]
s: 32700
i: [1,18222,0,12,5,0,0,255] fc11_xreal_0
s: 18222
i: [1,18223,0,12,5,0,0,255] fc12_xreal_0
s: 18223
d: [2,32696,3,2,1,1,40,31789,18060]
a: [3,32705,3,2,1,32696]
s: 32705
i: [1,18237,0,12,4,0,0,2] fc6_zf_lang
s: 18237
d: [2,33644,3,2,1,1,40,31823,18061]
a: [3,33649,3,2,1,33644]
s: 33649
i: [1,18238,0,12,4,0,0,1] dt_k13_zf_lang
s: 18238
i: [1,18240,0,12,4,0,0,1] dt_k13_zf_lang
s: 18240
d: [2,33645,3,2,1,1,40,31823,18060]
a: [3,33654,3,2,1,33645]
s: 33654
i: [1,18251,0,12,3,0,0,3] dt_k8_finseq_1
s: 18251
d: [2,33693,3,2,1,1,40,31975,18051]
a: [3,33698,3,2,1,33693]
s: 33698
i: [1,18052,0,13,6,0,0,255] t5_real
s: 18052
i: [1,18053,0,13,6,0,0,255] t6_real
s: 18053
d: [2,33694,3,2,1,1,40,31975,18050]
a: [3,33703,3,2,1,33694]
s: 33703
i: [1,18239,0,13,4,0,0,1] dt_k13_zf_lang
s: 18239
d: [2,31278,1,3,1,0,40,18063,18168]
s: 31278
i: [1,18241,0,13,4,0,0,1] dt_k13_zf_lang
s: 18241
i: [1,18096,0,14,4,0,0,1] d20_zf_lang
s: 18096
d: [2,34182,1,24,8,1,44,18312,18096]
d: [2,34183,1,21,7,1,21,34182,18092]
d: [2,34184,1,17,6,1,21,34183,18093]
d: [2,34185,1,14,5,1,21,34184,18094]
d: [2,34186,1,10,4,1,21,34185,18095]
s: 34186
i: [1,18102,0,14,4,0,0,1] d22_zf_lang
s: 18102
d: [2,34188,1,24,8,1,44,18326,18102]
d: [2,34189,1,21,7,1,21,34188,18098]
d: [2,34190,1,17,6,1,21,34189,18099]
d: [2,34191,1,14,5,1,21,34190,18100]
d: [2,34192,1,10,4,1,21,34191,18101]
s: 34192
i: [1,18108,0,14,4,0,0,1] d21_zf_lang
s: 18108
i: [1,18113,0,14,4,0,0,1] d23_zf_lang
s: 18113
d: [2,34194,1,24,8,1,44,18319,18108]
d: [2,34195,1,21,7,1,21,34194,18104]
d: [2,34196,1,17,6,1,21,34195,18105]
d: [2,34197,1,14,5,1,21,34196,18106]
d: [2,34198,1,10,4,1,21,34197,18107]
s: 34198
i: [1,18195,0,14,4,0,0,1] existence_m2_subset_1
s: 18195
d: [2,34201,1,22,7,1,44,18333,18113]
d: [2,34202,1,17,6,1,21,34201,18110]
d: [2,34203,1,14,5,1,21,34202,18111]
d: [2,34204,1,10,4,1,21,34203,18112]
s: 34204
d: [2,34207,2,7,3,1,47,34204]
d: [2,34208,2,5,2,2,21,34207,18336]
d: [2,34209,2,2,1,3,21,34208,30396]
d: [2,34210,2,0,0,4,21,34209,30401]
e: 34210
r: 18336
r: 18347
r: 18359
r: 34186
r: 18312
s: 18337
s: 18348
d: [2,32729,1,7,3,1,40,18348,18092]
d: [2,32739,1,5,2,2,21,32729,18354]
a: [3,32747,1,2,1,32739]
s: 32747
s: 18360
d: [2,32731,1,7,3,1,40,18348,18098]
d: [2,32755,1,5,2,2,21,32731,18354]
a: [3,32763,1,2,1,32755]
s: 32763
d: [2,34206,0,12,3,0,21,18195,18070]
s: 34206
i: [1,18209,0,14,5,0,0,2] fc5_zf_lang
s: 18209
d: [2,32733,1,7,3,1,40,18348,18104]
d: [2,32771,1,5,2,2,21,32733,18354]
a: [3,32779,1,2,1,32771]
s: 32779
i: [1,18210,0,14,5,0,0,1] dt_k10_zf_lang
s: 18210
d: [2,32735,1,7,3,1,40,18348,18111]
d: [2,32787,1,5,2,2,21,32735,18354]
a: [3,32795,1,2,1,32787]
s: 32795
i: [1,18212,0,14,5,0,0,1] dt_k10_zf_lang
s: 18212
i: [1,18214,0,14,5,0,0,1] dt_k12_zf_lang
s: 18214
d: [2,32737,1,2,1,1,40,18348,18143]
a: [3,32805,1,2,1,32737]
s: 32805
i: [1,18216,0,14,5,0,0,1] dt_k12_zf_lang
s: 18216
d: [2,32738,1,2,1,1,40,18348,18142]
a: [3,32810,1,2,1,32738]
s: 32810
i: [1,18218,0,14,5,0,0,1] dt_k11_zf_lang
s: 18218
i: [1,18220,0,14,5,0,0,1] dt_k11_zf_lang
s: 18220
d: [2,32812,1,7,3,1,40,18360,18092]
d: [2,32822,1,5,2,2,21,32812,18366]
a: [3,32830,1,2,1,32822]
s: 32830
i: [1,18252,0,14,3,0,0,3] redefinition_k8_finseq_1
s: 18252
d: [2,32814,1,7,3,1,40,18360,18098]
d: [2,32838,1,5,2,2,21,32814,18366]
a: [3,32846,1,2,1,32838]
s: 32846
i: [1,18192,0,15,5,0,0,1] dt_m2_subset_1
s: 18192
d: [2,34312,0,13,4,0,21,18192,18070]
s: 34312
d: [2,32816,1,7,3,1,40,18360,18104]
d: [2,32854,1,5,2,2,21,32816,18366]
a: [3,32862,1,2,1,32854]
s: 32862
i: [1,18193,0,15,5,0,0,1] redefinition_m2_subset_1
s: 18193
d: [2,32818,1,7,3,1,40,18360,18111]
d: [2,32870,1,5,2,2,21,32818,18366]
a: [3,32878,1,2,1,32870]
s: 32878
d: [2,34332,0,13,4,0,21,18193,18070]
s: 34332
i: [1,18194,0,15,5,0,0,1] redefinition_m2_subset_1
s: 18194
d: [2,32820,1,2,1,1,40,18360,18143]
a: [3,32888,1,2,1,32820]
s: 32888
d: [2,34352,0,13,4,0,21,18194,18070]
s: 34352
d: [2,32821,1,2,1,1,40,18360,18142]
a: [3,32893,1,2,1,32821]
s: 32893
i: [1,18211,0,15,5,0,0,1] dt_k10_zf_lang
s: 18211
i: [1,18213,0,15,5,0,0,1] dt_k10_zf_lang
s: 18213
d: [2,34276,2,6,3,1,40,32805,18126]
a: [3,34281,2,2,1,34276]
s: 34281
i: [1,18215,0,15,5,0,0,1] dt_k12_zf_lang
s: 18215
d: [2,34353,2,6,3,1,40,32888,18126]
a: [3,34358,2,2,1,34353]
s: 34358
i: [1,18217,0,15,5,0,0,1] dt_k12_zf_lang
s: 18217
i: [1,18219,0,15,5,0,0,1] dt_k11_zf_lang
s: 18219
d: [2,32736,1,3,1,1,40,18348,18146]
a: [3,32800,1,3,1,32736]
s: 32800
d: [2,34445,3,0,0,2,52,34281,32800,18139]
e: 34445
d: [2,34450,2,2,1,1,40,32800,18139]
d: [2,34453,2,0,0,2,21,34450,34281]
e: 34453
r: 34281
i: [1,18221,0,15,5,0,0,1] dt_k11_zf_lang
s: 18221
a: [3,34280,2,2,1,34276]
s: 34280
i: [1,18232,0,16,7,0,0,4] dt_k7_finseq_1
s: 18232
i: [1,18233,0,16,7,0,0,4] dt_k7_finseq_1
s: 18233
a: [3,34285,2,2,1,34276]
s: 34285
i: [1,18234,0,16,7,0,0,4] dt_k7_finseq_1
s: 18234
d: [2,34277,2,6,3,1,40,32805,18123]
a: [3,34290,2,2,1,34277]
s: 34290
i: [1,18248,0,17,4,0,0,255] spc7_arithm
s: 18248
i: [1,18249,0,17,4,0,0,255] spc6_arithm
s: 18249
d: [2,34461,3,4,2,1,40,34285,18090]
a: [3,34465,3,2,1,34461]
s: 34465
i: [1,18087,0,18,4,0,0,1] d19_zf_lang
s: 18087
d: [2,32819,1,3,1,1,40,18360,18146]
a: [3,32883,1,3,1,32819]
s: 32883
d: [2,34597,3,0,0,2,52,34358,32883,18139]
e: 34597
d: [2,34602,2,2,1,1,40,32883,18139]
d: [2,34605,2,0,0,2,21,34602,34358]
e: 34605
r: 34358
i: [1,18224,0,18,8,0,0,4] fc18_finseq_1
s: 18224
i: [1,18225,0,18,8,0,0,4] fc18_finseq_1
s: 18225
a: [3,34357,2,2,1,34353]
s: 34357
i: [1,18226,0,18,8,0,0,4] fc18_finseq_1
s: 18226
a: [3,34362,2,2,1,34353]
s: 34362
i: [1,18227,0,18,8,0,0,4] fc18_finseq_1
s: 18227
i: [1,18228,0,18,8,0,0,4] fc17_finseq_1
s: 18228
d: [2,34354,2,6,3,1,40,32888,18123]
a: [3,34367,2,2,1,34354]
s: 34367
i: [1,18229,0,18,8,0,0,4] fc17_finseq_1
s: 18229
d: [2,34613,3,4,2,1,40,34362,18090]
a: [3,34617,3,2,1,34613]
s: 34617
i: [1,18230,0,18,8,0,0,4] fc17_finseq_1
s: 18230
i: [1,18231,0,18,8,0,0,4] fc17_finseq_1
s: 18231
d: [2,34446,2,5,2,1,40,32800,18091]
d: [2,34452,2,3,1,2,21,34446,18354]
a: [3,34459,2,3,1,34452]
s: 34459
i: [1,18078,0,19,3,0,0,3] d6_zf_lang
s: 18078
d: [2,34598,2,5,2,1,40,32883,18091]
d: [2,34604,2,3,1,2,21,34598,18366]
a: [3,34611,2,3,1,34604]
s: 34611
i: [1,18115,0,19,5,0,0,1] d17_zf_lang
s: 18115
i: [1,18250,0,19,4,0,0,255] spc5_arithm
s: 18250
d: [2,31279,1,3,1,0,40,18063,18172]
s: 31279
i: [1,18116,0,20,5,0,0,1] d16_zf_lang
s: 18116
d: [2,31298,1,3,1,0,40,18064,18168]
s: 31298
i: [1,18086,0,21,3,0,0,2] d7_zf_lang
s: 18086
i: [1,18117,0,21,5,0,0,1] d18_zf_lang
s: 18117
d: [2,31299,1,3,1,0,40,18064,18172]
s: 31299
i: [1,18308,0,22,6,0,0,1] d23_zf_lang
s: 18308
d: [2,31318,1,3,1,0,40,18065,18168]
s: 31318
i: [1,18305,0,23,7,0,0,1] d20_zf_lang
s: 18305
d: [2,34793,1,17,7,1,44,18305,18342]
d: [2,34800,1,14,6,2,21,34793,30396]
d: [2,34801,1,11,5,3,21,34800,18348]
d: [2,34802,1,9,4,4,21,34801,18354]
d: [2,34803,1,6,3,5,21,34802,18360]
d: [2,34804,1,4,2,6,21,34803,18366]
d: [2,34805,1,2,1,7,21,34804,18316]
d: [2,34806,1,0,0,8,21,34805,30401]
e: 34806
r: 18366
r: 18397
r: 18449
r: 34198
r: 18319
a: [3,18365,0,2,1,17908]
s: 18365
d: [2,32895,1,7,3,1,40,18398,18092]
d: [2,32905,1,5,2,2,21,32895,18424]
a: [3,32913,1,2,1,32905]
s: 32913
s: 18398
d: [2,32897,1,7,3,1,40,18398,18098]
d: [2,32921,1,5,2,2,21,32897,18424]
a: [3,32929,1,2,1,32921]
s: 32929
s: 18450
i: [1,18306,0,23,7,0,0,1] d22_zf_lang
s: 18306
d: [2,32899,1,7,3,1,40,18398,18104]
d: [2,32937,1,5,2,2,21,32899,18424]
a: [3,32945,1,2,1,32937]
s: 32945
i: [1,18307,0,23,7,0,0,1] d21_zf_lang
s: 18307
d: [2,34887,1,17,7,1,44,18307,18372]
d: [2,34894,1,14,6,2,21,34887,30396]
d: [2,34895,1,11,5,3,21,34894,18398]
d: [2,34896,1,9,4,4,21,34895,18424]
d: [2,34897,1,6,3,5,21,34896,18450]
d: [2,34898,1,4,2,6,21,34897,18476]
d: [2,34899,1,2,1,7,21,34898,18323]
d: [2,34900,1,0,0,8,21,34899,30401]
e: 34900
r: 34204
r: 18333
r: 18342
r: 18365
r: 18424
r: 22397
r: 32992
r: 18354
r: 18450
r: 18747
r: 19239
r: 25064
r: 18330
r: 18360
r: 18502
r: 18316
r: 34192
r: 18326
r: 18398
r: 19732
s: 18315
a: [3,18329,0,2,1,17918]
s: 18329
s: 18424
d: [2,32901,1,7,3,1,40,18398,18111]
d: [2,32953,1,5,2,2,21,32901,18424]
a: [3,32961,1,2,1,32953]
s: 32961
s: 18360
d: [2,32903,1,2,1,1,40,18398,18143]
a: [3,32971,1,2,1,32903]
s: 32971
s: 18397
s: 18748
d: [2,32904,1,2,1,1,40,18398,18142]
a: [3,32976,1,2,1,32904]
s: 32976
s: 19240
d: [2,33003,1,7,3,1,40,18450,18092]
d: [2,33013,1,5,2,2,21,33003,18476]
a: [3,33021,1,2,1,33013]
s: 33021
s: 22398
a: [3,18341,0,5,1,17916]
s: 18341
d: [2,33005,1,7,3,1,40,18450,18098]
d: [2,33029,1,5,2,2,21,33005,18476]
a: [3,33037,1,2,1,33029]
s: 33037
a: [3,18501,0,5,1,17630]
s: 18501
d: [2,33007,1,7,3,1,40,18450,18104]
d: [2,33045,1,5,2,2,21,33007,18476]
a: [3,33053,1,2,1,33045]
s: 33053
a: [3,19731,0,5,1,14462]
s: 19731
s: 18312
d: [2,33009,1,7,3,1,40,18450,18111]
d: [2,33061,1,5,2,2,21,33009,18476]
a: [3,33069,1,2,1,33061]
s: 33069
i: [1,18253,0,27,11,0,0,255] fc28_finseq_1
s: 18253
d: [2,33011,1,2,1,1,40,18450,18143]
a: [3,33079,1,2,1,33011]
s: 33079
i: [1,18254,0,27,11,0,0,255] fc28_finseq_1
s: 18254
i: [1,18256,0,27,11,0,0,255] fc28_finseq_1
s: 18256
d: [2,33012,1,2,1,1,40,18450,18142]
a: [3,33084,1,2,1,33012]
s: 33084
i: [1,18255,0,30,11,0,0,255] fc28_finseq_1
s: 18255
d: [2,33230,1,7,3,1,40,18748,18092]
d: [2,33240,1,5,2,2,21,33230,18994]
a: [3,33248,1,2,1,33240]
s: 33248
d: [2,33232,1,7,3,1,40,18748,18098]
d: [2,33256,1,5,2,2,21,33232,18994]
a: [3,33264,1,2,1,33256]
s: 33264
d: [2,33234,1,7,3,1,40,18748,18104]
d: [2,33272,1,5,2,2,21,33234,18994]
a: [3,33280,1,2,1,33272]
s: 33280
d: [2,33236,1,7,3,1,40,18748,18111]
d: [2,33288,1,5,2,2,21,33236,18994]
a: [3,33296,1,2,1,33288]
s: 33296
d: [2,33238,1,2,1,1,40,18748,18143]
a: [3,33306,1,2,1,33238]
s: 33306
d: [2,33239,1,2,1,1,40,18748,18142]
a: [3,33311,1,2,1,33239]
s: 33311
d: [2,33325,1,7,3,1,40,19240,18092]
d: [2,33335,1,5,2,2,21,33325,19486]
a: [3,33343,1,2,1,33335]
s: 33343
d: [2,33327,1,7,3,1,40,19240,18098]
d: [2,33351,1,5,2,2,21,33327,19486]
a: [3,33359,1,2,1,33351]
s: 33359
d: [2,33329,1,7,3,1,40,19240,18104]
d: [2,33367,1,5,2,2,21,33329,19486]
a: [3,33375,1,2,1,33367]
s: 33375
d: [2,33331,1,7,3,1,40,19240,18111]
d: [2,33383,1,5,2,2,21,33331,19486]
a: [3,33391,1,2,1,33383]
s: 33391
d: [2,33333,1,2,1,1,40,19240,18143]
a: [3,33401,1,2,1,33333]
s: 33401
d: [2,33334,1,2,1,1,40,19240,18142]
a: [3,33406,1,2,1,33334]
s: 33406
s: 32586
a: [3,32991,2,2,1,32987]
s: 32991
d: [2,32902,1,3,1,1,40,18398,18146]
a: [3,32966,1,3,1,32902]
s: 32966
d: [2,35113,2,0,0,2,52,18397,32966,18147]
e: 35113
d: [2,35115,2,3,1,1,40,32966,18147]
d: [2,35121,2,0,0,2,21,35115,18397]
e: 35121
r: 32966
a: [3,32996,2,2,1,32987]
s: 32996
d: [2,32615,1,7,3,1,40,18094,30396]
d: [2,32978,1,5,2,2,21,32615,30401]
a: [3,32982,1,3,1,32978]
s: 32982
d: [2,33010,1,3,1,1,40,18450,18146]
a: [3,33074,1,3,1,33010]
s: 33074
d: [2,32988,2,6,3,1,40,32591,18123]
a: [3,33001,2,2,1,32988]
s: 33001
s: 18450
d: [2,34963,2,6,3,1,40,32971,18126]
a: [3,34968,2,2,1,34963]
s: 34968
d: [2,33237,1,3,1,1,40,18748,18146]
a: [3,33301,1,3,1,33237]
s: 33301
d: [2,32656,1,7,3,1,40,18100,30396]
d: [2,33313,1,5,2,2,21,32656,30401]
a: [3,33317,1,3,1,33313]
s: 33317
d: [2,35049,2,6,3,1,40,33079,18126]
a: [3,35053,2,2,1,35049]
s: 35053
d: [2,33332,1,3,1,1,40,19240,18146]
a: [3,33396,1,3,1,33332]
s: 33396
a: [3,35058,2,2,1,35049]
s: 35058
d: [2,34770,1,19,5,1,44,18117,18372]
d: [2,34811,1,17,4,2,21,34770,18424]
d: [2,34812,1,15,3,3,21,34811,18476]
a: [3,34816,1,9,1,34812]
s: 34816
d: [2,34772,1,19,5,1,44,18117,18372]
d: [2,34818,1,17,4,2,21,34772,18476]
d: [2,34819,1,15,3,3,21,34818,18424]
a: [3,34823,1,9,1,34819]
s: 34823
d: [2,35050,2,6,3,1,40,33079,18123]
a: [3,35063,2,2,1,35050]
s: 35063
d: [2,35044,1,24,8,1,44,18312,18096]
d: [2,35045,1,21,7,1,21,35044,18092]
d: [2,35046,1,17,6,1,21,35045,18093]
d: [2,35047,1,14,5,1,21,35046,18094]
d: [2,35048,1,10,4,1,21,35047,18095]
s: 35048
d: [2,35228,2,7,3,1,47,35048]
d: [2,35229,2,5,2,2,21,35228,18315]
d: [2,35230,2,2,1,3,21,35229,30396]
d: [2,35231,2,0,0,4,21,35230,30401]
e: 35231
r: 35048
r: 18312
r: 18315
r: 18323
r: 18337
r: 18341
r: 18348
r: 18397
r: 18450
r: 19731
r: 33074
r: 32800
s: 18322
s: 18336
s: 18347
d: [2,35070,2,6,3,1,40,33306,18126]
a: [3,35074,2,2,1,35070]
s: 35074
s: 18449
a: [3,35079,2,2,1,35070]
s: 35079
s: 18319
a: [3,32799,1,3,1,32736]
s: 32799
d: [2,35071,2,6,3,1,40,33306,18123]
a: [3,35084,2,2,1,35071]
s: 35084
s: 32966
d: [2,35254,3,0,0,2,52,34968,32966,18139]
e: 35254
d: [2,35259,2,2,1,1,40,32966,18139]
d: [2,35262,2,0,0,2,21,35259,34968]
e: 35262
r: 34968
a: [3,34967,2,2,1,34963]
s: 34967
s: 18398
a: [3,33073,1,3,1,33010]
s: 33073
a: [3,34972,2,2,1,34963]
s: 34972
d: [2,32688,1,7,3,1,40,18106,30396]
d: [2,33319,1,5,2,2,21,32688,30401]
a: [3,33323,1,3,1,33319]
s: 33323
d: [2,34964,2,6,3,1,40,32971,18123]
a: [3,34977,2,2,1,34964]
s: 34977
d: [2,35248,1,24,8,1,44,18319,18108]
d: [2,35250,1,21,7,1,21,35248,18104]
d: [2,35251,1,17,6,1,21,35250,18105]
d: [2,35252,1,14,5,1,21,35251,18106]
d: [2,35253,1,10,4,1,21,35252,18107]
s: 35253
d: [2,35291,2,7,3,1,47,35253]
d: [2,35292,2,5,2,2,21,35291,18322]
d: [2,35293,2,2,1,3,21,35292,30396]
d: [2,35294,2,0,0,4,21,35293,30401]
e: 35294
i: [1,17915,0,57,19,1,0,0] t6_zf_lang
i: [1,17913,0,55,19,1,0,0] t6_zf_lang
i: [1,17911,0,54,19,1,0,0] t6_zf_lang
i: [1,17909,0,55,19,1,0,0] t6_zf_lang
i: [1,17907,0,54,19,1,0,0] t6_zf_lang
i: [1,17882,0,57,19,1,0,0] t6_zf_lang
i: [1,17880,0,60,19,1,0,0] t6_zf_lang
i: [1,17879,0,47,15,1,0,0] t6_zf_lang
i: [1,17878,0,58,19,1,0,0] t6_zf_lang
i: [1,17877,0,45,15,1,0,0] t6_zf_lang
i: [1,17876,0,57,19,1,0,0] t6_zf_lang
i: [1,17875,0,44,15,1,0,0] t6_zf_lang
i: [1,17873,0,45,15,1,0,0] t6_zf_lang
i: [1,17872,0,57,19,1,0,0] t6_zf_lang
i: [1,17871,0,44,15,1,0,0] t6_zf_lang
i: [1,17858,0,55,19,1,0,0] t6_zf_lang
i: [1,17855,0,45,15,1,0,0] t6_zf_lang
i: [1,17853,0,43,15,1,0,0] t6_zf_lang
i: [1,17851,0,42,15,1,0,0] t6_zf_lang
i: [1,17849,0,43,15,1,0,0] t6_zf_lang
i: [1,17847,0,42,15,1,0,0] t6_zf_lang
i: [1,17834,0,54,19,1,0,0] t6_zf_lang
i: [1,17831,0,44,15,1,0,0] t6_zf_lang
i: [1,17829,0,42,15,1,0,0] t6_zf_lang
i: [1,17827,0,41,15,1,0,0] t6_zf_lang
i: [1,17825,0,42,15,1,0,0] t6_zf_lang
i: [1,17823,0,41,15,1,0,0] t6_zf_lang
i: [1,17810,0,55,19,1,0,0] t6_zf_lang
i: [1,17808,0,58,19,1,0,0] t6_zf_lang
i: [1,17807,0,45,15,1,0,0] t6_zf_lang
i: [1,17806,0,56,19,1,0,0] t6_zf_lang
i: [1,17805,0,43,15,1,0,0] t6_zf_lang
i: [1,17804,0,55,19,1,0,0] t6_zf_lang
i: [1,17803,0,42,15,1,0,0] t6_zf_lang
i: [1,17801,0,43,15,1,0,0] t6_zf_lang
i: [1,17800,0,55,19,1,0,0] t6_zf_lang
i: [1,17799,0,42,15,1,0,0] t6_zf_lang
i: [1,17786,0,54,19,1,0,0] t6_zf_lang
i: [1,17784,0,57,19,1,0,0] t6_zf_lang
i: [1,17783,0,44,15,1,0,0] t6_zf_lang
i: [1,17782,0,55,19,1,0,0] t6_zf_lang
i: [1,17781,0,42,15,1,0,0] t6_zf_lang
i: [1,17780,0,54,19,1,0,0] t6_zf_lang
i: [1,17779,0,41,15,1,0,0] t6_zf_lang
i: [1,17777,0,42,15,1,0,0] t6_zf_lang
i: [1,17776,0,54,19,1,0,0] t6_zf_lang
i: [1,17775,0,41,15,1,0,0] t6_zf_lang
i: [1,17486,0,57,19,1,0,0] t6_zf_lang
i: [1,17483,0,47,15,1,0,0] t6_zf_lang
i: [1,17481,0,45,15,1,0,0] t6_zf_lang
i: [1,17479,0,44,15,1,0,0] t6_zf_lang
i: [1,17477,0,45,15,1,0,0] t6_zf_lang
i: [1,17475,0,44,15,1,0,0] t6_zf_lang
i: [1,17450,0,47,15,1,0,0] t6_zf_lang
i: [1,17447,0,37,11,1,0,0] t6_zf_lang
i: [1,17445,0,35,11,1,0,0] t6_zf_lang
i: [1,17443,0,34,11,1,0,0] t6_zf_lang
i: [1,17441,0,35,11,1,0,0] t6_zf_lang
i: [1,17439,0,34,11,1,0,0] t6_zf_lang
i: [1,17426,0,45,15,1,0,0] t6_zf_lang
i: [1,17423,0,35,11,1,0,0] t6_zf_lang
i: [1,17421,0,33,11,1,0,0] t6_zf_lang
i: [1,17419,0,32,11,1,0,0] t6_zf_lang
i: [1,17417,0,33,11,1,0,0] t6_zf_lang
i: [1,17415,0,32,11,1,0,0] t6_zf_lang
i: [1,17402,0,44,15,1,0,0] t6_zf_lang
i: [1,17399,0,34,11,1,0,0] t6_zf_lang
i: [1,17397,0,32,11,1,0,0] t6_zf_lang
i: [1,17395,0,31,11,1,0,0] t6_zf_lang
i: [1,17393,0,32,11,1,0,0] t6_zf_lang
i: [1,17391,0,31,11,1,0,0] t6_zf_lang
i: [1,17378,0,45,15,1,0,0] t6_zf_lang
i: [1,17375,0,35,11,1,0,0] t6_zf_lang
i: [1,17373,0,33,11,1,0,0] t6_zf_lang
i: [1,17371,0,32,11,1,0,0] t6_zf_lang
i: [1,17369,0,33,11,1,0,0] t6_zf_lang
i: [1,17367,0,32,11,1,0,0] t6_zf_lang
i: [1,17354,0,44,15,1,0,0] t6_zf_lang
i: [1,17351,0,34,11,1,0,0] t6_zf_lang
i: [1,17349,0,32,11,1,0,0] t6_zf_lang
i: [1,17347,0,31,11,1,0,0] t6_zf_lang
i: [1,17345,0,32,11,1,0,0] t6_zf_lang
i: [1,17343,0,31,11,1,0,0] t6_zf_lang
i: [1,17198,0,55,19,1,0,0] t6_zf_lang
i: [1,17195,0,45,15,1,0,0] t6_zf_lang
i: [1,17193,0,43,15,1,0,0] t6_zf_lang
i: [1,17191,0,42,15,1,0,0] t6_zf_lang
i: [1,17189,0,43,15,1,0,0] t6_zf_lang
i: [1,17187,0,42,15,1,0,0] t6_zf_lang
i: [1,17162,0,45,15,1,0,0] t6_zf_lang
i: [1,17159,0,35,11,1,0,0] t6_zf_lang
i: [1,17157,0,33,11,1,0,0] t6_zf_lang
i: [1,17155,0,32,11,1,0,0] t6_zf_lang
i: [1,17153,0,33,11,1,0,0] t6_zf_lang
i: [1,17151,0,32,11,1,0,0] t6_zf_lang
i: [1,17138,0,43,15,1,0,0] t6_zf_lang
i: [1,17135,0,33,11,1,0,0] t6_zf_lang
i: [1,17133,0,31,11,1,0,0] t6_zf_lang
i: [1,17131,0,30,11,1,0,0] t6_zf_lang
i: [1,17129,0,31,11,1,0,0] t6_zf_lang
i: [1,17127,0,30,11,1,0,0] t6_zf_lang
i: [1,17114,0,42,15,1,0,0] t6_zf_lang
i: [1,17111,0,32,11,1,0,0] t6_zf_lang
i: [1,17109,0,30,11,1,0,0] t6_zf_lang
i: [1,17107,0,29,11,1,0,0] t6_zf_lang
i: [1,17105,0,30,11,1,0,0] t6_zf_lang
i: [1,17103,0,29,11,1,0,0] t6_zf_lang
i: [1,17090,0,43,15,1,0,0] t6_zf_lang
i: [1,17087,0,33,11,1,0,0] t6_zf_lang
i: [1,17085,0,31,11,1,0,0] t6_zf_lang
i: [1,17083,0,30,11,1,0,0] t6_zf_lang
i: [1,17081,0,31,11,1,0,0] t6_zf_lang
i: [1,17079,0,30,11,1,0,0] t6_zf_lang
i: [1,17066,0,42,15,1,0,0] t6_zf_lang
i: [1,17063,0,32,11,1,0,0] t6_zf_lang
i: [1,17061,0,30,11,1,0,0] t6_zf_lang
i: [1,17059,0,29,11,1,0,0] t6_zf_lang
i: [1,17057,0,30,11,1,0,0] t6_zf_lang
i: [1,17055,0,29,11,1,0,0] t6_zf_lang
i: [1,16910,0,54,19,1,0,0] t6_zf_lang
i: [1,16907,0,44,15,1,0,0] t6_zf_lang
i: [1,16905,0,42,15,1,0,0] t6_zf_lang
i: [1,16903,0,41,15,1,0,0] t6_zf_lang
i: [1,16901,0,42,15,1,0,0] t6_zf_lang
i: [1,16899,0,41,15,1,0,0] t6_zf_lang
i: [1,16874,0,44,15,1,0,0] t6_zf_lang
i: [1,16871,0,34,11,1,0,0] t6_zf_lang
i: [1,16869,0,32,11,1,0,0] t6_zf_lang
i: [1,16867,0,31,11,1,0,0] t6_zf_lang
i: [1,16865,0,32,11,1,0,0] t6_zf_lang
i: [1,16863,0,31,11,1,0,0] t6_zf_lang
i: [1,16850,0,42,15,1,0,0] t6_zf_lang
i: [1,16847,0,32,11,1,0,0] t6_zf_lang
i: [1,16845,0,30,11,1,0,0] t6_zf_lang
i: [1,16843,0,29,11,1,0,0] t6_zf_lang
i: [1,16841,0,30,11,1,0,0] t6_zf_lang
i: [1,16839,0,29,11,1,0,0] t6_zf_lang
i: [1,16826,0,41,15,1,0,0] t6_zf_lang
i: [1,16823,0,31,11,1,0,0] t6_zf_lang
i: [1,16821,0,29,11,1,0,0] t6_zf_lang
i: [1,16819,0,28,11,1,0,0] t6_zf_lang
i: [1,16817,0,29,11,1,0,0] t6_zf_lang
i: [1,16815,0,28,11,1,0,0] t6_zf_lang
i: [1,16802,0,42,15,1,0,0] t6_zf_lang
i: [1,16799,0,32,11,1,0,0] t6_zf_lang
i: [1,16797,0,30,11,1,0,0] t6_zf_lang
i: [1,16795,0,29,11,1,0,0] t6_zf_lang
i: [1,16793,0,30,11,1,0,0] t6_zf_lang
i: [1,16791,0,29,11,1,0,0] t6_zf_lang
i: [1,16778,0,41,15,1,0,0] t6_zf_lang
i: [1,16775,0,31,11,1,0,0] t6_zf_lang
i: [1,16773,0,29,11,1,0,0] t6_zf_lang
i: [1,16771,0,28,11,1,0,0] t6_zf_lang
i: [1,16769,0,29,11,1,0,0] t6_zf_lang
i: [1,16767,0,28,11,1,0,0] t6_zf_lang
i: [1,16757,0,56,19,1,0,0] t6_zf_lang
i: [1,16667,0,59,19,1,0,0] t6_zf_lang
i: [1,16665,0,57,19,1,0,0] t6_zf_lang
i: [1,16663,0,56,19,1,0,0] t6_zf_lang
i: [1,16659,0,56,19,1,0,0] t6_zf_lang
i: [1,16655,0,46,15,1,0,0] t6_zf_lang
i: [1,16653,0,44,15,1,0,0] t6_zf_lang
i: [1,16651,0,43,15,1,0,0] t6_zf_lang
i: [1,16647,0,43,15,1,0,0] t6_zf_lang
i: [1,16631,0,45,15,1,0,0] t6_zf_lang
i: [1,16629,0,43,15,1,0,0] t6_zf_lang
i: [1,16627,0,42,15,1,0,0] t6_zf_lang
i: [1,16623,0,42,15,1,0,0] t6_zf_lang
i: [1,16622,0,55,19,1,0,0] t6_zf_lang
i: [1,16619,0,45,15,1,0,0] t6_zf_lang
i: [1,16617,0,43,15,1,0,0] t6_zf_lang
i: [1,16615,0,42,15,1,0,0] t6_zf_lang
i: [1,16613,0,43,15,1,0,0] t6_zf_lang
i: [1,16611,0,42,15,1,0,0] t6_zf_lang
i: [1,16586,0,45,15,1,0,0] t6_zf_lang
i: [1,16583,0,35,11,1,0,0] t6_zf_lang
i: [1,16581,0,33,11,1,0,0] t6_zf_lang
i: [1,16579,0,32,11,1,0,0] t6_zf_lang
i: [1,16577,0,33,11,1,0,0] t6_zf_lang
i: [1,16575,0,32,11,1,0,0] t6_zf_lang
i: [1,16562,0,43,15,1,0,0] t6_zf_lang
i: [1,16559,0,33,11,1,0,0] t6_zf_lang
i: [1,16557,0,31,11,1,0,0] t6_zf_lang
i: [1,16555,0,30,11,1,0,0] t6_zf_lang
i: [1,16553,0,31,11,1,0,0] t6_zf_lang
i: [1,16551,0,30,11,1,0,0] t6_zf_lang
i: [1,16538,0,42,15,1,0,0] t6_zf_lang
i: [1,16535,0,32,11,1,0,0] t6_zf_lang
i: [1,16533,0,30,11,1,0,0] t6_zf_lang
i: [1,16531,0,29,11,1,0,0] t6_zf_lang
i: [1,16529,0,30,11,1,0,0] t6_zf_lang
i: [1,16527,0,29,11,1,0,0] t6_zf_lang
i: [1,16526,0,56,19,1,0,0] t6_zf_lang
i: [1,16523,0,46,15,1,0,0] t6_zf_lang
i: [1,16521,0,44,15,1,0,0] t6_zf_lang
i: [1,16519,0,43,15,1,0,0] t6_zf_lang
i: [1,16515,0,43,15,1,0,0] t6_zf_lang
i: [1,16514,0,43,15,1,0,0] t6_zf_lang
i: [1,16511,0,33,11,1,0,0] t6_zf_lang
i: [1,16509,0,31,11,1,0,0] t6_zf_lang
i: [1,16507,0,30,11,1,0,0] t6_zf_lang
i: [1,16505,0,31,11,1,0,0] t6_zf_lang
i: [1,16503,0,30,11,1,0,0] t6_zf_lang
i: [1,16490,0,42,15,1,0,0] t6_zf_lang
i: [1,16487,0,32,11,1,0,0] t6_zf_lang
i: [1,16485,0,30,11,1,0,0] t6_zf_lang
i: [1,16483,0,29,11,1,0,0] t6_zf_lang
i: [1,16481,0,30,11,1,0,0] t6_zf_lang
i: [1,16479,0,29,11,1,0,0] t6_zf_lang
i: [1,16334,0,54,19,1,0,0] t6_zf_lang
i: [1,16331,0,44,15,1,0,0] t6_zf_lang
i: [1,16329,0,42,15,1,0,0] t6_zf_lang
i: [1,16327,0,41,15,1,0,0] t6_zf_lang
i: [1,16325,0,42,15,1,0,0] t6_zf_lang
i: [1,16323,0,41,15,1,0,0] t6_zf_lang
i: [1,16298,0,44,15,1,0,0] t6_zf_lang
i: [1,16295,0,34,11,1,0,0] t6_zf_lang
i: [1,16293,0,32,11,1,0,0] t6_zf_lang
i: [1,16291,0,31,11,1,0,0] t6_zf_lang
i: [1,16289,0,32,11,1,0,0] t6_zf_lang
i: [1,16287,0,31,11,1,0,0] t6_zf_lang
i: [1,16274,0,42,15,1,0,0] t6_zf_lang
i: [1,16271,0,32,11,1,0,0] t6_zf_lang
i: [1,16269,0,30,11,1,0,0] t6_zf_lang
i: [1,16267,0,29,11,1,0,0] t6_zf_lang
i: [1,16265,0,30,11,1,0,0] t6_zf_lang
i: [1,16263,0,29,11,1,0,0] t6_zf_lang
i: [1,16250,0,41,15,1,0,0] t6_zf_lang
i: [1,16247,0,31,11,1,0,0] t6_zf_lang
i: [1,16245,0,29,11,1,0,0] t6_zf_lang
i: [1,16243,0,28,11,1,0,0] t6_zf_lang
i: [1,16241,0,29,11,1,0,0] t6_zf_lang
i: [1,16239,0,28,11,1,0,0] t6_zf_lang
i: [1,16226,0,42,15,1,0,0] t6_zf_lang
i: [1,16223,0,32,11,1,0,0] t6_zf_lang
i: [1,16221,0,30,11,1,0,0] t6_zf_lang
i: [1,16219,0,29,11,1,0,0] t6_zf_lang
i: [1,16217,0,30,11,1,0,0] t6_zf_lang
i: [1,16215,0,29,11,1,0,0] t6_zf_lang
i: [1,16202,0,41,15,1,0,0] t6_zf_lang
i: [1,16199,0,31,11,1,0,0] t6_zf_lang
i: [1,16197,0,29,11,1,0,0] t6_zf_lang
i: [1,16195,0,28,11,1,0,0] t6_zf_lang
i: [1,16193,0,29,11,1,0,0] t6_zf_lang
i: [1,16191,0,28,11,1,0,0] t6_zf_lang
i: [1,12734,0,58,20,1,0,0] t6_zf_lang
i: [1,12731,0,48,16,1,0,0] t6_zf_lang
i: [1,12729,0,46,16,1,0,0] t6_zf_lang
i: [1,12727,0,45,16,1,0,0] t6_zf_lang
i: [1,12725,0,46,16,1,0,0] t6_zf_lang
i: [1,12723,0,45,16,1,0,0] t6_zf_lang
i: [1,12698,0,48,16,1,0,0] t6_zf_lang
i: [1,12696,0,51,16,1,0,0] t6_zf_lang
i: [1,12695,0,38,12,1,0,0] t6_zf_lang
i: [1,12693,0,36,12,1,0,0] t6_zf_lang
i: [1,12692,0,48,16,1,0,0] t6_zf_lang
i: [1,12691,0,35,12,1,0,0] t6_zf_lang
i: [1,12689,0,36,12,1,0,0] t6_zf_lang
i: [1,12688,0,48,16,1,0,0] t6_zf_lang
i: [1,12687,0,35,12,1,0,0] t6_zf_lang
i: [1,12674,0,46,16,1,0,0] t6_zf_lang
i: [1,12671,0,36,12,1,0,0] t6_zf_lang
i: [1,12670,0,47,16,1,0,0] t6_zf_lang
i: [1,12669,0,34,12,1,0,0] t6_zf_lang
i: [1,12667,0,33,12,1,0,0] t6_zf_lang
i: [1,12666,0,47,16,1,0,0] t6_zf_lang
i: [1,12665,0,34,12,1,0,0] t6_zf_lang
i: [1,12663,0,33,12,1,0,0] t6_zf_lang
i: [1,12650,0,45,16,1,0,0] t6_zf_lang
i: [1,12648,0,48,16,1,0,0] t6_zf_lang
i: [1,12647,0,35,12,1,0,0] t6_zf_lang
i: [1,12646,0,46,16,1,0,0] t6_zf_lang
i: [1,12645,0,33,12,1,0,0] t6_zf_lang
i: [1,12644,0,45,16,1,0,0] t6_zf_lang
i: [1,12643,0,32,12,1,0,0] t6_zf_lang
i: [1,12642,0,46,16,1,0,0] t6_zf_lang
i: [1,12641,0,33,12,1,0,0] t6_zf_lang
i: [1,12640,0,45,16,1,0,0] t6_zf_lang
i: [1,12639,0,32,12,1,0,0] t6_zf_lang
i: [1,12626,0,46,16,1,0,0] t6_zf_lang
i: [1,12624,0,49,16,1,0,0] t6_zf_lang
i: [1,12623,0,36,12,1,0,0] t6_zf_lang
i: [1,12622,0,47,16,1,0,0] t6_zf_lang
i: [1,12621,0,34,12,1,0,0] t6_zf_lang
i: [1,12620,0,46,16,1,0,0] t6_zf_lang
i: [1,12619,0,33,12,1,0,0] t6_zf_lang
i: [1,12617,0,34,12,1,0,0] t6_zf_lang
i: [1,12616,0,46,16,1,0,0] t6_zf_lang
i: [1,12615,0,33,12,1,0,0] t6_zf_lang
i: [1,12602,0,45,16,1,0,0] t6_zf_lang
i: [1,12600,0,48,16,1,0,0] t6_zf_lang
i: [1,12599,0,35,12,1,0,0] t6_zf_lang
i: [1,12598,0,46,16,1,0,0] t6_zf_lang
i: [1,12597,0,33,12,1,0,0] t6_zf_lang
i: [1,12596,0,45,16,1,0,0] t6_zf_lang
i: [1,12595,0,32,12,1,0,0] t6_zf_lang
i: [1,12594,0,46,16,1,0,0] t6_zf_lang
i: [1,12593,0,33,12,1,0,0] t6_zf_lang
i: [1,12592,0,45,16,1,0,0] t6_zf_lang
i: [1,12591,0,32,12,1,0,0] t6_zf_lang
i: [1,12302,0,48,16,1,0,0] t6_zf_lang
i: [1,12299,0,38,12,1,0,0] t6_zf_lang
i: [1,12297,0,36,12,1,0,0] t6_zf_lang
i: [1,12295,0,35,12,1,0,0] t6_zf_lang
i: [1,12293,0,36,12,1,0,0] t6_zf_lang
i: [1,12291,0,35,12,1,0,0] t6_zf_lang
i: [1,12278,0,51,16,1,0,0] t6_zf_lang
i: [1,12273,0,39,12,1,0,0] t6_zf_lang
i: [1,12269,0,39,12,1,0,0] t6_zf_lang
i: [1,12267,0,38,12,1,0,0] t6_zf_lang
i: [1,12266,0,38,12,1,0,0] t6_zf_lang
i: [1,12263,0,28,8,1,0,0] t6_zf_lang
i: [1,12261,0,26,8,1,0,0] t6_zf_lang
i: [1,12259,0,25,8,1,0,0] t6_zf_lang
i: [1,12257,0,26,8,1,0,0] t6_zf_lang
i: [1,12255,0,25,8,1,0,0] t6_zf_lang
i: [1,12254,0,49,16,1,0,0] t6_zf_lang
i: [1,12251,0,39,12,1,0,0] t6_zf_lang
i: [1,12249,0,37,12,1,0,0] t6_zf_lang
i: [1,12247,0,36,12,1,0,0] t6_zf_lang
i: [1,12245,0,37,12,1,0,0] t6_zf_lang
i: [1,12243,0,36,12,1,0,0] t6_zf_lang
i: [1,12242,0,36,12,1,0,0] t6_zf_lang
i: [1,12239,0,26,8,1,0,0] t6_zf_lang
i: [1,12237,0,24,8,1,0,0] t6_zf_lang
i: [1,12235,0,23,8,1,0,0] t6_zf_lang
i: [1,12233,0,24,8,1,0,0] t6_zf_lang
i: [1,12231,0,23,8,1,0,0] t6_zf_lang
i: [1,12230,0,48,16,1,0,0] t6_zf_lang
i: [1,12227,0,38,12,1,0,0] t6_zf_lang
i: [1,12223,0,35,12,1,0,0] t6_zf_lang
i: [1,12219,0,35,12,1,0,0] t6_zf_lang
i: [1,12218,0,35,12,1,0,0] t6_zf_lang
i: [1,12215,0,25,8,1,0,0] t6_zf_lang
i: [1,12213,0,23,8,1,0,0] t6_zf_lang
i: [1,12211,0,22,8,1,0,0] t6_zf_lang
i: [1,12209,0,23,8,1,0,0] t6_zf_lang
i: [1,12207,0,22,8,1,0,0] t6_zf_lang
i: [1,12206,0,49,16,1,0,0] t6_zf_lang
i: [1,12201,0,37,12,1,0,0] t6_zf_lang
i: [1,12197,0,37,12,1,0,0] t6_zf_lang
i: [1,12194,0,36,12,1,0,0] t6_zf_lang
i: [1,12191,0,26,8,1,0,0] t6_zf_lang
i: [1,12189,0,24,8,1,0,0] t6_zf_lang
i: [1,12187,0,23,8,1,0,0] t6_zf_lang
i: [1,12185,0,24,8,1,0,0] t6_zf_lang
i: [1,12183,0,23,8,1,0,0] t6_zf_lang
i: [1,12182,0,48,16,1,0,0] t6_zf_lang
i: [1,12171,0,35,12,1,0,0] t6_zf_lang
i: [1,12170,0,35,12,1,0,0] t6_zf_lang
i: [1,12167,0,25,8,1,0,0] t6_zf_lang
i: [1,12165,0,23,8,1,0,0] t6_zf_lang
i: [1,12163,0,22,8,1,0,0] t6_zf_lang
i: [1,12161,0,23,8,1,0,0] t6_zf_lang
i: [1,12159,0,22,8,1,0,0] t6_zf_lang
i: [1,12149,0,47,16,1,0,0] t6_zf_lang
i: [1,12014,0,46,16,1,0,0] t6_zf_lang
i: [1,12011,0,36,12,1,0,0] t6_zf_lang
i: [1,12009,0,34,12,1,0,0] t6_zf_lang
i: [1,12007,0,33,12,1,0,0] t6_zf_lang
i: [1,12005,0,34,12,1,0,0] t6_zf_lang
i: [1,12003,0,33,12,1,0,0] t6_zf_lang
i: [1,11987,0,39,12,1,0,0] t6_zf_lang
i: [1,11985,0,37,12,1,0,0] t6_zf_lang
i: [1,11983,0,36,12,1,0,0] t6_zf_lang
i: [1,11978,0,36,12,1,0,0] t6_zf_lang
i: [1,11975,0,26,8,1,0,0] t6_zf_lang
i: [1,11973,0,24,8,1,0,0] t6_zf_lang
i: [1,11971,0,23,8,1,0,0] t6_zf_lang
i: [1,11969,0,24,8,1,0,0] t6_zf_lang
i: [1,11967,0,23,8,1,0,0] t6_zf_lang
i: [1,11963,0,37,12,1,0,0] t6_zf_lang
i: [1,11961,0,35,12,1,0,0] t6_zf_lang
i: [1,11959,0,34,12,1,0,0] t6_zf_lang
i: [1,11955,0,34,12,1,0,0] t6_zf_lang
i: [1,11954,0,34,12,1,0,0] t6_zf_lang
i: [1,11951,0,24,8,1,0,0] t6_zf_lang
i: [1,11949,0,22,8,1,0,0] t6_zf_lang
i: [1,11947,0,21,8,1,0,0] t6_zf_lang
i: [1,11945,0,22,8,1,0,0] t6_zf_lang
i: [1,11943,0,21,8,1,0,0] t6_zf_lang
i: [1,11939,0,36,12,1,0,0] t6_zf_lang
i: [1,11937,0,34,12,1,0,0] t6_zf_lang
i: [1,11935,0,33,12,1,0,0] t6_zf_lang
i: [1,11931,0,33,12,1,0,0] t6_zf_lang
i: [1,11930,0,33,12,1,0,0] t6_zf_lang
i: [1,11927,0,23,8,1,0,0] t6_zf_lang
i: [1,11925,0,21,8,1,0,0] t6_zf_lang
i: [1,11923,0,20,8,1,0,0] t6_zf_lang
i: [1,11921,0,21,8,1,0,0] t6_zf_lang
i: [1,11919,0,20,8,1,0,0] t6_zf_lang
i: [1,11915,0,37,12,1,0,0] t6_zf_lang
i: [1,11913,0,35,12,1,0,0] t6_zf_lang
i: [1,11911,0,34,12,1,0,0] t6_zf_lang
i: [1,11909,0,35,12,1,0,0] t6_zf_lang
i: [1,11907,0,34,12,1,0,0] t6_zf_lang
i: [1,11906,0,34,12,1,0,0] t6_zf_lang
i: [1,11903,0,24,8,1,0,0] t6_zf_lang
i: [1,11901,0,22,8,1,0,0] t6_zf_lang
i: [1,11899,0,21,8,1,0,0] t6_zf_lang
i: [1,11897,0,22,8,1,0,0] t6_zf_lang
i: [1,11895,0,21,8,1,0,0] t6_zf_lang
i: [1,11889,0,34,12,1,0,0] t6_zf_lang
i: [1,11887,0,33,12,1,0,0] t6_zf_lang
i: [1,11885,0,34,12,1,0,0] t6_zf_lang
i: [1,11882,0,33,12,1,0,0] t6_zf_lang
i: [1,11879,0,23,8,1,0,0] t6_zf_lang
i: [1,11877,0,21,8,1,0,0] t6_zf_lang
i: [1,11875,0,20,8,1,0,0] t6_zf_lang
i: [1,11873,0,21,8,1,0,0] t6_zf_lang
i: [1,11871,0,20,8,1,0,0] t6_zf_lang
i: [1,11726,0,45,16,1,0,0] t6_zf_lang
i: [1,11723,0,35,12,1,0,0] t6_zf_lang
i: [1,11721,0,33,12,1,0,0] t6_zf_lang
i: [1,11719,0,32,12,1,0,0] t6_zf_lang
i: [1,11717,0,33,12,1,0,0] t6_zf_lang
i: [1,11715,0,32,12,1,0,0] t6_zf_lang
i: [1,11702,0,48,16,1,0,0] t6_zf_lang
i: [1,11697,0,36,12,1,0,0] t6_zf_lang
i: [1,11693,0,36,12,1,0,0] t6_zf_lang
i: [1,11691,0,35,12,1,0,0] t6_zf_lang
i: [1,11690,0,35,12,1,0,0] t6_zf_lang
i: [1,11687,0,25,8,1,0,0] t6_zf_lang
i: [1,11685,0,23,8,1,0,0] t6_zf_lang
i: [1,11683,0,22,8,1,0,0] t6_zf_lang
i: [1,11681,0,23,8,1,0,0] t6_zf_lang
i: [1,11679,0,22,8,1,0,0] t6_zf_lang
i: [1,11673,0,34,12,1,0,0] t6_zf_lang
i: [1,11666,0,33,12,1,0,0] t6_zf_lang
i: [1,11663,0,23,8,1,0,0] t6_zf_lang
i: [1,11661,0,21,8,1,0,0] t6_zf_lang
i: [1,11659,0,20,8,1,0,0] t6_zf_lang
i: [1,11657,0,21,8,1,0,0] t6_zf_lang
i: [1,11655,0,20,8,1,0,0] t6_zf_lang
i: [1,11654,0,45,16,1,0,0] t6_zf_lang
i: [1,11651,0,35,12,1,0,0] t6_zf_lang
i: [1,11647,0,32,12,1,0,0] t6_zf_lang
i: [1,11643,0,32,12,1,0,0] t6_zf_lang
i: [1,11642,0,32,12,1,0,0] t6_zf_lang
i: [1,11639,0,22,8,1,0,0] t6_zf_lang
i: [1,11637,0,20,8,1,0,0] t6_zf_lang
i: [1,11635,0,19,8,1,0,0] t6_zf_lang
i: [1,11633,0,20,8,1,0,0] t6_zf_lang
i: [1,11631,0,19,8,1,0,0] t6_zf_lang
i: [1,11630,0,46,16,1,0,0] t6_zf_lang
i: [1,11618,0,33,12,1,0,0] t6_zf_lang
i: [1,11615,0,23,8,1,0,0] t6_zf_lang
i: [1,11613,0,21,8,1,0,0] t6_zf_lang
i: [1,11611,0,20,8,1,0,0] t6_zf_lang
i: [1,11609,0,21,8,1,0,0] t6_zf_lang
i: [1,11607,0,20,8,1,0,0] t6_zf_lang
i: [1,11606,0,45,16,1,0,0] t6_zf_lang
i: [1,11597,0,33,12,1,0,0] t6_zf_lang
i: [1,11594,0,32,12,1,0,0] t6_zf_lang
i: [1,11591,0,22,8,1,0,0] t6_zf_lang
i: [1,11589,0,20,8,1,0,0] t6_zf_lang
i: [1,11587,0,19,8,1,0,0] t6_zf_lang
i: [1,11585,0,20,8,1,0,0] t6_zf_lang
i: [1,11583,0,19,8,1,0,0] t6_zf_lang
i: [1,11573,0,47,16,1,0,0] t6_zf_lang
i: [1,11438,0,46,16,1,0,0] t6_zf_lang
i: [1,11435,0,36,12,1,0,0] t6_zf_lang
i: [1,11433,0,34,12,1,0,0] t6_zf_lang
i: [1,11431,0,33,12,1,0,0] t6_zf_lang
i: [1,11429,0,34,12,1,0,0] t6_zf_lang
i: [1,11427,0,33,12,1,0,0] t6_zf_lang
i: [1,11409,0,37,12,1,0,0] t6_zf_lang
i: [1,11403,0,36,12,1,0,0] t6_zf_lang
i: [1,11402,0,36,12,1,0,0] t6_zf_lang
i: [1,11399,0,26,8,1,0,0] t6_zf_lang
i: [1,11397,0,24,8,1,0,0] t6_zf_lang
i: [1,11395,0,23,8,1,0,0] t6_zf_lang
i: [1,11393,0,24,8,1,0,0] t6_zf_lang
i: [1,11391,0,23,8,1,0,0] t6_zf_lang
i: [1,11387,0,37,12,1,0,0] t6_zf_lang
i: [1,11385,0,35,12,1,0,0] t6_zf_lang
i: [1,11383,0,34,12,1,0,0] t6_zf_lang
i: [1,11381,0,35,12,1,0,0] t6_zf_lang
i: [1,11379,0,34,12,1,0,0] t6_zf_lang
i: [1,11378,0,34,12,1,0,0] t6_zf_lang
i: [1,11375,0,24,8,1,0,0] t6_zf_lang
i: [1,11373,0,22,8,1,0,0] t6_zf_lang
i: [1,11371,0,21,8,1,0,0] t6_zf_lang
i: [1,11369,0,22,8,1,0,0] t6_zf_lang
i: [1,11367,0,21,8,1,0,0] t6_zf_lang
i: [1,11363,0,36,12,1,0,0] t6_zf_lang
i: [1,11355,0,33,12,1,0,0] t6_zf_lang
i: [1,11354,0,33,12,1,0,0] t6_zf_lang
i: [1,11351,0,23,8,1,0,0] t6_zf_lang
i: [1,11349,0,21,8,1,0,0] t6_zf_lang
i: [1,11347,0,20,8,1,0,0] t6_zf_lang
i: [1,11345,0,21,8,1,0,0] t6_zf_lang
i: [1,11343,0,20,8,1,0,0] t6_zf_lang
i: [1,11337,0,35,12,1,0,0] t6_zf_lang
i: [1,11333,0,35,12,1,0,0] t6_zf_lang
i: [1,11330,0,34,12,1,0,0] t6_zf_lang
i: [1,11327,0,24,8,1,0,0] t6_zf_lang
i: [1,11325,0,22,8,1,0,0] t6_zf_lang
i: [1,11323,0,21,8,1,0,0] t6_zf_lang
i: [1,11321,0,22,8,1,0,0] t6_zf_lang
i: [1,11319,0,21,8,1,0,0] t6_zf_lang
i: [1,11315,0,36,12,1,0,0] t6_zf_lang
i: [1,11313,0,34,12,1,0,0] t6_zf_lang
i: [1,11309,0,34,12,1,0,0] t6_zf_lang
i: [1,11307,0,33,12,1,0,0] t6_zf_lang
i: [1,11306,0,33,12,1,0,0] t6_zf_lang
i: [1,11303,0,23,8,1,0,0] t6_zf_lang
i: [1,11301,0,21,8,1,0,0] t6_zf_lang
i: [1,11299,0,20,8,1,0,0] t6_zf_lang
i: [1,11297,0,21,8,1,0,0] t6_zf_lang
i: [1,11295,0,20,8,1,0,0] t6_zf_lang
i: [1,11150,0,45,16,1,0,0] t6_zf_lang
i: [1,11147,0,35,12,1,0,0] t6_zf_lang
i: [1,11145,0,33,12,1,0,0] t6_zf_lang
i: [1,11143,0,32,12,1,0,0] t6_zf_lang
i: [1,11141,0,33,12,1,0,0] t6_zf_lang
i: [1,11139,0,32,12,1,0,0] t6_zf_lang
i: [1,11126,0,48,16,1,0,0] t6_zf_lang
i: [1,11115,0,35,12,1,0,0] t6_zf_lang
i: [1,11114,0,35,12,1,0,0] t6_zf_lang
i: [1,11111,0,25,8,1,0,0] t6_zf_lang
i: [1,11109,0,23,8,1,0,0] t6_zf_lang
i: [1,11107,0,22,8,1,0,0] t6_zf_lang
i: [1,11105,0,23,8,1,0,0] t6_zf_lang
i: [1,11103,0,22,8,1,0,0] t6_zf_lang
i: [1,11102,0,46,16,1,0,0] t6_zf_lang
i: [1,11099,0,36,12,1,0,0] t6_zf_lang
i: [1,11090,0,33,12,1,0,0] t6_zf_lang
i: [1,11087,0,23,8,1,0,0] t6_zf_lang
i: [1,11085,0,21,8,1,0,0] t6_zf_lang
i: [1,11083,0,20,8,1,0,0] t6_zf_lang
i: [1,11081,0,21,8,1,0,0] t6_zf_lang
i: [1,11079,0,20,8,1,0,0] t6_zf_lang
i: [1,11075,0,35,12,1,0,0] t6_zf_lang
i: [1,11071,0,32,12,1,0,0] t6_zf_lang
i: [1,11067,0,32,12,1,0,0] t6_zf_lang
i: [1,11066,0,32,12,1,0,0] t6_zf_lang
i: [1,11063,0,22,8,1,0,0] t6_zf_lang
i: [1,11061,0,20,8,1,0,0] t6_zf_lang
i: [1,11059,0,19,8,1,0,0] t6_zf_lang
i: [1,11057,0,20,8,1,0,0] t6_zf_lang
i: [1,11055,0,19,8,1,0,0] t6_zf_lang
i: [1,11054,0,46,16,1,0,0] t6_zf_lang
i: [1,11051,0,36,12,1,0,0] t6_zf_lang
i: [1,11047,0,33,12,1,0,0] t6_zf_lang
i: [1,11045,0,34,12,1,0,0] t6_zf_lang
i: [1,11043,0,33,12,1,0,0] t6_zf_lang
i: [1,11042,0,33,12,1,0,0] t6_zf_lang
i: [1,11039,0,23,8,1,0,0] t6_zf_lang
i: [1,11037,0,21,8,1,0,0] t6_zf_lang
i: [1,11035,0,20,8,1,0,0] t6_zf_lang
i: [1,11033,0,21,8,1,0,0] t6_zf_lang
i: [1,11031,0,20,8,1,0,0] t6_zf_lang
i: [1,11027,0,35,12,1,0,0] t6_zf_lang
i: [1,11023,0,32,12,1,0,0] t6_zf_lang
i: [1,11019,0,32,12,1,0,0] t6_zf_lang
i: [1,11018,0,32,12,1,0,0] t6_zf_lang
i: [1,11015,0,22,8,1,0,0] t6_zf_lang
i: [1,11013,0,20,8,1,0,0] t6_zf_lang
i: [1,11011,0,19,8,1,0,0] t6_zf_lang
i: [1,11009,0,20,8,1,0,0] t6_zf_lang
i: [1,11007,0,19,8,1,0,0] t6_zf_lang
i: [1,9278,0,56,20,1,0,0] t6_zf_lang
i: [1,9275,0,46,16,1,0,0] t6_zf_lang
i: [1,9273,0,44,16,1,0,0] t6_zf_lang
i: [1,9271,0,43,16,1,0,0] t6_zf_lang
i: [1,9269,0,44,16,1,0,0] t6_zf_lang
i: [1,9267,0,43,16,1,0,0] t6_zf_lang
i: [1,9242,0,46,16,1,0,0] t6_zf_lang
i: [1,9240,0,49,16,1,0,0] t6_zf_lang
i: [1,9239,0,36,12,1,0,0] t6_zf_lang
i: [1,9238,0,47,16,1,0,0] t6_zf_lang
i: [1,9237,0,34,12,1,0,0] t6_zf_lang
i: [1,9236,0,46,16,1,0,0] t6_zf_lang
i: [1,9235,0,33,12,1,0,0] t6_zf_lang
i: [1,9234,0,47,16,1,0,0] t6_zf_lang
i: [1,9233,0,34,12,1,0,0] t6_zf_lang
i: [1,9232,0,46,16,1,0,0] t6_zf_lang
i: [1,9231,0,33,12,1,0,0] t6_zf_lang
i: [1,9218,0,44,16,1,0,0] t6_zf_lang
i: [1,9215,0,34,12,1,0,0] t6_zf_lang
i: [1,9214,0,45,16,1,0,0] t6_zf_lang
i: [1,9213,0,32,12,1,0,0] t6_zf_lang
i: [1,9211,0,31,12,1,0,0] t6_zf_lang
i: [1,9209,0,32,12,1,0,0] t6_zf_lang
i: [1,9207,0,31,12,1,0,0] t6_zf_lang
i: [1,9194,0,43,16,1,0,0] t6_zf_lang
i: [1,9192,0,46,16,1,0,0] t6_zf_lang
i: [1,9191,0,33,12,1,0,0] t6_zf_lang
i: [1,9190,0,44,16,1,0,0] t6_zf_lang
i: [1,9189,0,31,12,1,0,0] t6_zf_lang
i: [1,9188,0,43,16,1,0,0] t6_zf_lang
i: [1,9187,0,30,12,1,0,0] t6_zf_lang
i: [1,9186,0,44,16,1,0,0] t6_zf_lang
i: [1,9185,0,31,12,1,0,0] t6_zf_lang
i: [1,9184,0,43,16,1,0,0] t6_zf_lang
i: [1,9183,0,30,12,1,0,0] t6_zf_lang
i: [1,9170,0,44,16,1,0,0] t6_zf_lang
i: [1,9167,0,34,12,1,0,0] t6_zf_lang
i: [1,9165,0,32,12,1,0,0] t6_zf_lang
i: [1,9163,0,31,12,1,0,0] t6_zf_lang
i: [1,9161,0,32,12,1,0,0] t6_zf_lang
i: [1,9159,0,31,12,1,0,0] t6_zf_lang
i: [1,9146,0,43,16,1,0,0] t6_zf_lang
i: [1,9143,0,33,12,1,0,0] t6_zf_lang
i: [1,9141,0,31,12,1,0,0] t6_zf_lang
i: [1,9139,0,30,12,1,0,0] t6_zf_lang
i: [1,9138,0,44,16,1,0,0] t6_zf_lang
i: [1,9137,0,31,12,1,0,0] t6_zf_lang
i: [1,9135,0,30,12,1,0,0] t6_zf_lang
i: [1,8985,0,47,16,1,0,0] t6_zf_lang
i: [1,8983,0,46,16,1,0,0] t6_zf_lang
i: [1,8846,0,46,16,1,0,0] t6_zf_lang
i: [1,8843,0,36,12,1,0,0] t6_zf_lang
i: [1,8841,0,34,12,1,0,0] t6_zf_lang
i: [1,8839,0,33,12,1,0,0] t6_zf_lang
i: [1,8837,0,34,12,1,0,0] t6_zf_lang
i: [1,8835,0,33,12,1,0,0] t6_zf_lang
i: [1,8822,0,49,16,1,0,0] t6_zf_lang
i: [1,8819,0,39,12,1,0,0] t6_zf_lang
i: [1,8815,0,36,12,1,0,0] t6_zf_lang
i: [1,8811,0,36,12,1,0,0] t6_zf_lang
i: [1,8810,0,36,12,1,0,0] t6_zf_lang
i: [1,8807,0,26,8,1,0,0] t6_zf_lang
i: [1,8805,0,24,8,1,0,0] t6_zf_lang
i: [1,8803,0,23,8,1,0,0] t6_zf_lang
i: [1,8801,0,24,8,1,0,0] t6_zf_lang
i: [1,8799,0,23,8,1,0,0] t6_zf_lang
i: [1,8798,0,47,16,1,0,0] t6_zf_lang
i: [1,8795,0,37,12,1,0,0] t6_zf_lang
i: [1,8787,0,34,12,1,0,0] t6_zf_lang
i: [1,8786,0,34,12,1,0,0] t6_zf_lang
i: [1,8783,0,24,8,1,0,0] t6_zf_lang
i: [1,8781,0,22,8,1,0,0] t6_zf_lang
i: [1,8779,0,21,8,1,0,0] t6_zf_lang
i: [1,8777,0,22,8,1,0,0] t6_zf_lang
i: [1,8775,0,21,8,1,0,0] t6_zf_lang
i: [1,8774,0,46,16,1,0,0] t6_zf_lang
i: [1,8762,0,33,12,1,0,0] t6_zf_lang
i: [1,8759,0,23,8,1,0,0] t6_zf_lang
i: [1,8757,0,21,8,1,0,0] t6_zf_lang
i: [1,8755,0,20,8,1,0,0] t6_zf_lang
i: [1,8753,0,21,8,1,0,0] t6_zf_lang
i: [1,8751,0,20,8,1,0,0] t6_zf_lang
i: [1,8750,0,47,16,1,0,0] t6_zf_lang
i: [1,8747,0,37,12,1,0,0] t6_zf_lang
i: [1,8741,0,35,12,1,0,0] t6_zf_lang
i: [1,8738,0,34,12,1,0,0] t6_zf_lang
i: [1,8735,0,24,8,1,0,0] t6_zf_lang
i: [1,8733,0,22,8,1,0,0] t6_zf_lang
i: [1,8731,0,21,8,1,0,0] t6_zf_lang
i: [1,8729,0,22,8,1,0,0] t6_zf_lang
i: [1,8727,0,21,8,1,0,0] t6_zf_lang
i: [1,8726,0,46,16,1,0,0] t6_zf_lang
i: [1,8721,0,34,12,1,0,0] t6_zf_lang
i: [1,8715,0,33,12,1,0,0] t6_zf_lang
i: [1,8714,0,33,12,1,0,0] t6_zf_lang
i: [1,8711,0,23,8,1,0,0] t6_zf_lang
i: [1,8709,0,21,8,1,0,0] t6_zf_lang
i: [1,8707,0,20,8,1,0,0] t6_zf_lang
i: [1,8705,0,21,8,1,0,0] t6_zf_lang
i: [1,8703,0,20,8,1,0,0] t6_zf_lang
i: [1,8697,0,45,16,1,0,0] t6_zf_lang
i: [1,8695,0,44,16,1,0,0] t6_zf_lang
i: [1,8558,0,44,16,1,0,0] t6_zf_lang
i: [1,8555,0,34,12,1,0,0] t6_zf_lang
i: [1,8553,0,32,12,1,0,0] t6_zf_lang
i: [1,8551,0,31,12,1,0,0] t6_zf_lang
i: [1,8549,0,32,12,1,0,0] t6_zf_lang
i: [1,8547,0,31,12,1,0,0] t6_zf_lang
i: [1,8534,0,47,16,1,0,0] t6_zf_lang
i: [1,8529,0,35,12,1,0,0] t6_zf_lang
i: [1,8527,0,34,12,1,0,0] t6_zf_lang
i: [1,8525,0,35,12,1,0,0] t6_zf_lang
i: [1,8523,0,34,12,1,0,0] t6_zf_lang
i: [1,8522,0,34,12,1,0,0] t6_zf_lang
i: [1,8519,0,24,8,1,0,0] t6_zf_lang
i: [1,8517,0,22,8,1,0,0] t6_zf_lang
i: [1,8515,0,21,8,1,0,0] t6_zf_lang
i: [1,8513,0,22,8,1,0,0] t6_zf_lang
i: [1,8511,0,21,8,1,0,0] t6_zf_lang
i: [1,8510,0,45,16,1,0,0] t6_zf_lang
i: [1,8507,0,35,12,1,0,0] t6_zf_lang
i: [1,8505,0,33,12,1,0,0] t6_zf_lang
i: [1,8499,0,32,12,1,0,0] t6_zf_lang
i: [1,8498,0,32,12,1,0,0] t6_zf_lang
i: [1,8495,0,22,8,1,0,0] t6_zf_lang
i: [1,8493,0,20,8,1,0,0] t6_zf_lang
i: [1,8491,0,19,8,1,0,0] t6_zf_lang
i: [1,8489,0,20,8,1,0,0] t6_zf_lang
i: [1,8487,0,19,8,1,0,0] t6_zf_lang
i: [1,8486,0,44,16,1,0,0] t6_zf_lang
i: [1,8474,0,31,12,1,0,0] t6_zf_lang
i: [1,8471,0,21,8,1,0,0] t6_zf_lang
i: [1,8469,0,19,8,1,0,0] t6_zf_lang
i: [1,8467,0,18,8,1,0,0] t6_zf_lang
i: [1,8465,0,19,8,1,0,0] t6_zf_lang
i: [1,8463,0,18,8,1,0,0] t6_zf_lang
i: [1,8457,0,33,12,1,0,0] t6_zf_lang
i: [1,8455,0,32,12,1,0,0] t6_zf_lang
i: [1,8453,0,33,12,1,0,0] t6_zf_lang
i: [1,8450,0,32,12,1,0,0] t6_zf_lang
i: [1,8447,0,22,8,1,0,0] t6_zf_lang
i: [1,8445,0,20,8,1,0,0] t6_zf_lang
i: [1,8443,0,19,8,1,0,0] t6_zf_lang
i: [1,8441,0,20,8,1,0,0] t6_zf_lang
i: [1,8439,0,19,8,1,0,0] t6_zf_lang
i: [1,8438,0,44,16,1,0,0] t6_zf_lang
i: [1,8433,0,32,12,1,0,0] t6_zf_lang
i: [1,8431,0,31,12,1,0,0] t6_zf_lang
i: [1,8426,0,31,12,1,0,0] t6_zf_lang
i: [1,8423,0,21,8,1,0,0] t6_zf_lang
i: [1,8421,0,19,8,1,0,0] t6_zf_lang
i: [1,8419,0,18,8,1,0,0] t6_zf_lang
i: [1,8417,0,19,8,1,0,0] t6_zf_lang
i: [1,8415,0,18,8,1,0,0] t6_zf_lang
i: [1,8409,0,44,16,1,0,0] t6_zf_lang
i: [1,8407,0,43,16,1,0,0] t6_zf_lang
i: [1,8270,0,43,16,1,0,0] t6_zf_lang
i: [1,8267,0,33,12,1,0,0] t6_zf_lang
i: [1,8265,0,31,12,1,0,0] t6_zf_lang
i: [1,8263,0,30,12,1,0,0] t6_zf_lang
i: [1,8261,0,31,12,1,0,0] t6_zf_lang
i: [1,8259,0,30,12,1,0,0] t6_zf_lang
i: [1,8246,0,46,16,1,0,0] t6_zf_lang
i: [1,8243,0,36,12,1,0,0] t6_zf_lang
i: [1,8241,0,34,12,1,0,0] t6_zf_lang
i: [1,8239,0,33,12,1,0,0] t6_zf_lang
i: [1,8237,0,34,12,1,0,0] t6_zf_lang
i: [1,8234,0,33,12,1,0,0] t6_zf_lang
i: [1,8231,0,23,8,1,0,0] t6_zf_lang
i: [1,8229,0,21,8,1,0,0] t6_zf_lang
i: [1,8227,0,20,8,1,0,0] t6_zf_lang
i: [1,8225,0,21,8,1,0,0] t6_zf_lang
i: [1,8223,0,20,8,1,0,0] t6_zf_lang
i: [1,8222,0,44,16,1,0,0] t6_zf_lang
i: [1,8219,0,34,12,1,0,0] t6_zf_lang
i: [1,8215,0,31,12,1,0,0] t6_zf_lang
i: [1,8213,0,32,12,1,0,0] t6_zf_lang
i: [1,8211,0,31,12,1,0,0] t6_zf_lang
i: [1,8210,0,31,12,1,0,0] t6_zf_lang
i: [1,8207,0,21,8,1,0,0] t6_zf_lang
i: [1,8205,0,19,8,1,0,0] t6_zf_lang
i: [1,8203,0,18,8,1,0,0] t6_zf_lang
i: [1,8201,0,19,8,1,0,0] t6_zf_lang
i: [1,8199,0,18,8,1,0,0] t6_zf_lang
i: [1,8186,0,30,12,1,0,0] t6_zf_lang
i: [1,8183,0,20,8,1,0,0] t6_zf_lang
i: [1,8181,0,18,8,1,0,0] t6_zf_lang
i: [1,8179,0,17,8,1,0,0] t6_zf_lang
i: [1,8177,0,18,8,1,0,0] t6_zf_lang
i: [1,8175,0,17,8,1,0,0] t6_zf_lang
i: [1,8174,0,44,16,1,0,0] t6_zf_lang
i: [1,8171,0,34,12,1,0,0] t6_zf_lang
i: [1,8165,0,32,12,1,0,0] t6_zf_lang
i: [1,8162,0,31,12,1,0,0] t6_zf_lang
i: [1,8159,0,21,8,1,0,0] t6_zf_lang
i: [1,8157,0,19,8,1,0,0] t6_zf_lang
i: [1,8155,0,18,8,1,0,0] t6_zf_lang
i: [1,8153,0,19,8,1,0,0] t6_zf_lang
i: [1,8151,0,18,8,1,0,0] t6_zf_lang
i: [1,8150,0,43,16,1,0,0] t6_zf_lang
i: [1,8147,0,33,12,1,0,0] t6_zf_lang
i: [1,8143,0,30,12,1,0,0] t6_zf_lang
i: [1,8141,0,31,12,1,0,0] t6_zf_lang
i: [1,8138,0,30,12,1,0,0] t6_zf_lang
i: [1,8135,0,20,8,1,0,0] t6_zf_lang
i: [1,8133,0,18,8,1,0,0] t6_zf_lang
i: [1,8131,0,17,8,1,0,0] t6_zf_lang
i: [1,8129,0,18,8,1,0,0] t6_zf_lang
i: [1,8127,0,17,8,1,0,0] t6_zf_lang
i: [1,8123,0,47,16,1,0,0] t6_zf_lang
i: [1,8121,0,45,16,1,0,0] t6_zf_lang
i: [1,8119,0,44,16,1,0,0] t6_zf_lang
i: [1,8117,0,45,16,1,0,0] t6_zf_lang
i: [1,7982,0,44,16,1,0,0] t6_zf_lang
i: [1,7979,0,34,12,1,0,0] t6_zf_lang
i: [1,7977,0,32,12,1,0,0] t6_zf_lang
i: [1,7975,0,31,12,1,0,0] t6_zf_lang
i: [1,7973,0,32,12,1,0,0] t6_zf_lang
i: [1,7971,0,31,12,1,0,0] t6_zf_lang
i: [1,7953,0,35,12,1,0,0] t6_zf_lang
i: [1,7951,0,34,12,1,0,0] t6_zf_lang
i: [1,7949,0,35,12,1,0,0] t6_zf_lang
i: [1,7946,0,34,12,1,0,0] t6_zf_lang
i: [1,7943,0,24,8,1,0,0] t6_zf_lang
i: [1,7941,0,22,8,1,0,0] t6_zf_lang
i: [1,7939,0,21,8,1,0,0] t6_zf_lang
i: [1,7937,0,22,8,1,0,0] t6_zf_lang
i: [1,7935,0,21,8,1,0,0] t6_zf_lang
i: [1,7931,0,35,12,1,0,0] t6_zf_lang
i: [1,7929,0,33,12,1,0,0] t6_zf_lang
i: [1,7927,0,32,12,1,0,0] t6_zf_lang
i: [1,7925,0,33,12,1,0,0] t6_zf_lang
i: [1,7922,0,32,12,1,0,0] t6_zf_lang
i: [1,7919,0,22,8,1,0,0] t6_zf_lang
i: [1,7917,0,20,8,1,0,0] t6_zf_lang
i: [1,7915,0,19,8,1,0,0] t6_zf_lang
i: [1,7913,0,20,8,1,0,0] t6_zf_lang
i: [1,7911,0,19,8,1,0,0] t6_zf_lang
i: [1,7898,0,31,12,1,0,0] t6_zf_lang
i: [1,7895,0,21,8,1,0,0] t6_zf_lang
i: [1,7893,0,19,8,1,0,0] t6_zf_lang
i: [1,7891,0,18,8,1,0,0] t6_zf_lang
i: [1,7889,0,19,8,1,0,0] t6_zf_lang
i: [1,7887,0,18,8,1,0,0] t6_zf_lang
i: [1,7883,0,35,12,1,0,0] t6_zf_lang
i: [1,7881,0,33,12,1,0,0] t6_zf_lang
i: [1,7879,0,32,12,1,0,0] t6_zf_lang
i: [1,7877,0,33,12,1,0,0] t6_zf_lang
i: [1,7875,0,32,12,1,0,0] t6_zf_lang
i: [1,7874,0,32,12,1,0,0] t6_zf_lang
i: [1,7871,0,22,8,1,0,0] t6_zf_lang
i: [1,7869,0,20,8,1,0,0] t6_zf_lang
i: [1,7867,0,19,8,1,0,0] t6_zf_lang
i: [1,7865,0,20,8,1,0,0] t6_zf_lang
i: [1,7863,0,19,8,1,0,0] t6_zf_lang
i: [1,7862,0,44,16,1,0,0] t6_zf_lang
i: [1,7850,0,31,12,1,0,0] t6_zf_lang
i: [1,7847,0,21,8,1,0,0] t6_zf_lang
i: [1,7845,0,19,8,1,0,0] t6_zf_lang
i: [1,7843,0,18,8,1,0,0] t6_zf_lang
i: [1,7841,0,19,8,1,0,0] t6_zf_lang
i: [1,7839,0,18,8,1,0,0] t6_zf_lang
i: [1,7835,0,46,16,1,0,0] t6_zf_lang
i: [1,7833,0,44,16,1,0,0] t6_zf_lang
i: [1,7831,0,43,16,1,0,0] t6_zf_lang
i: [1,7694,0,43,16,1,0,0] t6_zf_lang
i: [1,7691,0,33,12,1,0,0] t6_zf_lang
i: [1,7689,0,31,12,1,0,0] t6_zf_lang
i: [1,7687,0,30,12,1,0,0] t6_zf_lang
i: [1,7685,0,31,12,1,0,0] t6_zf_lang
i: [1,7683,0,30,12,1,0,0] t6_zf_lang
i: [1,7670,0,46,16,1,0,0] t6_zf_lang
i: [1,7659,0,33,12,1,0,0] t6_zf_lang
i: [1,7658,0,33,12,1,0,0] t6_zf_lang
i: [1,7655,0,23,8,1,0,0] t6_zf_lang
i: [1,7653,0,21,8,1,0,0] t6_zf_lang
i: [1,7651,0,20,8,1,0,0] t6_zf_lang
i: [1,7649,0,21,8,1,0,0] t6_zf_lang
i: [1,7647,0,20,8,1,0,0] t6_zf_lang
i: [1,7646,0,44,16,1,0,0] t6_zf_lang
i: [1,7643,0,34,12,1,0,0] t6_zf_lang
i: [1,7634,0,31,12,1,0,0] t6_zf_lang
i: [1,7631,0,21,8,1,0,0] t6_zf_lang
i: [1,7629,0,19,8,1,0,0] t6_zf_lang
i: [1,7627,0,18,8,1,0,0] t6_zf_lang
i: [1,7625,0,19,8,1,0,0] t6_zf_lang
i: [1,7623,0,18,8,1,0,0] t6_zf_lang
i: [1,7622,0,43,16,1,0,0] t6_zf_lang
i: [1,7610,0,30,12,1,0,0] t6_zf_lang
i: [1,7607,0,20,8,1,0,0] t6_zf_lang
i: [1,7605,0,18,8,1,0,0] t6_zf_lang
i: [1,7603,0,17,8,1,0,0] t6_zf_lang
i: [1,7601,0,18,8,1,0,0] t6_zf_lang
i: [1,7599,0,17,8,1,0,0] t6_zf_lang
i: [1,7598,0,44,16,1,0,0] t6_zf_lang
i: [1,7591,0,31,12,1,0,0] t6_zf_lang
i: [1,7589,0,32,12,1,0,0] t6_zf_lang
i: [1,7586,0,31,12,1,0,0] t6_zf_lang
i: [1,7583,0,21,8,1,0,0] t6_zf_lang
i: [1,7581,0,19,8,1,0,0] t6_zf_lang
i: [1,7579,0,18,8,1,0,0] t6_zf_lang
i: [1,7577,0,19,8,1,0,0] t6_zf_lang
i: [1,7575,0,18,8,1,0,0] t6_zf_lang
i: [1,7574,0,43,16,1,0,0] t6_zf_lang
i: [1,7567,0,30,12,1,0,0] t6_zf_lang
i: [1,7565,0,31,12,1,0,0] t6_zf_lang
i: [1,7562,0,30,12,1,0,0] t6_zf_lang
i: [1,7559,0,20,8,1,0,0] t6_zf_lang
i: [1,7557,0,18,8,1,0,0] t6_zf_lang
i: [1,7555,0,17,8,1,0,0] t6_zf_lang
i: [1,7553,0,18,8,1,0,0] t6_zf_lang
i: [1,7551,0,17,8,1,0,0] t6_zf_lang
i: [1,5822,0,55,20,1,0,0] t6_zf_lang
i: [1,5819,0,45,16,1,0,0] t6_zf_lang
i: [1,5817,0,43,16,1,0,0] t6_zf_lang
i: [1,5815,0,42,16,1,0,0] t6_zf_lang
i: [1,5813,0,43,16,1,0,0] t6_zf_lang
i: [1,5811,0,42,16,1,0,0] t6_zf_lang
i: [1,5786,0,45,16,1,0,0] t6_zf_lang
i: [1,5783,0,35,12,1,0,0] t6_zf_lang
i: [1,5781,0,33,12,1,0,0] t6_zf_lang
i: [1,5779,0,32,12,1,0,0] t6_zf_lang
i: [1,5778,0,46,16,1,0,0] t6_zf_lang
i: [1,5777,0,33,12,1,0,0] t6_zf_lang
i: [1,5775,0,32,12,1,0,0] t6_zf_lang
i: [1,5762,0,43,16,1,0,0] t6_zf_lang
i: [1,5759,0,33,12,1,0,0] t6_zf_lang
i: [1,5758,0,44,16,1,0,0] t6_zf_lang
i: [1,5757,0,31,12,1,0,0] t6_zf_lang
i: [1,5755,0,30,12,1,0,0] t6_zf_lang
i: [1,5753,0,31,12,1,0,0] t6_zf_lang
i: [1,5751,0,30,12,1,0,0] t6_zf_lang
i: [1,5738,0,42,16,1,0,0] t6_zf_lang
i: [1,5735,0,32,12,1,0,0] t6_zf_lang
i: [1,5734,0,43,16,1,0,0] t6_zf_lang
i: [1,5733,0,30,12,1,0,0] t6_zf_lang
i: [1,5731,0,29,12,1,0,0] t6_zf_lang
i: [1,5730,0,43,16,1,0,0] t6_zf_lang
i: [1,5729,0,30,12,1,0,0] t6_zf_lang
i: [1,5727,0,29,12,1,0,0] t6_zf_lang
i: [1,5714,0,43,16,1,0,0] t6_zf_lang
i: [1,5712,0,46,16,1,0,0] t6_zf_lang
i: [1,5711,0,33,12,1,0,0] t6_zf_lang
i: [1,5709,0,31,12,1,0,0] t6_zf_lang
i: [1,5708,0,43,16,1,0,0] t6_zf_lang
i: [1,5707,0,30,12,1,0,0] t6_zf_lang
i: [1,5705,0,31,12,1,0,0] t6_zf_lang
i: [1,5704,0,43,16,1,0,0] t6_zf_lang
i: [1,5703,0,30,12,1,0,0] t6_zf_lang
i: [1,5690,0,42,16,1,0,0] t6_zf_lang
i: [1,5687,0,32,12,1,0,0] t6_zf_lang
i: [1,5685,0,30,12,1,0,0] t6_zf_lang
i: [1,5683,0,29,12,1,0,0] t6_zf_lang
i: [1,5682,0,43,16,1,0,0] t6_zf_lang
i: [1,5681,0,30,12,1,0,0] t6_zf_lang
i: [1,5680,0,42,16,1,0,0] t6_zf_lang
i: [1,5679,0,29,12,1,0,0] t6_zf_lang
i: [1,5529,0,46,16,1,0,0] t6_zf_lang
i: [1,5527,0,45,16,1,0,0] t6_zf_lang
i: [1,5390,0,45,16,1,0,0] t6_zf_lang
i: [1,5387,0,35,12,1,0,0] t6_zf_lang
i: [1,5385,0,33,12,1,0,0] t6_zf_lang
i: [1,5383,0,32,12,1,0,0] t6_zf_lang
i: [1,5381,0,33,12,1,0,0] t6_zf_lang
i: [1,5379,0,32,12,1,0,0] t6_zf_lang
i: [1,5366,0,48,16,1,0,0] t6_zf_lang
i: [1,5361,0,36,12,1,0,0] t6_zf_lang
i: [1,5355,0,35,12,1,0,0] t6_zf_lang
i: [1,5354,0,35,12,1,0,0] t6_zf_lang
i: [1,5351,0,25,8,1,0,0] t6_zf_lang
i: [1,5349,0,23,8,1,0,0] t6_zf_lang
i: [1,5347,0,22,8,1,0,0] t6_zf_lang
i: [1,5345,0,23,8,1,0,0] t6_zf_lang
i: [1,5343,0,22,8,1,0,0] t6_zf_lang
i: [1,5337,0,34,12,1,0,0] t6_zf_lang
i: [1,5335,0,33,12,1,0,0] t6_zf_lang
i: [1,5331,0,33,12,1,0,0] t6_zf_lang
i: [1,5330,0,33,12,1,0,0] t6_zf_lang
i: [1,5327,0,23,8,1,0,0] t6_zf_lang
i: [1,5325,0,21,8,1,0,0] t6_zf_lang
i: [1,5323,0,20,8,1,0,0] t6_zf_lang
i: [1,5321,0,21,8,1,0,0] t6_zf_lang
i: [1,5319,0,20,8,1,0,0] t6_zf_lang
i: [1,5318,0,45,16,1,0,0] t6_zf_lang
i: [1,5315,0,35,12,1,0,0] t6_zf_lang
i: [1,5313,0,33,12,1,0,0] t6_zf_lang
i: [1,5311,0,32,12,1,0,0] t6_zf_lang
i: [1,5306,0,32,12,1,0,0] t6_zf_lang
i: [1,5303,0,22,8,1,0,0] t6_zf_lang
i: [1,5301,0,20,8,1,0,0] t6_zf_lang
i: [1,5299,0,19,8,1,0,0] t6_zf_lang
i: [1,5297,0,20,8,1,0,0] t6_zf_lang
i: [1,5295,0,19,8,1,0,0] t6_zf_lang
i: [1,5291,0,36,12,1,0,0] t6_zf_lang
i: [1,5289,0,34,12,1,0,0] t6_zf_lang
i: [1,5287,0,33,12,1,0,0] t6_zf_lang
i: [1,5282,0,33,12,1,0,0] t6_zf_lang
i: [1,5279,0,23,8,1,0,0] t6_zf_lang
i: [1,5277,0,21,8,1,0,0] t6_zf_lang
i: [1,5275,0,20,8,1,0,0] t6_zf_lang
i: [1,5273,0,21,8,1,0,0] t6_zf_lang
i: [1,5271,0,20,8,1,0,0] t6_zf_lang
i: [1,5270,0,45,16,1,0,0] t6_zf_lang
i: [1,5267,0,35,12,1,0,0] t6_zf_lang
i: [1,5263,0,32,12,1,0,0] t6_zf_lang
i: [1,5259,0,32,12,1,0,0] t6_zf_lang
i: [1,5258,0,32,12,1,0,0] t6_zf_lang
i: [1,5255,0,22,8,1,0,0] t6_zf_lang
i: [1,5253,0,20,8,1,0,0] t6_zf_lang
i: [1,5251,0,19,8,1,0,0] t6_zf_lang
i: [1,5249,0,20,8,1,0,0] t6_zf_lang
i: [1,5247,0,19,8,1,0,0] t6_zf_lang
i: [1,5241,0,44,16,1,0,0] t6_zf_lang
i: [1,5239,0,43,16,1,0,0] t6_zf_lang
i: [1,5102,0,43,16,1,0,0] t6_zf_lang
i: [1,5099,0,33,12,1,0,0] t6_zf_lang
i: [1,5097,0,31,12,1,0,0] t6_zf_lang
i: [1,5095,0,30,12,1,0,0] t6_zf_lang
i: [1,5093,0,31,12,1,0,0] t6_zf_lang
i: [1,5091,0,30,12,1,0,0] t6_zf_lang
i: [1,5078,0,46,16,1,0,0] t6_zf_lang
i: [1,5069,0,34,12,1,0,0] t6_zf_lang
i: [1,5066,0,33,12,1,0,0] t6_zf_lang
i: [1,5063,0,23,8,1,0,0] t6_zf_lang
i: [1,5061,0,21,8,1,0,0] t6_zf_lang
i: [1,5059,0,20,8,1,0,0] t6_zf_lang
i: [1,5057,0,21,8,1,0,0] t6_zf_lang
i: [1,5055,0,20,8,1,0,0] t6_zf_lang
i: [1,5054,0,44,16,1,0,0] t6_zf_lang
i: [1,5047,0,31,12,1,0,0] t6_zf_lang
i: [1,5045,0,32,12,1,0,0] t6_zf_lang
i: [1,5042,0,31,12,1,0,0] t6_zf_lang
i: [1,5039,0,21,8,1,0,0] t6_zf_lang
i: [1,5037,0,19,8,1,0,0] t6_zf_lang
i: [1,5035,0,18,8,1,0,0] t6_zf_lang
i: [1,5033,0,19,8,1,0,0] t6_zf_lang
i: [1,5031,0,18,8,1,0,0] t6_zf_lang
i: [1,5030,0,43,16,1,0,0] t6_zf_lang
i: [1,5023,0,30,12,1,0,0] t6_zf_lang
i: [1,5018,0,30,12,1,0,0] t6_zf_lang
i: [1,5015,0,20,8,1,0,0] t6_zf_lang
i: [1,5013,0,18,8,1,0,0] t6_zf_lang
i: [1,5011,0,17,8,1,0,0] t6_zf_lang
i: [1,5009,0,18,8,1,0,0] t6_zf_lang
i: [1,5007,0,17,8,1,0,0] t6_zf_lang
i: [1,5006,0,44,16,1,0,0] t6_zf_lang
i: [1,4995,0,31,12,1,0,0] t6_zf_lang
i: [1,4994,0,31,12,1,0,0] t6_zf_lang
i: [1,4991,0,21,8,1,0,0] t6_zf_lang
i: [1,4989,0,19,8,1,0,0] t6_zf_lang
i: [1,4987,0,18,8,1,0,0] t6_zf_lang
i: [1,4985,0,19,8,1,0,0] t6_zf_lang
i: [1,4983,0,18,8,1,0,0] t6_zf_lang
i: [1,4982,0,43,16,1,0,0] t6_zf_lang
i: [1,4979,0,33,12,1,0,0] t6_zf_lang
i: [1,4977,0,31,12,1,0,0] t6_zf_lang
i: [1,4970,0,30,12,1,0,0] t6_zf_lang
i: [1,4967,0,20,8,1,0,0] t6_zf_lang
i: [1,4965,0,18,8,1,0,0] t6_zf_lang
i: [1,4963,0,17,8,1,0,0] t6_zf_lang
i: [1,4961,0,18,8,1,0,0] t6_zf_lang
i: [1,4959,0,17,8,1,0,0] t6_zf_lang
i: [1,4951,0,42,16,1,0,0] t6_zf_lang
i: [1,4814,0,42,16,1,0,0] t6_zf_lang
i: [1,4811,0,32,12,1,0,0] t6_zf_lang
i: [1,4809,0,30,12,1,0,0] t6_zf_lang
i: [1,4807,0,29,12,1,0,0] t6_zf_lang
i: [1,4805,0,30,12,1,0,0] t6_zf_lang
i: [1,4803,0,29,12,1,0,0] t6_zf_lang
i: [1,4790,0,45,16,1,0,0] t6_zf_lang
i: [1,4787,0,35,12,1,0,0] t6_zf_lang
i: [1,4785,0,33,12,1,0,0] t6_zf_lang
i: [1,4783,0,32,12,1,0,0] t6_zf_lang
i: [1,4781,0,33,12,1,0,0] t6_zf_lang
i: [1,4779,0,32,12,1,0,0] t6_zf_lang
i: [1,4778,0,32,12,1,0,0] t6_zf_lang
i: [1,4775,0,22,8,1,0,0] t6_zf_lang
i: [1,4773,0,20,8,1,0,0] t6_zf_lang
i: [1,4771,0,19,8,1,0,0] t6_zf_lang
i: [1,4769,0,20,8,1,0,0] t6_zf_lang
i: [1,4767,0,19,8,1,0,0] t6_zf_lang
i: [1,4766,0,43,16,1,0,0] t6_zf_lang
i: [1,4763,0,33,12,1,0,0] t6_zf_lang
i: [1,4759,0,30,12,1,0,0] t6_zf_lang
i: [1,4757,0,31,12,1,0,0] t6_zf_lang
i: [1,4755,0,30,12,1,0,0] t6_zf_lang
i: [1,4754,0,30,12,1,0,0] t6_zf_lang
i: [1,4751,0,20,8,1,0,0] t6_zf_lang
i: [1,4749,0,18,8,1,0,0] t6_zf_lang
i: [1,4747,0,17,8,1,0,0] t6_zf_lang
i: [1,4745,0,18,8,1,0,0] t6_zf_lang
i: [1,4743,0,17,8,1,0,0] t6_zf_lang
i: [1,4742,0,42,16,1,0,0] t6_zf_lang
i: [1,4739,0,32,12,1,0,0] t6_zf_lang
i: [1,4735,0,29,12,1,0,0] t6_zf_lang
i: [1,4730,0,29,12,1,0,0] t6_zf_lang
i: [1,4727,0,19,8,1,0,0] t6_zf_lang
i: [1,4725,0,17,8,1,0,0] t6_zf_lang
i: [1,4723,0,16,8,1,0,0] t6_zf_lang
i: [1,4721,0,17,8,1,0,0] t6_zf_lang
i: [1,4719,0,16,8,1,0,0] t6_zf_lang
i: [1,4718,0,43,16,1,0,0] t6_zf_lang
i: [1,4715,0,33,12,1,0,0] t6_zf_lang
i: [1,4709,0,31,12,1,0,0] t6_zf_lang
i: [1,4707,0,30,12,1,0,0] t6_zf_lang
i: [1,4706,0,30,12,1,0,0] t6_zf_lang
i: [1,4703,0,20,8,1,0,0] t6_zf_lang
i: [1,4701,0,18,8,1,0,0] t6_zf_lang
i: [1,4699,0,17,8,1,0,0] t6_zf_lang
i: [1,4697,0,18,8,1,0,0] t6_zf_lang
i: [1,4695,0,17,8,1,0,0] t6_zf_lang
i: [1,4694,0,42,16,1,0,0] t6_zf_lang
i: [1,4691,0,32,12,1,0,0] t6_zf_lang
i: [1,4687,0,29,12,1,0,0] t6_zf_lang
i: [1,4683,0,29,12,1,0,0] t6_zf_lang
i: [1,4682,0,29,12,1,0,0] t6_zf_lang
i: [1,4679,0,19,8,1,0,0] t6_zf_lang
i: [1,4677,0,17,8,1,0,0] t6_zf_lang
i: [1,4675,0,16,8,1,0,0] t6_zf_lang
i: [1,4673,0,17,8,1,0,0] t6_zf_lang
i: [1,4671,0,16,8,1,0,0] t6_zf_lang
i: [1,4663,0,43,16,1,0,0] t6_zf_lang
i: [1,4661,0,44,16,1,0,0] t6_zf_lang
i: [1,4526,0,43,16,1,0,0] t6_zf_lang
i: [1,4523,0,33,12,1,0,0] t6_zf_lang
i: [1,4521,0,31,12,1,0,0] t6_zf_lang
i: [1,4519,0,30,12,1,0,0] t6_zf_lang
i: [1,4517,0,31,12,1,0,0] t6_zf_lang
i: [1,4515,0,30,12,1,0,0] t6_zf_lang
i: [1,4502,0,46,16,1,0,0] t6_zf_lang
i: [1,4499,0,36,12,1,0,0] t6_zf_lang
i: [1,4497,0,34,12,1,0,0] t6_zf_lang
i: [1,4493,0,34,12,1,0,0] t6_zf_lang
i: [1,4491,0,33,12,1,0,0] t6_zf_lang
i: [1,4490,0,33,12,1,0,0] t6_zf_lang
i: [1,4487,0,23,8,1,0,0] t6_zf_lang
i: [1,4485,0,21,8,1,0,0] t6_zf_lang
i: [1,4483,0,20,8,1,0,0] t6_zf_lang
i: [1,4481,0,21,8,1,0,0] t6_zf_lang
i: [1,4479,0,20,8,1,0,0] t6_zf_lang
i: [1,4478,0,44,16,1,0,0] t6_zf_lang
i: [1,4469,0,32,12,1,0,0] t6_zf_lang
i: [1,4467,0,31,12,1,0,0] t6_zf_lang
i: [1,4466,0,31,12,1,0,0] t6_zf_lang
i: [1,4463,0,21,8,1,0,0] t6_zf_lang
i: [1,4461,0,19,8,1,0,0] t6_zf_lang
i: [1,4459,0,18,8,1,0,0] t6_zf_lang
i: [1,4457,0,19,8,1,0,0] t6_zf_lang
i: [1,4455,0,18,8,1,0,0] t6_zf_lang
i: [1,4454,0,43,16,1,0,0] t6_zf_lang
i: [1,4451,0,33,12,1,0,0] t6_zf_lang
i: [1,4449,0,31,12,1,0,0] t6_zf_lang
i: [1,4442,0,30,12,1,0,0] t6_zf_lang
i: [1,4439,0,20,8,1,0,0] t6_zf_lang
i: [1,4437,0,18,8,1,0,0] t6_zf_lang
i: [1,4435,0,17,8,1,0,0] t6_zf_lang
i: [1,4433,0,18,8,1,0,0] t6_zf_lang
i: [1,4431,0,17,8,1,0,0] t6_zf_lang
i: [1,4430,0,44,16,1,0,0] t6_zf_lang
i: [1,4427,0,34,12,1,0,0] t6_zf_lang
i: [1,4421,0,32,12,1,0,0] t6_zf_lang
i: [1,4418,0,31,12,1,0,0] t6_zf_lang
i: [1,4415,0,21,8,1,0,0] t6_zf_lang
i: [1,4413,0,19,8,1,0,0] t6_zf_lang
i: [1,4411,0,18,8,1,0,0] t6_zf_lang
i: [1,4410,0,32,12,1,0,0] t6_zf_lang
i: [1,4409,0,19,8,1,0,0] t6_zf_lang
i: [1,4407,0,18,8,1,0,0] t6_zf_lang
i: [1,4406,0,43,16,1,0,0] t6_zf_lang
i: [1,4401,0,31,12,1,0,0] t6_zf_lang
i: [1,4397,0,31,12,1,0,0] t6_zf_lang
i: [1,4394,0,30,12,1,0,0] t6_zf_lang
i: [1,4391,0,20,8,1,0,0] t6_zf_lang
i: [1,4389,0,18,8,1,0,0] t6_zf_lang
i: [1,4387,0,17,8,1,0,0] t6_zf_lang
i: [1,4385,0,18,8,1,0,0] t6_zf_lang
i: [1,4383,0,17,8,1,0,0] t6_zf_lang
i: [1,4377,0,43,16,1,0,0] t6_zf_lang
i: [1,4238,0,42,16,1,0,0] t6_zf_lang
i: [1,4235,0,32,12,1,0,0] t6_zf_lang
i: [1,4233,0,30,12,1,0,0] t6_zf_lang
i: [1,4231,0,29,12,1,0,0] t6_zf_lang
i: [1,4229,0,30,12,1,0,0] t6_zf_lang
i: [1,4227,0,29,12,1,0,0] t6_zf_lang
i: [1,4214,0,45,16,1,0,0] t6_zf_lang
i: [1,4209,0,33,12,1,0,0] t6_zf_lang
i: [1,4203,0,32,12,1,0,0] t6_zf_lang
i: [1,4202,0,32,12,1,0,0] t6_zf_lang
i: [1,4199,0,22,8,1,0,0] t6_zf_lang
i: [1,4197,0,20,8,1,0,0] t6_zf_lang
i: [1,4195,0,19,8,1,0,0] t6_zf_lang
i: [1,4193,0,20,8,1,0,0] t6_zf_lang
i: [1,4191,0,19,8,1,0,0] t6_zf_lang
i: [1,4190,0,43,16,1,0,0] t6_zf_lang
i: [1,4185,0,31,12,1,0,0] t6_zf_lang
i: [1,4183,0,30,12,1,0,0] t6_zf_lang
i: [1,4178,0,30,12,1,0,0] t6_zf_lang
i: [1,4175,0,20,8,1,0,0] t6_zf_lang
i: [1,4173,0,18,8,1,0,0] t6_zf_lang
i: [1,4171,0,17,8,1,0,0] t6_zf_lang
i: [1,4169,0,18,8,1,0,0] t6_zf_lang
i: [1,4167,0,17,8,1,0,0] t6_zf_lang
i: [1,4166,0,42,16,1,0,0] t6_zf_lang
i: [1,4161,0,30,12,1,0,0] t6_zf_lang
i: [1,4154,0,29,12,1,0,0] t6_zf_lang
i: [1,4151,0,19,8,1,0,0] t6_zf_lang
i: [1,4149,0,17,8,1,0,0] t6_zf_lang
i: [1,4147,0,16,8,1,0,0] t6_zf_lang
i: [1,4145,0,17,8,1,0,0] t6_zf_lang
i: [1,4143,0,16,8,1,0,0] t6_zf_lang
i: [1,4139,0,33,12,1,0,0] t6_zf_lang
i: [1,4137,0,31,12,1,0,0] t6_zf_lang
i: [1,4133,0,31,12,1,0,0] t6_zf_lang
i: [1,4131,0,30,12,1,0,0] t6_zf_lang
i: [1,4130,0,30,12,1,0,0] t6_zf_lang
i: [1,4127,0,20,8,1,0,0] t6_zf_lang
i: [1,4125,0,18,8,1,0,0] t6_zf_lang
i: [1,4123,0,17,8,1,0,0] t6_zf_lang
i: [1,4121,0,18,8,1,0,0] t6_zf_lang
i: [1,4119,0,17,8,1,0,0] t6_zf_lang
i: [1,4118,0,42,16,1,0,0] t6_zf_lang
i: [1,4106,0,29,12,1,0,0] t6_zf_lang
i: [1,4103,0,19,8,1,0,0] t6_zf_lang
i: [1,4101,0,17,8,1,0,0] t6_zf_lang
i: [1,4099,0,16,8,1,0,0] t6_zf_lang
i: [1,4097,0,17,8,1,0,0] t6_zf_lang
i: [1,4095,0,16,8,1,0,0] t6_zf_lang
i: [1,2366,0,57,20,1,0,0] t6_zf_lang
i: [1,2363,0,47,16,1,0,0] t6_zf_lang
i: [1,2361,0,45,16,1,0,0] t6_zf_lang
i: [1,2359,0,44,16,1,0,0] t6_zf_lang
i: [1,2357,0,45,16,1,0,0] t6_zf_lang
i: [1,2355,0,44,16,1,0,0] t6_zf_lang
i: [1,2330,0,47,16,1,0,0] t6_zf_lang
i: [1,2328,0,50,16,1,0,0] t6_zf_lang
i: [1,2327,0,37,12,1,0,0] t6_zf_lang
i: [1,2325,0,35,12,1,0,0] t6_zf_lang
i: [1,2324,0,47,16,1,0,0] t6_zf_lang
i: [1,2323,0,34,12,1,0,0] t6_zf_lang
i: [1,2322,0,48,16,1,0,0] t6_zf_lang
i: [1,2321,0,35,12,1,0,0] t6_zf_lang
i: [1,2320,0,47,16,1,0,0] t6_zf_lang
i: [1,2319,0,34,12,1,0,0] t6_zf_lang
i: [1,2306,0,45,16,1,0,0] t6_zf_lang
i: [1,2304,0,48,16,1,0,0] t6_zf_lang
i: [1,2303,0,35,12,1,0,0] t6_zf_lang
i: [1,2302,0,46,16,1,0,0] t6_zf_lang
i: [1,2301,0,33,12,1,0,0] t6_zf_lang
i: [1,2300,0,45,16,1,0,0] t6_zf_lang
i: [1,2299,0,32,12,1,0,0] t6_zf_lang
i: [1,2297,0,33,12,1,0,0] t6_zf_lang
i: [1,2296,0,45,16,1,0,0] t6_zf_lang
i: [1,2295,0,32,12,1,0,0] t6_zf_lang
i: [1,2282,0,44,16,1,0,0] t6_zf_lang
i: [1,2280,0,47,16,1,0,0] t6_zf_lang
i: [1,2279,0,34,12,1,0,0] t6_zf_lang
i: [1,2278,0,45,16,1,0,0] t6_zf_lang
i: [1,2277,0,32,12,1,0,0] t6_zf_lang
i: [1,2276,0,44,16,1,0,0] t6_zf_lang
i: [1,2275,0,31,12,1,0,0] t6_zf_lang
i: [1,2274,0,45,16,1,0,0] t6_zf_lang
i: [1,2273,0,32,12,1,0,0] t6_zf_lang
i: [1,2272,0,44,16,1,0,0] t6_zf_lang
i: [1,2271,0,31,12,1,0,0] t6_zf_lang
i: [1,2258,0,45,16,1,0,0] t6_zf_lang
i: [1,2256,0,48,16,1,0,0] t6_zf_lang
i: [1,2255,0,35,12,1,0,0] t6_zf_lang
i: [1,2253,0,33,12,1,0,0] t6_zf_lang
i: [1,2252,0,45,16,1,0,0] t6_zf_lang
i: [1,2251,0,32,12,1,0,0] t6_zf_lang
i: [1,2249,0,33,12,1,0,0] t6_zf_lang
i: [1,2248,0,45,16,1,0,0] t6_zf_lang
i: [1,2247,0,32,12,1,0,0] t6_zf_lang
i: [1,2234,0,44,16,1,0,0] t6_zf_lang
i: [1,2232,0,47,16,1,0,0] t6_zf_lang
i: [1,2231,0,34,12,1,0,0] t6_zf_lang
i: [1,2230,0,45,16,1,0,0] t6_zf_lang
i: [1,2229,0,32,12,1,0,0] t6_zf_lang
i: [1,2228,0,44,16,1,0,0] t6_zf_lang
i: [1,2227,0,31,12,1,0,0] t6_zf_lang
i: [1,2226,0,45,16,1,0,0] t6_zf_lang
i: [1,2225,0,32,12,1,0,0] t6_zf_lang
i: [1,2224,0,44,16,1,0,0] t6_zf_lang
i: [1,2223,0,31,12,1,0,0] t6_zf_lang
i: [1,2073,0,48,16,1,0,0] t6_zf_lang
i: [1,2071,0,47,16,1,0,0] t6_zf_lang
i: [1,1934,0,47,16,1,0,0] t6_zf_lang
i: [1,1931,0,37,12,1,0,0] t6_zf_lang
i: [1,1929,0,35,12,1,0,0] t6_zf_lang
i: [1,1927,0,34,12,1,0,0] t6_zf_lang
i: [1,1925,0,35,12,1,0,0] t6_zf_lang
i: [1,1923,0,34,12,1,0,0] t6_zf_lang
i: [1,1907,0,40,12,1,0,0] t6_zf_lang
i: [1,1905,0,38,12,1,0,0] t6_zf_lang
i: [1,1903,0,37,12,1,0,0] t6_zf_lang
i: [1,1901,0,38,12,1,0,0] t6_zf_lang
i: [1,1898,0,37,12,1,0,0] t6_zf_lang
i: [1,1895,0,27,8,1,0,0] t6_zf_lang
i: [1,1893,0,25,8,1,0,0] t6_zf_lang
i: [1,1891,0,24,8,1,0,0] t6_zf_lang
i: [1,1889,0,25,8,1,0,0] t6_zf_lang
i: [1,1887,0,24,8,1,0,0] t6_zf_lang
i: [1,1881,0,36,12,1,0,0] t6_zf_lang
i: [1,1874,0,35,12,1,0,0] t6_zf_lang
i: [1,1871,0,25,8,1,0,0] t6_zf_lang
i: [1,1869,0,23,8,1,0,0] t6_zf_lang
i: [1,1867,0,22,8,1,0,0] t6_zf_lang
i: [1,1865,0,23,8,1,0,0] t6_zf_lang
i: [1,1863,0,22,8,1,0,0] t6_zf_lang
i: [1,1859,0,37,12,1,0,0] t6_zf_lang
i: [1,1857,0,35,12,1,0,0] t6_zf_lang
i: [1,1855,0,34,12,1,0,0] t6_zf_lang
i: [1,1850,0,34,12,1,0,0] t6_zf_lang
i: [1,1847,0,24,8,1,0,0] t6_zf_lang
i: [1,1845,0,22,8,1,0,0] t6_zf_lang
i: [1,1843,0,21,8,1,0,0] t6_zf_lang
i: [1,1841,0,22,8,1,0,0] t6_zf_lang
i: [1,1839,0,21,8,1,0,0] t6_zf_lang
i: [1,1838,0,48,16,1,0,0] t6_zf_lang
i: [1,1833,0,36,12,1,0,0] t6_zf_lang
i: [1,1829,0,36,12,1,0,0] t6_zf_lang
i: [1,1827,0,35,12,1,0,0] t6_zf_lang
i: [1,1826,0,35,12,1,0,0] t6_zf_lang
i: [1,1823,0,25,8,1,0,0] t6_zf_lang
i: [1,1821,0,23,8,1,0,0] t6_zf_lang
i: [1,1819,0,22,8,1,0,0] t6_zf_lang
i: [1,1817,0,23,8,1,0,0] t6_zf_lang
i: [1,1815,0,22,8,1,0,0] t6_zf_lang
i: [1,1811,0,37,12,1,0,0] t6_zf_lang
i: [1,1807,0,34,12,1,0,0] t6_zf_lang
i: [1,1805,0,35,12,1,0,0] t6_zf_lang
i: [1,1803,0,34,12,1,0,0] t6_zf_lang
i: [1,1802,0,34,12,1,0,0] t6_zf_lang
i: [1,1799,0,24,8,1,0,0] t6_zf_lang
i: [1,1797,0,22,8,1,0,0] t6_zf_lang
i: [1,1795,0,21,8,1,0,0] t6_zf_lang
i: [1,1793,0,22,8,1,0,0] t6_zf_lang
i: [1,1791,0,21,8,1,0,0] t6_zf_lang
i: [1,1646,0,45,16,1,0,0] t6_zf_lang
i: [1,1643,0,35,12,1,0,0] t6_zf_lang
i: [1,1641,0,33,12,1,0,0] t6_zf_lang
i: [1,1639,0,32,12,1,0,0] t6_zf_lang
i: [1,1637,0,33,12,1,0,0] t6_zf_lang
i: [1,1635,0,32,12,1,0,0] t6_zf_lang
i: [1,1617,0,36,12,1,0,0] t6_zf_lang
i: [1,1610,0,35,12,1,0,0] t6_zf_lang
i: [1,1607,0,25,8,1,0,0] t6_zf_lang
i: [1,1605,0,23,8,1,0,0] t6_zf_lang
i: [1,1603,0,22,8,1,0,0] t6_zf_lang
i: [1,1601,0,23,8,1,0,0] t6_zf_lang
i: [1,1599,0,22,8,1,0,0] t6_zf_lang
i: [1,1595,0,36,12,1,0,0] t6_zf_lang
i: [1,1593,0,34,12,1,0,0] t6_zf_lang
i: [1,1589,0,34,12,1,0,0] t6_zf_lang
i: [1,1586,0,33,12,1,0,0] t6_zf_lang
i: [1,1583,0,23,8,1,0,0] t6_zf_lang
i: [1,1581,0,21,8,1,0,0] t6_zf_lang
i: [1,1579,0,20,8,1,0,0] t6_zf_lang
i: [1,1577,0,21,8,1,0,0] t6_zf_lang
i: [1,1575,0,20,8,1,0,0] t6_zf_lang
i: [1,1574,0,45,16,1,0,0] t6_zf_lang
i: [1,1571,0,35,12,1,0,0] t6_zf_lang
i: [1,1567,0,32,12,1,0,0] t6_zf_lang
i: [1,1562,0,32,12,1,0,0] t6_zf_lang
i: [1,1559,0,22,8,1,0,0] t6_zf_lang
i: [1,1557,0,20,8,1,0,0] t6_zf_lang
i: [1,1555,0,19,8,1,0,0] t6_zf_lang
i: [1,1553,0,20,8,1,0,0] t6_zf_lang
i: [1,1551,0,19,8,1,0,0] t6_zf_lang
i: [1,1547,0,36,12,1,0,0] t6_zf_lang
i: [1,1545,0,34,12,1,0,0] t6_zf_lang
i: [1,1541,0,34,12,1,0,0] t6_zf_lang
i: [1,1539,0,33,12,1,0,0] t6_zf_lang
i: [1,1538,0,33,12,1,0,0] t6_zf_lang
i: [1,1535,0,23,8,1,0,0] t6_zf_lang
i: [1,1533,0,21,8,1,0,0] t6_zf_lang
i: [1,1531,0,20,8,1,0,0] t6_zf_lang
i: [1,1529,0,21,8,1,0,0] t6_zf_lang
i: [1,1527,0,20,8,1,0,0] t6_zf_lang
i: [1,1517,0,33,12,1,0,0] t6_zf_lang
i: [1,1515,0,32,12,1,0,0] t6_zf_lang
i: [1,1514,0,32,12,1,0,0] t6_zf_lang
i: [1,1511,0,22,8,1,0,0] t6_zf_lang
i: [1,1509,0,20,8,1,0,0] t6_zf_lang
i: [1,1507,0,19,8,1,0,0] t6_zf_lang
i: [1,1505,0,20,8,1,0,0] t6_zf_lang
i: [1,1503,0,19,8,1,0,0] t6_zf_lang
i: [1,1495,0,44,16,1,0,0] t6_zf_lang
i: [1,1358,0,44,16,1,0,0] t6_zf_lang
i: [1,1355,0,34,12,1,0,0] t6_zf_lang
i: [1,1353,0,32,12,1,0,0] t6_zf_lang
i: [1,1351,0,31,12,1,0,0] t6_zf_lang
i: [1,1349,0,32,12,1,0,0] t6_zf_lang
i: [1,1347,0,31,12,1,0,0] t6_zf_lang
i: [1,1331,0,37,12,1,0,0] t6_zf_lang
i: [1,1329,0,35,12,1,0,0] t6_zf_lang
i: [1,1327,0,34,12,1,0,0] t6_zf_lang
i: [1,1325,0,35,12,1,0,0] t6_zf_lang
i: [1,1323,0,34,12,1,0,0] t6_zf_lang
i: [1,1322,0,34,12,1,0,0] t6_zf_lang
i: [1,1319,0,24,8,1,0,0] t6_zf_lang
i: [1,1317,0,22,8,1,0,0] t6_zf_lang
i: [1,1315,0,21,8,1,0,0] t6_zf_lang
i: [1,1313,0,22,8,1,0,0] t6_zf_lang
i: [1,1311,0,21,8,1,0,0] t6_zf_lang
i: [1,1307,0,35,12,1,0,0] t6_zf_lang
i: [1,1305,0,33,12,1,0,0] t6_zf_lang
i: [1,1303,0,32,12,1,0,0] t6_zf_lang
i: [1,1301,0,33,12,1,0,0] t6_zf_lang
i: [1,1299,0,32,12,1,0,0] t6_zf_lang
i: [1,1298,0,32,12,1,0,0] t6_zf_lang
i: [1,1295,0,22,8,1,0,0] t6_zf_lang
i: [1,1293,0,20,8,1,0,0] t6_zf_lang
i: [1,1291,0,19,8,1,0,0] t6_zf_lang
i: [1,1289,0,20,8,1,0,0] t6_zf_lang
i: [1,1287,0,19,8,1,0,0] t6_zf_lang
i: [1,1286,0,44,16,1,0,0] t6_zf_lang
i: [1,1283,0,34,12,1,0,0] t6_zf_lang
i: [1,1279,0,31,12,1,0,0] t6_zf_lang
i: [1,1275,0,31,12,1,0,0] t6_zf_lang
i: [1,1274,0,31,12,1,0,0] t6_zf_lang
i: [1,1271,0,21,8,1,0,0] t6_zf_lang
i: [1,1269,0,19,8,1,0,0] t6_zf_lang
i: [1,1267,0,18,8,1,0,0] t6_zf_lang
i: [1,1265,0,19,8,1,0,0] t6_zf_lang
i: [1,1263,0,18,8,1,0,0] t6_zf_lang
i: [1,1262,0,45,16,1,0,0] t6_zf_lang
i: [1,1259,0,35,12,1,0,0] t6_zf_lang
i: [1,1253,0,33,12,1,0,0] t6_zf_lang
i: [1,1251,0,32,12,1,0,0] t6_zf_lang
i: [1,1250,0,32,12,1,0,0] t6_zf_lang
i: [1,1247,0,22,8,1,0,0] t6_zf_lang
i: [1,1245,0,20,8,1,0,0] t6_zf_lang
i: [1,1243,0,19,8,1,0,0] t6_zf_lang
i: [1,1241,0,20,8,1,0,0] t6_zf_lang
i: [1,1239,0,19,8,1,0,0] t6_zf_lang
i: [1,1238,0,44,16,1,0,0] t6_zf_lang
i: [1,1235,0,34,12,1,0,0] t6_zf_lang
i: [1,1231,0,31,12,1,0,0] t6_zf_lang
i: [1,1227,0,31,12,1,0,0] t6_zf_lang
i: [1,1226,0,31,12,1,0,0] t6_zf_lang
i: [1,1223,0,21,8,1,0,0] t6_zf_lang
i: [1,1221,0,19,8,1,0,0] t6_zf_lang
i: [1,1219,0,18,8,1,0,0] t6_zf_lang
i: [1,1217,0,19,8,1,0,0] t6_zf_lang
i: [1,1215,0,18,8,1,0,0] t6_zf_lang
i: [1,1205,0,46,16,1,0,0] t6_zf_lang
i: [1,1070,0,45,16,1,0,0] t6_zf_lang
i: [1,1067,0,35,12,1,0,0] t6_zf_lang
i: [1,1065,0,33,12,1,0,0] t6_zf_lang
i: [1,1063,0,32,12,1,0,0] t6_zf_lang
i: [1,1061,0,33,12,1,0,0] t6_zf_lang
i: [1,1059,0,32,12,1,0,0] t6_zf_lang
i: [1,1043,0,38,12,1,0,0] t6_zf_lang
i: [1,1041,0,36,12,1,0,0] t6_zf_lang
i: [1,1039,0,35,12,1,0,0] t6_zf_lang
i: [1,1035,0,35,12,1,0,0] t6_zf_lang
i: [1,1034,0,35,12,1,0,0] t6_zf_lang
i: [1,1031,0,25,8,1,0,0] t6_zf_lang
i: [1,1029,0,23,8,1,0,0] t6_zf_lang
i: [1,1027,0,22,8,1,0,0] t6_zf_lang
i: [1,1025,0,23,8,1,0,0] t6_zf_lang
i: [1,1023,0,22,8,1,0,0] t6_zf_lang
i: [1,1019,0,36,12,1,0,0] t6_zf_lang
i: [1,1017,0,34,12,1,0,0] t6_zf_lang
i: [1,1013,0,34,12,1,0,0] t6_zf_lang
i: [1,1011,0,33,12,1,0,0] t6_zf_lang
i: [1,1010,0,33,12,1,0,0] t6_zf_lang
i: [1,1007,0,23,8,1,0,0] t6_zf_lang
i: [1,1005,0,21,8,1,0,0] t6_zf_lang
i: [1,1003,0,20,8,1,0,0] t6_zf_lang
i: [1,1001,0,21,8,1,0,0] t6_zf_lang
i: [1,999,0,20,8,1,0,0] t6_zf_lang
i: [1,991,0,32,12,1,0,0] t6_zf_lang
i: [1,986,0,32,12,1,0,0] t6_zf_lang
i: [1,983,0,22,8,1,0,0] t6_zf_lang
i: [1,981,0,20,8,1,0,0] t6_zf_lang
i: [1,979,0,19,8,1,0,0] t6_zf_lang
i: [1,977,0,20,8,1,0,0] t6_zf_lang
i: [1,975,0,19,8,1,0,0] t6_zf_lang
i: [1,971,0,36,12,1,0,0] t6_zf_lang
i: [1,969,0,34,12,1,0,0] t6_zf_lang
i: [1,965,0,34,12,1,0,0] t6_zf_lang
i: [1,962,0,33,12,1,0,0] t6_zf_lang
i: [1,959,0,23,8,1,0,0] t6_zf_lang
i: [1,957,0,21,8,1,0,0] t6_zf_lang
i: [1,955,0,20,8,1,0,0] t6_zf_lang
i: [1,953,0,21,8,1,0,0] t6_zf_lang
i: [1,951,0,20,8,1,0,0] t6_zf_lang
i: [1,943,0,32,12,1,0,0] t6_zf_lang
i: [1,941,0,33,12,1,0,0] t6_zf_lang
i: [1,938,0,32,12,1,0,0] t6_zf_lang
i: [1,935,0,22,8,1,0,0] t6_zf_lang
i: [1,933,0,20,8,1,0,0] t6_zf_lang
i: [1,931,0,19,8,1,0,0] t6_zf_lang
i: [1,929,0,20,8,1,0,0] t6_zf_lang
i: [1,927,0,19,8,1,0,0] t6_zf_lang
i: [1,923,0,47,16,1,0,0] t6_zf_lang
i: [1,919,0,44,16,1,0,0] t6_zf_lang
i: [1,782,0,44,16,1,0,0] t6_zf_lang
i: [1,779,0,34,12,1,0,0] t6_zf_lang
i: [1,777,0,32,12,1,0,0] t6_zf_lang
i: [1,775,0,31,12,1,0,0] t6_zf_lang
i: [1,773,0,32,12,1,0,0] t6_zf_lang
i: [1,771,0,31,12,1,0,0] t6_zf_lang
i: [1,758,0,47,16,1,0,0] t6_zf_lang
i: [1,753,0,35,12,1,0,0] t6_zf_lang
i: [1,746,0,34,12,1,0,0] t6_zf_lang
i: [1,743,0,24,8,1,0,0] t6_zf_lang
i: [1,741,0,22,8,1,0,0] t6_zf_lang
i: [1,739,0,21,8,1,0,0] t6_zf_lang
i: [1,737,0,22,8,1,0,0] t6_zf_lang
i: [1,735,0,21,8,1,0,0] t6_zf_lang
i: [1,734,0,45,16,1,0,0] t6_zf_lang
i: [1,729,0,33,12,1,0,0] t6_zf_lang
i: [1,727,0,32,12,1,0,0] t6_zf_lang
i: [1,722,0,32,12,1,0,0] t6_zf_lang
i: [1,719,0,22,8,1,0,0] t6_zf_lang
i: [1,717,0,20,8,1,0,0] t6_zf_lang
i: [1,715,0,19,8,1,0,0] t6_zf_lang
i: [1,713,0,20,8,1,0,0] t6_zf_lang
i: [1,711,0,19,8,1,0,0] t6_zf_lang
i: [1,705,0,32,12,1,0,0] t6_zf_lang
i: [1,698,0,31,12,1,0,0] t6_zf_lang
i: [1,695,0,21,8,1,0,0] t6_zf_lang
i: [1,693,0,19,8,1,0,0] t6_zf_lang
i: [1,691,0,18,8,1,0,0] t6_zf_lang
i: [1,689,0,19,8,1,0,0] t6_zf_lang
i: [1,687,0,18,8,1,0,0] t6_zf_lang
i: [1,686,0,45,16,1,0,0] t6_zf_lang
i: [1,683,0,35,12,1,0,0] t6_zf_lang
i: [1,679,0,32,12,1,0,0] t6_zf_lang
i: [1,674,0,32,12,1,0,0] t6_zf_lang
i: [1,671,0,22,8,1,0,0] t6_zf_lang
i: [1,669,0,20,8,1,0,0] t6_zf_lang
i: [1,667,0,19,8,1,0,0] t6_zf_lang
i: [1,665,0,20,8,1,0,0] t6_zf_lang
i: [1,663,0,19,8,1,0,0] t6_zf_lang
i: [1,659,0,34,12,1,0,0] t6_zf_lang
i: [1,650,0,31,12,1,0,0] t6_zf_lang
i: [1,647,0,21,8,1,0,0] t6_zf_lang
i: [1,645,0,19,8,1,0,0] t6_zf_lang
i: [1,643,0,18,8,1,0,0] t6_zf_lang
i: [1,641,0,19,8,1,0,0] t6_zf_lang
i: [1,639,0,18,8,1,0,0] t6_zf_lang
d: [2,34212,2,4,2,2,21,34207,30396]
d: [2,34213,2,2,1,3,21,34212,30401]
d: [2,33139,1,13,5,2,44,18312,18342]
d: [2,33140,1,10,4,2,19,33139]
d: [2,34215,1,8,3,3,21,33140,18366]
d: [2,34216,1,6,2,4,21,34215,18354]
d: [2,34808,1,4,2,6,21,34803,18316]
d: [2,34809,1,2,1,7,21,34808,30401]
d: [2,33141,1,13,5,2,44,18319,18372]
d: [2,33142,1,10,4,2,19,33141]
d: [2,34825,1,8,3,3,21,33142,18476]
d: [2,34826,1,6,2,4,21,34825,18424]
d: [2,34879,1,17,7,1,44,18306,18502]
d: [2,34929,1,14,6,2,21,34879,30396]
d: [2,34930,1,12,5,3,21,34929,18994]
d: [2,34931,1,10,4,4,21,34930,19486]
d: [2,34932,1,8,3,5,21,34931,18330]
d: [2,34933,1,6,2,6,21,34932,30401]
d: [2,33408,1,13,5,2,44,18326,18502]
d: [2,33409,1,10,4,2,19,33408]
d: [2,34935,1,8,3,3,21,33409,19486]
d: [2,34936,1,6,2,4,21,34935,18994]
d: [2,34785,1,16,6,1,44,18308,19732]
d: [2,34915,1,13,5,2,21,34785,30396]
d: [2,34938,1,9,4,3,21,34915,27730]
d: [2,34939,1,7,3,4,21,34938,18337]
d: [2,34940,1,5,2,5,21,34939,30401]
d: [2,34193,2,7,3,1,47,34192]
d: [2,34942,2,4,2,2,21,34193,30396]
d: [2,34943,2,2,1,3,21,34942,30401]
d: [2,34949,1,12,5,3,21,34894,18476]
d: [2,35233,1,10,4,4,21,34949,18424]
d: [2,35234,1,8,3,5,21,35233,30401]
f: 35294,35234,35231,34943,34940,34936,34933,34826,34809,34216,34213,637,638,639,641,643,645,647,650,659,663,665,667,669,671,674,679,683,686,687,689,691,693,695,698,705,711,713,715,717,719,722,727,729,734,735,737,739,741,743,746,753,758,771,773,775,777,779,782,919,923,927,929,931,933,935,938,941,943,951,953,955,957,959,962,965,969,971,975,977,979,981,983,986,991,999,1001,1003,1005,1007,1010,1011,1013,1017,1019,1023,1025,1027,1029,1031,1034,1035,1039,1041,1043,1059,1061,1063,1065,1067,1070,1205,1215,1217,1219,1221,1223,1226,1227,1231,1235,1238,1239,1241,1243,1245,1247,1250,1251,1253,1259,1262,1263,1265,1267,1269,1271,1274,1275,1279,1283,1286,1287,1289,1291,1293,1295,1298,1299,1301,1303,1305,1307,1311,1313,1315,1317,1319,1322,1323,1325,1327,1329,1331,1347,1349,1351,1353,1355,1358,1495,1503,1505,1507,1509,1511,1514,1515,1517,1527,1529,1531,1533,1535,1538,1539,1541,1545,1547,1551,1553,1555,1557,1559,1562,1567,1571,1574,1575,1577,1579,1581,1583,1586,1589,1593,1595,1599,1601,1603,1605,1607,1610,1617,1635,1637,1639,1641,1643,1646,1791,1793,1795,1797,1799,1802,1803,1805,1807,1811,1815,1817,1819,1821,1823,1826,1827,1829,1833,1838,1839,1841,1843,1845,1847,1850,1855,1857,1859,1863,1865,1867,1869,1871,1874,1881,1887,1889,1891,1893,1895,1898,1901,1903,1905,1907,1923,1925,1927,1929,1931,1934,2071,2073,2223,2224,2225,2226,2227,2228,2229,2230,2231,2232,2234,2247,2248,2249,2251,2252,2253,2255,2256,2258,2271,2272,2273,2274,2275,2276,2277,2278,2279,2280,2282,2295,2296,2297,2299,2300,2301,2302,2303,2304,2306,2319,2320,2321,2322,2323,2324,2325,2327,2328,2330,2355,2357,2359,2361,2363,2366,4095,4097,4099,4101,4103,4106,4118,4119,4121,4123,4125,4127,4130,4131,4133,4137,4139,4143,4145,4147,4149,4151,4154,4161,4166,4167,4169,4171,4173,4175,4178,4183,4185,4190,4191,4193,4195,4197,4199,4202,4203,4209,4214,4227,4229,4231,4233,4235,4238,4377,4383,4385,4387,4389,4391,4394,4397,4401,4406,4407,4409,4410,4411,4413,4415,4418,4421,4427,4430,4431,4433,4435,4437,4439,4442,4449,4451,4454,4455,4457,4459,4461,4463,4466,4467,4469,4478,4479,4481,4483,4485,4487,4490,4491,4493,4497,4499,4502,4515,4517,4519,4521,4523,4526,4661,4663,4671,4673,4675,4677,4679,4682,4683,4687,4691,4694,4695,4697,4699,4701,4703,4706,4707,4709,4715,4718,4719,4721,4723,4725,4727,4730,4735,4739,4742,4743,4745,4747,4749,4751,4754,4755,4757,4759,4763,4766,4767,4769,4771,4773,4775,4778,4779,4781,4783,4785,4787,4790,4803,4805,4807,4809,4811,4814,4951,4959,4961,4963,4965,4967,4970,4977,4979,4982,4983,4985,4987,4989,4991,4994,4995,5006,5007,5009,5011,5013,5015,5018,5023,5030,5031,5033,5035,5037,5039,5042,5045,5047,5054,5055,5057,5059,5061,5063,5066,5069,5078,5091,5093,5095,5097,5099,5102,5239,5241,5247,5249,5251,5253,5255,5258,5259,5263,5267,5270,5271,5273,5275,5277,5279,5282,5287,5289,5291,5295,5297,5299,5301,5303,5306,5311,5313,5315,5318,5319,5321,5323,5325,5327,5330,5331,5335,5337,5343,5345,5347,5349,5351,5354,5355,5361,5366,5379,5381,5383,5385,5387,5390,5527,5529,5679,5680,5681,5682,5683,5685,5687,5690,5703,5704,5705,5707,5708,5709,5711,5712,5714,5727,5729,5730,5731,5733,5734,5735,5738,5751,5753,5755,5757,5758,5759,5762,5775,5777,5778,5779,5781,5783,5786,5811,5813,5815,5817,5819,5822,7551,7553,7555,7557,7559,7562,7565,7567,7574,7575,7577,7579,7581,7583,7586,7589,7591,7598,7599,7601,7603,7605,7607,7610,7622,7623,7625,7627,7629,7631,7634,7643,7646,7647,7649,7651,7653,7655,7658,7659,7670,7683,7685,7687,7689,7691,7694,7831,7833,7835,7839,7841,7843,7845,7847,7850,7862,7863,7865,7867,7869,7871,7874,7875,7877,7879,7881,7883,7887,7889,7891,7893,7895,7898,7911,7913,7915,7917,7919,7922,7925,7927,7929,7931,7935,7937,7939,7941,7943,7946,7949,7951,7953,7971,7973,7975,7977,7979,7982,8117,8119,8121,8123,8127,8129,8131,8133,8135,8138,8141,8143,8147,8150,8151,8153,8155,8157,8159,8162,8165,8171,8174,8175,8177,8179,8181,8183,8186,8199,8201,8203,8205,8207,8210,8211,8213,8215,8219,8222,8223,8225,8227,8229,8231,8234,8237,8239,8241,8243,8246,8259,8261,8263,8265,8267,8270,8407,8409,8415,8417,8419,8421,8423,8426,8431,8433,8438,8439,8441,8443,8445,8447,8450,8453,8455,8457,8463,8465,8467,8469,8471,8474,8486,8487,8489,8491,8493,8495,8498,8499,8505,8507,8510,8511,8513,8515,8517,8519,8522,8523,8525,8527,8529,8534,8547,8549,8551,8553,8555,8558,8695,8697,8703,8705,8707,8709,8711,8714,8715,8721,8726,8727,8729,8731,8733,8735,8738,8741,8747,8750,8751,8753,8755,8757,8759,8762,8774,8775,8777,8779,8781,8783,8786,8787,8795,8798,8799,8801,8803,8805,8807,8810,8811,8815,8819,8822,8835,8837,8839,8841,8843,8846,8983,8985,9135,9137,9138,9139,9141,9143,9146,9159,9161,9163,9165,9167,9170,9183,9184,9185,9186,9187,9188,9189,9190,9191,9192,9194,9207,9209,9211,9213,9214,9215,9218,9231,9232,9233,9234,9235,9236,9237,9238,9239,9240,9242,9267,9269,9271,9273,9275,9278,11007,11009,11011,11013,11015,11018,11019,11023,11027,11031,11033,11035,11037,11039,11042,11043,11045,11047,11051,11054,11055,11057,11059,11061,11063,11066,11067,11071,11075,11079,11081,11083,11085,11087,11090,11099,11102,11103,11105,11107,11109,11111,11114,11115,11126,11139,11141,11143,11145,11147,11150,11295,11297,11299,11301,11303,11306,11307,11309,11313,11315,11319,11321,11323,11325,11327,11330,11333,11337,11343,11345,11347,11349,11351,11354,11355,11363,11367,11369,11371,11373,11375,11378,11379,11381,11383,11385,11387,11391,11393,11395,11397,11399,11402,11403,11409,11427,11429,11431,11433,11435,11438,11573,11583,11585,11587,11589,11591,11594,11597,11606,11607,11609,11611,11613,11615,11618,11630,11631,11633,11635,11637,11639,11642,11643,11647,11651,11654,11655,11657,11659,11661,11663,11666,11673,11679,11681,11683,11685,11687,11690,11691,11693,11697,11702,11715,11717,11719,11721,11723,11726,11871,11873,11875,11877,11879,11882,11885,11887,11889,11895,11897,11899,11901,11903,11906,11907,11909,11911,11913,11915,11919,11921,11923,11925,11927,11930,11931,11935,11937,11939,11943,11945,11947,11949,11951,11954,11955,11959,11961,11963,11967,11969,11971,11973,11975,11978,11983,11985,11987,12003,12005,12007,12009,12011,12014,12149,12159,12161,12163,12165,12167,12170,12171,12182,12183,12185,12187,12189,12191,12194,12197,12201,12206,12207,12209,12211,12213,12215,12218,12219,12223,12227,12230,12231,12233,12235,12237,12239,12242,12243,12245,12247,12249,12251,12254,12255,12257,12259,12261,12263,12266,12267,12269,12273,12278,12291,12293,12295,12297,12299,12302,12591,12592,12593,12594,12595,12596,12597,12598,12599,12600,12602,12615,12616,12617,12619,12620,12621,12622,12623,12624,12626,12639,12640,12641,12642,12643,12644,12645,12646,12647,12648,12650,12663,12665,12666,12667,12669,12670,12671,12674,12687,12688,12689,12691,12692,12693,12695,12696,12698,12723,12725,12727,12729,12731,12734,16191,16193,16195,16197,16199,16202,16215,16217,16219,16221,16223,16226,16239,16241,16243,16245,16247,16250,16263,16265,16267,16269,16271,16274,16287,16289,16291,16293,16295,16298,16323,16325,16327,16329,16331,16334,16479,16481,16483,16485,16487,16490,16503,16505,16507,16509,16511,16514,16515,16519,16521,16523,16526,16527,16529,16531,16533,16535,16538,16551,16553,16555,16557,16559,16562,16575,16577,16579,16581,16583,16586,16611,16613,16615,16617,16619,16622,16623,16627,16629,16631,16647,16651,16653,16655,16659,16663,16665,16667,16757,16766,16767,16769,16771,16773,16775,16778,16791,16793,16795,16797,16799,16802,16815,16817,16819,16821,16823,16826,16839,16841,16843,16845,16847,16850,16863,16865,16867,16869,16871,16874,16899,16901,16903,16905,16907,16910,17055,17057,17059,17061,17063,17066,17079,17081,17083,17085,17087,17090,17103,17105,17107,17109,17111,17114,17127,17129,17131,17133,17135,17138,17151,17153,17155,17157,17159,17162,17187,17189,17191,17193,17195,17198,17343,17345,17347,17349,17351,17354,17367,17369,17371,17373,17375,17378,17391,17393,17395,17397,17399,17402,17415,17417,17419,17421,17423,17426,17439,17441,17443,17445,17447,17450,17475,17477,17479,17481,17483,17486,17775,17776,17777,17779,17780,17781,17782,17783,17784,17786,17799,17800,17801,17803,17804,17805,17806,17807,17808,17810,17823,17825,17827,17829,17831,17834,17847,17849,17851,17853,17855,17858,17871,17872,17873,17875,17876,17877,17878,17879,17880,17882,17907,17909,17911,17913,17915,17918
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zf_lang__t6_zf_lang
% SZS output start Proof for zf_lang__t6_zf_lang
1. ! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => ((? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)) => v11_zf_lang(X0)) & ~(! [X1] : (m2_subset_1(X1,k5_numbers,k1_zf_lang) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k13_zf_lang(X1,X2) != X0)) & v11_zf_lang(X0)) & (? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v10_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k12_zf_lang(X1,X2) != X0)) & v10_zf_lang(X0)) & (? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v9_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k11_zf_lang(X1,X2) != X0)) & v9_zf_lang(X0)) & (? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v8_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k10_zf_lang(X1,X2) != X0)) & v8_zf_lang(X0)))) [input t6_zf_lang]
2. ~! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => ((? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)) => v11_zf_lang(X0)) & ~(! [X1] : (m2_subset_1(X1,k5_numbers,k1_zf_lang) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k13_zf_lang(X1,X2) != X0)) & v11_zf_lang(X0)) & (? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v10_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k12_zf_lang(X1,X2) != X0)) & v10_zf_lang(X0)) & (? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v9_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k11_zf_lang(X1,X2) != X0)) & v9_zf_lang(X0)) & (? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => v8_zf_lang(X0)) & ~(! [X1] : ((m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) => ! [X2] : ((m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => k10_zf_lang(X1,X2) != X0)) & v8_zf_lang(X0)))) [negated conjecture 1]
252. ! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => (v8_zf_lang(X0) <=> ? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)))) [input d20_zf_lang]
253. ! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => (v9_zf_lang(X0) <=> ? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)))) [input d21_zf_lang]
254. ! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => (v10_zf_lang(X0) <=> ? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)))) [input d22_zf_lang]
255. ! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => (v11_zf_lang(X0) <=> ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)))) [input d23_zf_lang]
258. ~! [X0] : ((m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => ((? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)) => v11_zf_lang(X0)) & ~(! [X3] : (m2_subset_1(X3,k5_numbers,k1_zf_lang) => ! [X4] : ((m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) => k13_zf_lang(X3,X4) != X0)) & v11_zf_lang(X0)) & (? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = X0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5)) => v10_zf_lang(X0)) & ~(! [X7] : ((m2_finseq_1(X7,k5_numbers) & v1_zf_lang(X7)) => ! [X8] : ((m2_finseq_1(X8,k5_numbers) & v1_zf_lang(X8)) => k12_zf_lang(X7,X8) != X0)) & v10_zf_lang(X0)) & (? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = X0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9)) => v9_zf_lang(X0)) & ~(! [X11] : ((m2_finseq_1(X11,k5_numbers) & v1_zf_lang(X11)) => ! [X12] : ((m2_finseq_1(X12,k5_numbers) & v1_zf_lang(X12)) => k11_zf_lang(X11,X12) != X0)) & v9_zf_lang(X0)) & (? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = X0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13)) => v8_zf_lang(X0)) & ~(! [X15] : ((m2_finseq_1(X15,k5_numbers) & v1_zf_lang(X15)) => ! [X16] : ((m2_finseq_1(X16,k5_numbers) & v1_zf_lang(X16)) => k10_zf_lang(X15,X16) != X0)) & v8_zf_lang(X0)))) [rectify 2]
299. ? [X0] : (((~v11_zf_lang(X0) & ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | (! [X3] : (! [X4] : (k13_zf_lang(X3,X4) != X0 | (~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4))) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang)) & v11_zf_lang(X0)) | (~v10_zf_lang(X0) & ? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = X0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5))) | (! [X7] : (! [X8] : (k12_zf_lang(X7,X8) != X0 | (~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8))) | (~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7))) & v10_zf_lang(X0)) | (~v9_zf_lang(X0) & ? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = X0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9))) | (! [X11] : (! [X12] : (k11_zf_lang(X11,X12) != X0 | (~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12))) | (~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11))) & v9_zf_lang(X0)) | (~v8_zf_lang(X0) & ? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = X0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13))) | (! [X15] : (! [X16] : (k10_zf_lang(X15,X16) != X0 | (~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16))) | (~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15))) & v8_zf_lang(X0))) & (m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0))) [ennf transformation 258]
300. ? [X0] : (((~v11_zf_lang(X0) & ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | (! [X3] : (! [X4] : (k13_zf_lang(X3,X4) != X0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4)) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang)) & v11_zf_lang(X0)) | (~v10_zf_lang(X0) & ? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = X0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5))) | (! [X7] : (! [X8] : (k12_zf_lang(X7,X8) != X0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8)) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7)) & v10_zf_lang(X0)) | (~v9_zf_lang(X0) & ? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = X0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9))) | (! [X11] : (! [X12] : (k11_zf_lang(X11,X12) != X0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12)) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11)) & v9_zf_lang(X0)) | (~v8_zf_lang(X0) & ? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = X0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13))) | (! [X15] : (! [X16] : (k10_zf_lang(X15,X16) != X0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16)) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15)) & v8_zf_lang(X0))) & m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) [flattening 299]
379. ! [X0] : ((v8_zf_lang(X0) <=> ? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | (~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0))) [ennf transformation 252]
380. ! [X0] : ((v8_zf_lang(X0) <=> ? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [flattening 379]
381. ! [X0] : ((v10_zf_lang(X0) <=> ? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | (~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0))) [ennf transformation 254]
382. ! [X0] : ((v10_zf_lang(X0) <=> ? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [flattening 381]
383. ! [X0] : ((v9_zf_lang(X0) <=> ? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | (~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0))) [ennf transformation 253]
384. ! [X0] : ((v9_zf_lang(X0) <=> ? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [flattening 383]
385. ! [X0] : ((v11_zf_lang(X0) <=> ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | (~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0))) [ennf transformation 255]
386. ! [X0] : ((v11_zf_lang(X0) <=> ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [flattening 385]
521. ? [X0] : (((~v11_zf_lang(X0) & ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | (! [X3] : (! [X4] : (k13_zf_lang(X3,X4) != X0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4)) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang)) & v11_zf_lang(X0)) | (~v10_zf_lang(X0) & ? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = X0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5))) | (! [X7] : (! [X8] : (k12_zf_lang(X7,X8) != X0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8)) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7)) & v10_zf_lang(X0)) | (~v9_zf_lang(X0) & ? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = X0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9))) | (! [X11] : (! [X12] : (k11_zf_lang(X11,X12) != X0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12)) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11)) & v9_zf_lang(X0)) | (~v8_zf_lang(X0) & ? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = X0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13))) | (! [X15] : (! [X16] : (k10_zf_lang(X15,X16) != X0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16)) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15)) & v8_zf_lang(X0))) & m2_finseq_1(X0,k5_numbers) & v1_zf_lang(X0)) => (((~v11_zf_lang(sK0) & ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = sK0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang))) | (! [X3] : (! [X4] : (k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4)) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang)) & v11_zf_lang(sK0)) | (~v10_zf_lang(sK0) & ? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = sK0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5))) | (! [X7] : (! [X8] : (k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8)) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7)) & v10_zf_lang(sK0)) | (~v9_zf_lang(sK0) & ? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = sK0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9))) | (! [X11] : (! [X12] : (k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12)) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11)) & v9_zf_lang(sK0)) | (~v8_zf_lang(sK0) & ? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = sK0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13))) | (! [X15] : (! [X16] : (k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16)) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15)) & v8_zf_lang(sK0))) & m2_finseq_1(sK0,k5_numbers) & v1_zf_lang(sK0)) [choice axiom]
522. ? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = sK0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)) => (? [X2] : (sK0 = k13_zf_lang(sK1,X2) & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(sK1,k5_numbers,k1_zf_lang)) [choice axiom]
523. ? [X2] : (sK0 = k13_zf_lang(sK1,X2) & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) => (sK0 = k13_zf_lang(sK1,sK2) & m2_finseq_1(sK2,k5_numbers) & v1_zf_lang(sK2)) [choice axiom]
524. ? [X5] : (? [X6] : (k12_zf_lang(X5,X6) = sK0 & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(X5,k5_numbers) & v1_zf_lang(X5)) => (? [X6] : (sK0 = k12_zf_lang(sK3,X6) & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) & m2_finseq_1(sK3,k5_numbers) & v1_zf_lang(sK3)) [choice axiom]
525. ? [X6] : (sK0 = k12_zf_lang(sK3,X6) & m2_finseq_1(X6,k5_numbers) & v1_zf_lang(X6)) => (sK0 = k12_zf_lang(sK3,sK4) & m2_finseq_1(sK4,k5_numbers) & v1_zf_lang(sK4)) [choice axiom]
526. ? [X9] : (? [X10] : (k11_zf_lang(X9,X10) = sK0 & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(X9,k5_numbers) & v1_zf_lang(X9)) => (? [X10] : (sK0 = k11_zf_lang(sK5,X10) & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) & m2_finseq_1(sK5,k5_numbers) & v1_zf_lang(sK5)) [choice axiom]
527. ? [X10] : (sK0 = k11_zf_lang(sK5,X10) & m2_finseq_1(X10,k5_numbers) & v1_zf_lang(X10)) => (sK0 = k11_zf_lang(sK5,sK6) & m2_finseq_1(sK6,k5_numbers) & v1_zf_lang(sK6)) [choice axiom]
528. ? [X13] : (? [X14] : (k10_zf_lang(X13,X14) = sK0 & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(X13,k5_numbers) & v1_zf_lang(X13)) => (? [X14] : (sK0 = k10_zf_lang(sK7,X14) & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) & m2_finseq_1(sK7,k5_numbers) & v1_zf_lang(sK7)) [choice axiom]
529. ? [X14] : (sK0 = k10_zf_lang(sK7,X14) & m2_finseq_1(X14,k5_numbers) & v1_zf_lang(X14)) => (sK0 = k10_zf_lang(sK7,sK8) & m2_finseq_1(sK8,k5_numbers) & v1_zf_lang(sK8)) [choice axiom]
530. ((~v11_zf_lang(sK0) & ((sK0 = k13_zf_lang(sK1,sK2) & m2_finseq_1(sK2,k5_numbers) & v1_zf_lang(sK2)) & m2_subset_1(sK1,k5_numbers,k1_zf_lang))) | (! [X3] : (! [X4] : (k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4)) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang)) & v11_zf_lang(sK0)) | (~v10_zf_lang(sK0) & ((sK0 = k12_zf_lang(sK3,sK4) & m2_finseq_1(sK4,k5_numbers) & v1_zf_lang(sK4)) & m2_finseq_1(sK3,k5_numbers) & v1_zf_lang(sK3))) | (! [X7] : (! [X8] : (k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8)) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7)) & v10_zf_lang(sK0)) | (~v9_zf_lang(sK0) & ((sK0 = k11_zf_lang(sK5,sK6) & m2_finseq_1(sK6,k5_numbers) & v1_zf_lang(sK6)) & m2_finseq_1(sK5,k5_numbers) & v1_zf_lang(sK5))) | (! [X11] : (! [X12] : (k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12)) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11)) & v9_zf_lang(sK0)) | (~v8_zf_lang(sK0) & ((sK0 = k10_zf_lang(sK7,sK8) & m2_finseq_1(sK8,k5_numbers) & v1_zf_lang(sK8)) & m2_finseq_1(sK7,k5_numbers) & v1_zf_lang(sK7))) | (! [X15] : (! [X16] : (k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16)) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15)) & v8_zf_lang(sK0))) & m2_finseq_1(sK0,k5_numbers) & v1_zf_lang(sK0) [skolemisation 300,529,528,527,526,525,524,523,522,521]
551. ! [X0] : (((v8_zf_lang(X0) | ! [X1] : (! [X2] : (k10_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X1] : (? [X2] : (k10_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) | ~v8_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [nnf transformation 380]
552. ! [X0] : (((v8_zf_lang(X0) | ! [X1] : (! [X2] : (k10_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X3] : (? [X4] : (k10_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) | ~v8_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [rectify 551]
553. ! [X0] : (? [X3] : (? [X4] : (k10_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) => (? [X4] : (k10_zf_lang(sK19(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(sK19(X0),k5_numbers) & v1_zf_lang(sK19(X0)))) [choice axiom]
554. ! [X0] : (? [X4] : (k10_zf_lang(sK19(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) => (k10_zf_lang(sK19(X0),sK20(X0)) = X0 & m2_finseq_1(sK20(X0),k5_numbers) & v1_zf_lang(sK20(X0)))) [choice axiom]
555. ! [X0] : (((v8_zf_lang(X0) | ! [X1] : (! [X2] : (k10_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (((k10_zf_lang(sK19(X0),sK20(X0)) = X0 & m2_finseq_1(sK20(X0),k5_numbers) & v1_zf_lang(sK20(X0))) & m2_finseq_1(sK19(X0),k5_numbers) & v1_zf_lang(sK19(X0))) | ~v8_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [skolemisation 552,554,553]
556. ! [X0] : (((v10_zf_lang(X0) | ! [X1] : (! [X2] : (k12_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X1] : (? [X2] : (k12_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) | ~v10_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [nnf transformation 382]
557. ! [X0] : (((v10_zf_lang(X0) | ! [X1] : (! [X2] : (k12_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X3] : (? [X4] : (k12_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) | ~v10_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [rectify 556]
558. ! [X0] : (? [X3] : (? [X4] : (k12_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) => (? [X4] : (k12_zf_lang(sK21(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(sK21(X0),k5_numbers) & v1_zf_lang(sK21(X0)))) [choice axiom]
559. ! [X0] : (? [X4] : (k12_zf_lang(sK21(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) => (k12_zf_lang(sK21(X0),sK22(X0)) = X0 & m2_finseq_1(sK22(X0),k5_numbers) & v1_zf_lang(sK22(X0)))) [choice axiom]
560. ! [X0] : (((v10_zf_lang(X0) | ! [X1] : (! [X2] : (k12_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (((k12_zf_lang(sK21(X0),sK22(X0)) = X0 & m2_finseq_1(sK22(X0),k5_numbers) & v1_zf_lang(sK22(X0))) & m2_finseq_1(sK21(X0),k5_numbers) & v1_zf_lang(sK21(X0))) | ~v10_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [skolemisation 557,559,558]
561. ! [X0] : (((v9_zf_lang(X0) | ! [X1] : (! [X2] : (k11_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X1] : (? [X2] : (k11_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_finseq_1(X1,k5_numbers) & v1_zf_lang(X1)) | ~v9_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [nnf transformation 384]
562. ! [X0] : (((v9_zf_lang(X0) | ! [X1] : (! [X2] : (k11_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (? [X3] : (? [X4] : (k11_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) | ~v9_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [rectify 561]
563. ! [X0] : (? [X3] : (? [X4] : (k11_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(X3,k5_numbers) & v1_zf_lang(X3)) => (? [X4] : (k11_zf_lang(sK23(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_finseq_1(sK23(X0),k5_numbers) & v1_zf_lang(sK23(X0)))) [choice axiom]
564. ! [X0] : (? [X4] : (k11_zf_lang(sK23(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) => (k11_zf_lang(sK23(X0),sK24(X0)) = X0 & m2_finseq_1(sK24(X0),k5_numbers) & v1_zf_lang(sK24(X0)))) [choice axiom]
565. ! [X0] : (((v9_zf_lang(X0) | ! [X1] : (! [X2] : (k11_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1))) & (((k11_zf_lang(sK23(X0),sK24(X0)) = X0 & m2_finseq_1(sK24(X0),k5_numbers) & v1_zf_lang(sK24(X0))) & m2_finseq_1(sK23(X0),k5_numbers) & v1_zf_lang(sK23(X0))) | ~v9_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [skolemisation 562,564,563]
566. ! [X0] : (((v11_zf_lang(X0) | ! [X1] : (! [X2] : (k13_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_subset_1(X1,k5_numbers,k1_zf_lang))) & (? [X1] : (? [X2] : (k13_zf_lang(X1,X2) = X0 & m2_finseq_1(X2,k5_numbers) & v1_zf_lang(X2)) & m2_subset_1(X1,k5_numbers,k1_zf_lang)) | ~v11_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [nnf transformation 386]
567. ! [X0] : (((v11_zf_lang(X0) | ! [X1] : (! [X2] : (k13_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_subset_1(X1,k5_numbers,k1_zf_lang))) & (? [X3] : (? [X4] : (k13_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_subset_1(X3,k5_numbers,k1_zf_lang)) | ~v11_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [rectify 566]
568. ! [X0] : (? [X3] : (? [X4] : (k13_zf_lang(X3,X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_subset_1(X3,k5_numbers,k1_zf_lang)) => (? [X4] : (k13_zf_lang(sK25(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) & m2_subset_1(sK25(X0),k5_numbers,k1_zf_lang))) [choice axiom]
569. ! [X0] : (? [X4] : (k13_zf_lang(sK25(X0),X4) = X0 & m2_finseq_1(X4,k5_numbers) & v1_zf_lang(X4)) => (k13_zf_lang(sK25(X0),sK26(X0)) = X0 & m2_finseq_1(sK26(X0),k5_numbers) & v1_zf_lang(sK26(X0)))) [choice axiom]
570. ! [X0] : (((v11_zf_lang(X0) | ! [X1] : (! [X2] : (k13_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2)) | ~m2_subset_1(X1,k5_numbers,k1_zf_lang))) & (((k13_zf_lang(sK25(X0),sK26(X0)) = X0 & m2_finseq_1(sK26(X0),k5_numbers) & v1_zf_lang(sK26(X0))) & m2_subset_1(sK25(X0),k5_numbers,k1_zf_lang)) | ~v11_zf_lang(X0))) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0)) [skolemisation 567,569,568]
637. v1_zf_lang(sK0) [cnf transformation 530] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
638. m2_finseq_1(sK0,k5_numbers) [cnf transformation 530] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
639. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
641. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
643. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
645. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
647. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
650. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
659. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
663. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
665. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
667. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
669. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
671. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
674. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
679. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
683. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
686. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
687. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
689. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
691. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
693. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
695. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
698. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
705. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
711. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
713. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
715. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
717. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
719. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
722. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
727. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
729. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
734. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
735. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
737. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
739. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
741. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
743. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
746. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
753. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
758. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
771. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
773. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
775. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
777. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
779. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
782. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
919. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
923. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
927. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
929. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
931. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
933. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
935. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
938. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
941. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
943. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
951. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
953. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
955. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
957. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
959. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
962. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
965. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
969. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
971. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
975. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
977. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
979. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
981. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
983. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
986. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
991. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
999. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1001. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1003. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1005. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1007. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1010. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1011. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1013. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1017. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1019. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1023. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1025. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1027. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1029. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1031. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1034. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1035. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1039. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1041. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1043. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1059. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1061. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1063. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1065. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1067. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1070. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1205. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1215. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1217. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1219. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1221. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1223. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1226. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1227. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1231. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1235. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1238. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1239. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1241. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1243. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1245. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1247. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1250. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1251. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1253. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1259. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1262. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1263. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1265. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1267. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1269. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1271. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1274. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1275. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1279. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1283. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1286. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1287. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1289. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1291. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1293. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1295. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1298. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1299. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1301. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1303. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1305. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1307. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1311. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1313. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1315. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1317. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1319. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1322. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1323. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1325. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1327. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1329. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1331. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1347. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1349. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1351. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1353. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1355. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1358. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1495. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1503. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1505. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1507. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1509. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1511. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1514. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1515. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1517. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1527. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1529. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1531. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1533. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1535. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1538. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1539. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1541. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1545. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1547. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1551. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1553. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1555. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1557. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1559. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1562. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1567. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1571. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1574. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1575. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1577. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1579. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1581. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1583. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1586. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1589. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1593. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1595. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1599. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1601. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1603. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1605. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1607. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1610. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1617. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1635. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1637. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1639. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1641. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1643. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1646. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1791. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1793. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1795. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1797. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1799. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1802. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1803. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1805. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1807. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1811. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1815. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1817. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1819. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1821. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1823. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1826. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1827. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1829. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1833. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1838. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1839. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1841. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1843. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1845. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1847. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1850. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1855. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1857. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1859. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1863. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1865. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1867. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1869. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1871. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1874. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1881. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1887. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1889. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1891. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1893. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1895. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1898. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1901. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1903. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1905. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1907. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:40,wCS:4000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1923. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1925. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1927. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1929. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1931. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
1934. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2071. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2073. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2223. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2224. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2225. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2226. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2227. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2228. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2229. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2230. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2231. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2232. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2234. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2247. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2248. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2249. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2251. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2252. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2253. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2255. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2256. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2258. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2271. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2272. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2273. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2274. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2275. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2276. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2277. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2278. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2279. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2280. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2282. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2295. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2296. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2297. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2299. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2300. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2301. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2302. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2303. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2304. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2306. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2319. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2320. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2321. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2322. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2323. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2324. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2325. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2327. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2328. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:50,wCS:5000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2330. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2355. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2357. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2359. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2361. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2363. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
2366. m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4095. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4097. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4099. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4101. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4103. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4106. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4118. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4119. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4121. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4123. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4125. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4127. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4130. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4131. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4133. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4137. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4139. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4143. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4145. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4147. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4149. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4151. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4154. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4161. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4166. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4167. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4169. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4171. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4173. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4175. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4178. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4183. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4185. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4190. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4191. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4193. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4195. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4197. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4199. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4202. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4203. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4209. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4214. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4227. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4229. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4231. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4233. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4235. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4238. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4377. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4383. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4385. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4387. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4389. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4391. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4394. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4397. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4401. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4406. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4407. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4409. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4410. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4411. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4413. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4415. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4418. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4421. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4427. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4430. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4431. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4433. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4435. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4437. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4439. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4442. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4449. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4451. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4454. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4455. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4457. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4459. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4461. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4463. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4466. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4467. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4469. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4478. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4479. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4481. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4483. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4485. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4487. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4490. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4491. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4493. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4497. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4499. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4502. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4515. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4517. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4519. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4521. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4523. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4526. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4661. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4663. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4671. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4673. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4675. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4677. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4679. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4682. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4683. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4687. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4691. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4694. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4695. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4697. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4699. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4701. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4703. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4706. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4707. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4709. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4715. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4718. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4719. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4721. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4723. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4725. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4727. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4730. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4735. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4739. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4742. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4743. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4745. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4747. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4749. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4751. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4754. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4755. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4757. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4759. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4763. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4766. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4767. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4769. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4771. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4773. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4775. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4778. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4779. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4781. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4783. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4785. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4787. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4790. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4803. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4805. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4807. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4809. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4811. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4814. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4951. v1_zf_lang(sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4959. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4961. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4963. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4965. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4967. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4970. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4977. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4979. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4982. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4983. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4985. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4987. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4989. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4991. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4994. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
4995. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5006. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5007. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5009. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5011. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5013. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5015. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5018. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5023. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5030. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5031. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5033. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5035. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5037. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5039. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5042. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5045. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5047. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5054. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5055. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5057. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5059. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5061. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5063. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5066. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5069. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5078. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5091. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5093. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5095. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5097. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5099. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5102. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5239. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5241. v1_zf_lang(sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5247. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5249. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5251. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5253. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5255. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5258. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5259. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5263. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5267. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5270. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5271. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5273. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5275. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5277. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5279. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5282. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5287. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5289. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5291. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5295. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5297. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5299. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5301. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5303. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5306. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5311. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5313. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5315. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5318. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5319. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5321. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5323. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5325. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5327. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5330. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5331. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5335. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5337. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5343. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5345. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5347. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5349. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5351. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5354. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5355. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5361. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5366. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5379. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5381. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5383. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5385. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5387. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5390. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5527. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5529. v1_zf_lang(sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5679. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5680. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5681. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5682. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5683. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5685. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5687. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5690. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5703. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5704. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5705. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5707. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5708. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5709. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5711. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5712. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5714. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5727. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5729. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5730. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5731. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5733. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5734. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5735. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5738. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5751. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5753. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5755. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5757. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5758. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5759. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5762. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5775. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5777. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5778. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5779. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5781. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5783. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5786. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5811. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5813. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5815. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5817. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5819. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
5822. v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7551. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7553. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7555. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7557. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7559. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7562. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7565. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7567. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7574. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7575. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7577. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7579. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7581. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7583. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7586. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7589. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7591. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7598. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7599. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7601. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7603. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7605. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7607. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7610. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7622. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7623. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7625. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7627. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7629. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7631. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7634. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7643. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7646. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7647. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7649. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7651. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7653. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7655. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7658. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7659. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7670. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7683. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7685. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7687. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7689. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7691. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7694. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7831. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7833. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7835. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7839. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7841. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7843. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7845. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7847. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7850. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7862. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7863. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7865. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7867. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7869. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7871. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7874. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7875. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7877. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7879. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7881. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7883. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7887. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7889. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7891. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7893. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7895. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7898. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7911. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7913. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7915. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7917. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7919. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7922. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7925. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7927. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7929. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7931. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7935. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7937. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7939. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7941. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7943. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7946. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7949. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7951. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7953. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7971. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7973. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7975. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7977. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7979. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
7982. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8117. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8119. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8121. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8123. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8127. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8129. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8131. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8133. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8135. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8138. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8141. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8143. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8147. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8150. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8151. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8153. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8155. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8157. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8159. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8162. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8165. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8171. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8174. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8175. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8177. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8179. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8181. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8183. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8186. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8199. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8201. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8203. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8205. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8207. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8210. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8211. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8213. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8215. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8219. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8222. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8223. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8225. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8227. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8229. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8231. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8234. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8237. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8239. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8241. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8243. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8246. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8259. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8261. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8263. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8265. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8267. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8270. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8407. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8409. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8415. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8417. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8419. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8421. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8423. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8426. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8431. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8433. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8438. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8439. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8441. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8443. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8445. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8447. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8450. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8453. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8455. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8457. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8463. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8465. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8467. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8469. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8471. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8474. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8486. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8487. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8489. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8491. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8493. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8495. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8498. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8499. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8505. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8507. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8510. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8511. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8513. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8515. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8517. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8519. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8522. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8523. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8525. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8527. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8529. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8534. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8547. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8549. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8551. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8553. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8555. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8558. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8695. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8697. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8703. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8705. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8707. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8709. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8711. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8714. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8715. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8721. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8726. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8727. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8729. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8731. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8733. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8735. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8738. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8741. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8747. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8750. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8751. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8753. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8755. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8757. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8759. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8762. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8774. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8775. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8777. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8779. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8781. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8783. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8786. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8787. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8795. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8798. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8799. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8801. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8803. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8805. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8807. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8810. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8811. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8815. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8819. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8822. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8835. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8837. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8839. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8841. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8843. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8846. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8983. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
8985. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9135. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9137. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9138. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9139. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9141. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9143. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9146. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9159. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9161. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9163. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9165. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9167. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9170. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9183. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9184. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9185. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9186. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9187. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9188. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9189. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9190. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9191. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9192. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9194. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9207. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9209. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9211. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9213. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9214. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9215. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9218. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9231. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9232. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9233. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9234. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9235. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9236. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9237. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9238. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9239. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9240. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9242. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9267. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9269. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9271. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9273. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9275. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
9278. m2_finseq_1(sK2,k5_numbers) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11007. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11009. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11011. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11013. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11015. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11018. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11019. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11023. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11027. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11031. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11033. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11035. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11037. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11039. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11042. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11043. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11045. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11047. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11051. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11054. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11055. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11057. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11059. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11061. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11063. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11066. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11067. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11071. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11075. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11079. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11081. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11083. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11085. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11087. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11090. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11099. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11102. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11103. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11105. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11107. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11109. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11111. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11114. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11115. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11126. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11139. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11141. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11143. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11145. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11147. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11150. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11295. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11297. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11299. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11301. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11303. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11306. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11307. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11309. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11313. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11315. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11319. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11321. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11323. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11325. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11327. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11330. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11333. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11337. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11343. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11345. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11347. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11349. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11351. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11354. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11355. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11363. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11367. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11369. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11371. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11373. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11375. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11378. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11379. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11381. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11383. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11385. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11387. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11391. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11393. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11395. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11397. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11399. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11402. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11403. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11409. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11427. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11429. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11431. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11433. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11435. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11438. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11573. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11583. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11585. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11587. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11589. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11591. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11594. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11597. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11606. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11607. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11609. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11611. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11613. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11615. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11618. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11630. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11631. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11633. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11635. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11637. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11639. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11642. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11643. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11647. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11651. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11654. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11655. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11657. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11659. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11661. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11663. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11666. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11673. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11679. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11681. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11683. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11685. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11687. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11690. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11691. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11693. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11697. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11702. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11715. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11717. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11719. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11721. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11723. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11726. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11871. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11873. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11875. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11877. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11879. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11882. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11885. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11887. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11889. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11895. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11897. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11899. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11901. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11903. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11906. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11907. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11909. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11911. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11913. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11915. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11919. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11921. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11923. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11925. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11927. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11930. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11931. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11935. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11937. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11939. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11943. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11945. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11947. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11949. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11951. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11954. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11955. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11959. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11961. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11963. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11967. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11969. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11971. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11973. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11975. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11978. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11983. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11985. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
11987. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12003. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12005. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12007. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12009. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12011. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12014. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12149. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | m2_finseq_1(sK4,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12159. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12161. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12163. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12165. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12167. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12170. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12171. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12182. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12183. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12185. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12187. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12189. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12191. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12194. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12197. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12201. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12206. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12207. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12209. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12211. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12213. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12215. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12218. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12219. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12223. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12227. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12230. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12231. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12233. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12235. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:23,wCS:2300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12237. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12239. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12242. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12243. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12245. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12247. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12249. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12251. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12254. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12255. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12257. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12259. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:25,wCS:2500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12261. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:26,wCS:2600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12263. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12266. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12267. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12269. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12273. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12278. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12291. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12293. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12295. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12297. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12299. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12302. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12591. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12592. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12593. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12594. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12595. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12596. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12597. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12598. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12599. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12600. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12602. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12615. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12616. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12617. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12619. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12620. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12621. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12622. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12623. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12624. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12626. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12639. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12640. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12641. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12642. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12643. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12644. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12645. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12646. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12647. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12648. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12650. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12663. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12665. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12666. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12667. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12669. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12670. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12671. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12674. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12687. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12688. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12689. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12691. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12692. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12693. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:36,wCS:3600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12695. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:38,wCS:3800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12696. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12698. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12723. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12725. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12727. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12729. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12731. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:48,wCS:4800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
12734. sK0 = k13_zf_lang(sK1,sK2) | v11_zf_lang(sK0) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:58,wCS:5800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16191. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16193. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16195. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16197. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16199. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16202. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16215. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16217. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16219. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16221. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16223. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16226. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16239. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16241. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16243. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16245. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16247. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16250. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16263. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16265. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16267. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16269. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16271. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16274. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16287. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16289. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16291. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16293. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16295. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16298. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16323. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16325. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16327. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16329. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16331. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16334. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16479. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16481. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16483. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16485. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16487. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16490. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16503. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16505. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16507. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16509. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16511. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16514. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16515. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16519. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16521. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16523. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16526. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16527. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16529. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16531. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16533. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16535. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16538. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16551. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16553. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16555. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16557. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16559. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16562. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16575. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16577. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16579. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16581. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16583. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16586. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16611. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16613. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16615. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16617. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16619. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16622. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16623. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16627. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16629. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16631. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16647. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16651. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16653. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16655. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:46,wCS:4600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16659. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16663. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16665. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16667. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:59,wCS:5900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16757. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16766. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK3,k5_numbers) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16767. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16769. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16771. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16773. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16775. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16778. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16791. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16793. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16795. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16797. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16799. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16802. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16815. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16817. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16819. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:28,wCS:2800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16821. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16823. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16826. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16839. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16841. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16843. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16845. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16847. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16850. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16863. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16865. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16867. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16869. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16871. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16874. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16899. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16901. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16903. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16905. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16907. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
16910. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | v1_zf_lang(sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17055. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17057. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17059. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17061. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17063. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17066. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17079. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17081. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17083. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17085. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17087. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17090. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17103. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17105. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17107. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:29,wCS:2900,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17109. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17111. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17114. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17127. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17129. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17131. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:30,wCS:3000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17133. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17135. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17138. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17151. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17153. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17155. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17157. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17159. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17162. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17187. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17189. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17191. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17193. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17195. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17198. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | m2_finseq_1(sK4,k5_numbers) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17343. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17345. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17347. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17349. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17351. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17354. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17367. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17369. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17371. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17373. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17375. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17378. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17391. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17393. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17395. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17397. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17399. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17402. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17415. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17417. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17419. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17421. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17423. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17426. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17439. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17441. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17443. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17445. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17447. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:37,wCS:3700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17450. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17475. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17477. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17479. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17481. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17483. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17486. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | sK0 = k12_zf_lang(sK3,sK4) | v10_zf_lang(sK0) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17775. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17776. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17777. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17779. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17780. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17781. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17782. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17783. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17784. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17786. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17799. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17800. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17801. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17803. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17804. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17805. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17806. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17807. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17808. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:58,wCS:5800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17810. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17823. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17825. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17827. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:41,wCS:4100,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17829. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17831. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17834. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | v1_zf_lang(sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17847. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17849. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17851. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:42,wCS:4200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17853. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17855. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17858. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | m2_finseq_1(sK6,k5_numbers) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17871. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17872. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK7) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17873. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17875. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:44,wCS:4400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17876. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | v1_zf_lang(sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17877. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17878. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | m2_finseq_1(sK8,k5_numbers) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:58,wCS:5800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17879. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17880. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | sK0 = k10_zf_lang(sK7,sK8) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17882. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | sK0 = k11_zf_lang(sK5,sK6) | v9_zf_lang(sK0) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17907. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK7) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17909. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK7,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17911. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | v1_zf_lang(sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17913. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | m2_finseq_1(sK8,k5_numbers) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17915. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | sK0 = k10_zf_lang(sK7,sK8) | v8_zf_lang(sK0) [cnf transformation 530] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
17918. ~v11_zf_lang(sK0) | k13_zf_lang(X3,X4) != sK0 | ~m2_finseq_1(X4,k5_numbers) | ~v1_zf_lang(X4) | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v10_zf_lang(sK0) | k12_zf_lang(X7,X8) != sK0 | ~m2_finseq_1(X8,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X7) | ~v9_zf_lang(sK0) | k11_zf_lang(X11,X12) != sK0 | ~m2_finseq_1(X12,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X11) | ~v8_zf_lang(sK0) | k10_zf_lang(X15,X16) != sK0 | ~m2_finseq_1(X16,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X15) [cnf transformation 530] {a:0,w:67,wCS:6700,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
18092. ~m2_finseq_1(X0,k5_numbers) | ~v8_zf_lang(X0) | v1_zf_lang(sK19(X0)) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18093. m2_finseq_1(sK19(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18094. ~m2_finseq_1(X0,k5_numbers) | ~v8_zf_lang(X0) | v1_zf_lang(sK20(X0)) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18095. m2_finseq_1(sK20(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18096. k10_zf_lang(sK19(X0),sK20(X0)) = X0 | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18097. v8_zf_lang(X0) | k10_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 555] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
18098. ~m2_finseq_1(X0,k5_numbers) | ~v10_zf_lang(X0) | v1_zf_lang(sK21(X0)) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18099. m2_finseq_1(sK21(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18100. ~m2_finseq_1(X0,k5_numbers) | ~v10_zf_lang(X0) | v1_zf_lang(sK22(X0)) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18101. m2_finseq_1(sK22(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18102. k12_zf_lang(sK21(X0),sK22(X0)) = X0 | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18103. v10_zf_lang(X0) | k12_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 560] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
18104. ~m2_finseq_1(X0,k5_numbers) | ~v9_zf_lang(X0) | v1_zf_lang(sK23(X0)) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18105. m2_finseq_1(sK23(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18106. ~m2_finseq_1(X0,k5_numbers) | ~v9_zf_lang(X0) | v1_zf_lang(sK24(X0)) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18107. m2_finseq_1(sK24(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18108. k11_zf_lang(sK23(X0),sK24(X0)) = X0 | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18109. v9_zf_lang(X0) | k11_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 565] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
18110. m2_subset_1(sK25(X0),k5_numbers,k1_zf_lang) | ~v11_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 570] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18111. ~m2_finseq_1(X0,k5_numbers) | ~v11_zf_lang(X0) | v1_zf_lang(sK26(X0)) | ~v1_zf_lang(X0) [cnf transformation 570] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18112. m2_finseq_1(sK26(X0),k5_numbers) | ~v11_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 570] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18113. k13_zf_lang(sK25(X0),sK26(X0)) = X0 | ~v11_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 570] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18114. v11_zf_lang(X0) | k13_zf_lang(X1,X2) != X0 | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_subset_1(X1,k5_numbers,k1_zf_lang) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) [cnf transformation 570] {a:0,w:21,wCS:8400,sine:1,thAx:0,allAx:3,thDist:-3}
18305. ~m2_finseq_1(k10_zf_lang(X1,X2),k5_numbers) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | v8_zf_lang(k10_zf_lang(X1,X2)) | ~v1_zf_lang(k10_zf_lang(X1,X2)) [equality resolution 18097] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18306. ~m2_finseq_1(k12_zf_lang(X1,X2),k5_numbers) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | v10_zf_lang(k12_zf_lang(X1,X2)) | ~v1_zf_lang(k12_zf_lang(X1,X2)) [equality resolution 18103] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18307. ~m2_finseq_1(k11_zf_lang(X1,X2),k5_numbers) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_finseq_1(X1,k5_numbers) | ~v1_zf_lang(X1) | v9_zf_lang(k11_zf_lang(X1,X2)) | ~v1_zf_lang(k11_zf_lang(X1,X2)) [equality resolution 18109] {a:0,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18308. ~m2_finseq_1(k13_zf_lang(X1,X2),k5_numbers) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) | ~m2_subset_1(X1,k5_numbers,k1_zf_lang) | v11_zf_lang(k13_zf_lang(X1,X2)) | ~v1_zf_lang(k13_zf_lang(X1,X2)) [equality resolution 18114] {a:0,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
18311. 1 <=> ! [X16,X15] : (k10_zf_lang(X15,X16) != sK0 | ~v1_zf_lang(X15) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X16,k5_numbers)) [avatar definition]
18312. k10_zf_lang(X15,X16) != sK0 | ~v1_zf_lang(X15) | ~m2_finseq_1(X15,k5_numbers) | ~v1_zf_lang(X16) | ~m2_finseq_1(X16,k5_numbers) <- (1) [avatar component clause 18311] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18314. 2 <=> v8_zf_lang(sK0) [avatar definition]
18315. v8_zf_lang(sK0) <- (2) [avatar component clause 18314] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18316. ~v8_zf_lang(sK0) <- (~2) [avatar component clause 18314] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18318. 3 <=> ! [X11,X12] : (k11_zf_lang(X11,X12) != sK0 | ~v1_zf_lang(X11) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X12,k5_numbers)) [avatar definition]
18319. k11_zf_lang(X11,X12) != sK0 | ~v1_zf_lang(X11) | ~m2_finseq_1(X11,k5_numbers) | ~v1_zf_lang(X12) | ~m2_finseq_1(X12,k5_numbers) <- (3) [avatar component clause 18318] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18321. 4 <=> v9_zf_lang(sK0) [avatar definition]
18322. v9_zf_lang(sK0) <- (4) [avatar component clause 18321] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18325. 5 <=> ! [X7,X8] : (k12_zf_lang(X7,X8) != sK0 | ~v1_zf_lang(X7) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X8,k5_numbers)) [avatar definition]
18326. k12_zf_lang(X7,X8) != sK0 | ~v1_zf_lang(X7) | ~m2_finseq_1(X7,k5_numbers) | ~v1_zf_lang(X8) | ~m2_finseq_1(X8,k5_numbers) <- (5) [avatar component clause 18325] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18328. 6 <=> v10_zf_lang(sK0) [avatar definition]
18330. ~v10_zf_lang(sK0) <- (~6) [avatar component clause 18328] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18332. 7 <=> ! [X3,X4] : (k13_zf_lang(X3,X4) != sK0 | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v1_zf_lang(X4) | ~m2_finseq_1(X4,k5_numbers)) [avatar definition]
18333. k13_zf_lang(X3,X4) != sK0 | ~m2_subset_1(X3,k5_numbers,k1_zf_lang) | ~v1_zf_lang(X4) | ~m2_finseq_1(X4,k5_numbers) <- (7) [avatar component clause 18332] {a:0,w:14,wCS:1400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18335. 8 <=> v11_zf_lang(sK0) [avatar definition]
18337. ~v11_zf_lang(sK0) <- (~8) [avatar component clause 18335] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18338. 1 | ~2 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17918,18335,18332,18328,18325,18321,18318,18314,18311]
18340. 9 <=> sK0 = k10_zf_lang(sK7,sK8) [avatar definition]
18342. sK0 = k10_zf_lang(sK7,sK8) <- (9) [avatar component clause 18340] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18344. 2 | 9 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17915,18335,18332,18328,18325,18321,18318,18340,18314]
18346. 10 <=> m2_finseq_1(sK8,k5_numbers) [avatar definition]
18348. m2_finseq_1(sK8,k5_numbers) <- (10) [avatar component clause 18346] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18350. 2 | 10 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17913,18335,18332,18328,18325,18321,18318,18346,18314]
18352. 11 <=> v1_zf_lang(sK8) [avatar definition]
18354. v1_zf_lang(sK8) <- (11) [avatar component clause 18352] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18356. 2 | 11 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17911,18335,18332,18328,18325,18321,18318,18352,18314]
18358. 12 <=> m2_finseq_1(sK7,k5_numbers) [avatar definition]
18360. m2_finseq_1(sK7,k5_numbers) <- (12) [avatar component clause 18358] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18362. 2 | 12 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17909,18335,18332,18328,18325,18321,18318,18358,18314]
18364. 13 <=> v1_zf_lang(sK7) [avatar definition]
18366. v1_zf_lang(sK7) <- (13) [avatar component clause 18364] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18368. 2 | 13 | 3 | ~4 | 5 | ~6 | 7 | ~8 [avatar split clause 17907,18335,18332,18328,18325,18321,18318,18364,18314]
18370. 14 <=> sK0 = k11_zf_lang(sK5,sK6) [avatar definition]
18372. sK0 = k11_zf_lang(sK5,sK6) <- (14) [avatar component clause 18370] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18384. 1 | ~2 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17882,18335,18332,18328,18325,18370,18321,18314,18311]
18385. 1 | 9 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17880,18335,18332,18328,18325,18370,18321,18340,18311]
18386. 2 | 9 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17879,18335,18332,18328,18325,18370,18321,18340,18314]
18387. 1 | 10 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17878,18335,18332,18328,18325,18370,18321,18346,18311]
18388. 2 | 10 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17877,18335,18332,18328,18325,18370,18321,18346,18314]
18389. 1 | 11 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17876,18335,18332,18328,18325,18370,18321,18352,18311]
18390. 2 | 11 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17875,18335,18332,18328,18325,18370,18321,18352,18314]
18392. 2 | 12 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17873,18335,18332,18328,18325,18370,18321,18358,18314]
18393. 1 | 13 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17872,18335,18332,18328,18325,18370,18321,18364,18311]
18394. 2 | 13 | 4 | 14 | 5 | ~6 | 7 | ~8 [avatar split clause 17871,18335,18332,18328,18325,18370,18321,18364,18314]
18396. 15 <=> m2_finseq_1(sK6,k5_numbers) [avatar definition]
18410. 1 | ~2 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17858,18335,18332,18328,18325,18396,18321,18314,18311]
18412. 2 | 9 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17855,18335,18332,18328,18325,18396,18321,18340,18314]
18414. 2 | 10 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17853,18335,18332,18328,18325,18396,18321,18346,18314]
18416. 2 | 11 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17851,18335,18332,18328,18325,18396,18321,18352,18314]
18418. 2 | 12 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17849,18335,18332,18328,18325,18396,18321,18358,18314]
18420. 2 | 13 | 4 | 15 | 5 | ~6 | 7 | ~8 [avatar split clause 17847,18335,18332,18328,18325,18396,18321,18364,18314]
18422. 16 <=> v1_zf_lang(sK6) [avatar definition]
18424. v1_zf_lang(sK6) <- (16) [avatar component clause 18422] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18436. 1 | ~2 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17834,18335,18332,18328,18325,18422,18321,18314,18311]
18438. 2 | 9 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17831,18335,18332,18328,18325,18422,18321,18340,18314]
18440. 2 | 10 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17829,18335,18332,18328,18325,18422,18321,18346,18314]
18442. 2 | 11 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17827,18335,18332,18328,18325,18422,18321,18352,18314]
18444. 2 | 12 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17825,18335,18332,18328,18325,18422,18321,18358,18314]
18446. 2 | 13 | 4 | 16 | 5 | ~6 | 7 | ~8 [avatar split clause 17823,18335,18332,18328,18325,18422,18321,18364,18314]
18448. 17 <=> m2_finseq_1(sK5,k5_numbers) [avatar definition]
18462. 1 | ~2 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17810,18335,18332,18328,18325,18448,18321,18314,18311]
18463. 1 | 9 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17808,18335,18332,18328,18325,18448,18321,18340,18311]
18464. 2 | 9 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17807,18335,18332,18328,18325,18448,18321,18340,18314]
18465. 1 | 10 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17806,18335,18332,18328,18325,18448,18321,18346,18311]
18466. 2 | 10 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17805,18335,18332,18328,18325,18448,18321,18346,18314]
18467. 1 | 11 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17804,18335,18332,18328,18325,18448,18321,18352,18311]
18468. 2 | 11 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17803,18335,18332,18328,18325,18448,18321,18352,18314]
18470. 2 | 12 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17801,18335,18332,18328,18325,18448,18321,18358,18314]
18471. 1 | 13 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17800,18335,18332,18328,18325,18448,18321,18364,18311]
18472. 2 | 13 | 4 | 17 | 5 | ~6 | 7 | ~8 [avatar split clause 17799,18335,18332,18328,18325,18448,18321,18364,18314]
18474. 18 <=> v1_zf_lang(sK5) [avatar definition]
18476. v1_zf_lang(sK5) <- (18) [avatar component clause 18474] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18488. 1 | ~2 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17786,18335,18332,18328,18325,18474,18321,18314,18311]
18489. 1 | 9 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17784,18335,18332,18328,18325,18474,18321,18340,18311]
18490. 2 | 9 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17783,18335,18332,18328,18325,18474,18321,18340,18314]
18491. 1 | 10 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17782,18335,18332,18328,18325,18474,18321,18346,18311]
18492. 2 | 10 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17781,18335,18332,18328,18325,18474,18321,18346,18314]
18493. 1 | 11 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17780,18335,18332,18328,18325,18474,18321,18352,18311]
18494. 2 | 11 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17779,18335,18332,18328,18325,18474,18321,18352,18314]
18496. 2 | 12 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17777,18335,18332,18328,18325,18474,18321,18358,18314]
18497. 1 | 13 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17776,18335,18332,18328,18325,18474,18321,18364,18311]
18498. 2 | 13 | 4 | 18 | 5 | ~6 | 7 | ~8 [avatar split clause 17775,18335,18332,18328,18325,18474,18321,18364,18314]
18500. 19 <=> sK0 = k12_zf_lang(sK3,sK4) [avatar definition]
18502. sK0 = k12_zf_lang(sK3,sK4) <- (19) [avatar component clause 18500] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
18624. 1 | ~2 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17486,18335,18332,18500,18328,18321,18318,18314,18311]
18626. 2 | 9 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17483,18335,18332,18500,18328,18321,18318,18340,18314]
18628. 2 | 10 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17481,18335,18332,18500,18328,18321,18318,18346,18314]
18630. 2 | 11 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17479,18335,18332,18500,18328,18321,18318,18352,18314]
18632. 2 | 12 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17477,18335,18332,18500,18328,18321,18318,18358,18314]
18634. 2 | 13 | 3 | ~4 | 6 | 19 | 7 | ~8 [avatar split clause 17475,18335,18332,18500,18328,18321,18318,18364,18314]
18646. 1 | ~2 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17450,18335,18332,18500,18328,18370,18321,18314,18311]
18648. 2 | 9 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17447,18335,18332,18500,18328,18370,18321,18340,18314]
18650. 2 | 10 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17445,18335,18332,18500,18328,18370,18321,18346,18314]
18652. 2 | 11 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17443,18335,18332,18500,18328,18370,18321,18352,18314]
18654. 2 | 12 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17441,18335,18332,18500,18328,18370,18321,18358,18314]
18656. 2 | 13 | 4 | 14 | 6 | 19 | 7 | ~8 [avatar split clause 17439,18335,18332,18500,18328,18370,18321,18364,18314]
18668. 1 | ~2 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17426,18335,18332,18500,18328,18396,18321,18314,18311]
18670. 2 | 9 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17423,18335,18332,18500,18328,18396,18321,18340,18314]
18672. 2 | 10 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17421,18335,18332,18500,18328,18396,18321,18346,18314]
18674. 2 | 11 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17419,18335,18332,18500,18328,18396,18321,18352,18314]
18676. 2 | 12 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17417,18335,18332,18500,18328,18396,18321,18358,18314]
18678. 2 | 13 | 4 | 15 | 6 | 19 | 7 | ~8 [avatar split clause 17415,18335,18332,18500,18328,18396,18321,18364,18314]
18690. 1 | ~2 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17402,18335,18332,18500,18328,18422,18321,18314,18311]
18692. 2 | 9 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17399,18335,18332,18500,18328,18422,18321,18340,18314]
18694. 2 | 10 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17397,18335,18332,18500,18328,18422,18321,18346,18314]
18696. 2 | 11 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17395,18335,18332,18500,18328,18422,18321,18352,18314]
18698. 2 | 12 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17393,18335,18332,18500,18328,18422,18321,18358,18314]
18700. 2 | 13 | 4 | 16 | 6 | 19 | 7 | ~8 [avatar split clause 17391,18335,18332,18500,18328,18422,18321,18364,18314]
18712. 1 | ~2 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17378,18335,18332,18500,18328,18448,18321,18314,18311]
18714. 2 | 9 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17375,18335,18332,18500,18328,18448,18321,18340,18314]
18716. 2 | 10 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17373,18335,18332,18500,18328,18448,18321,18346,18314]
18718. 2 | 11 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17371,18335,18332,18500,18328,18448,18321,18352,18314]
18720. 2 | 12 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17369,18335,18332,18500,18328,18448,18321,18358,18314]
18722. 2 | 13 | 4 | 17 | 6 | 19 | 7 | ~8 [avatar split clause 17367,18335,18332,18500,18328,18448,18321,18364,18314]
18734. 1 | ~2 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17354,18335,18332,18500,18328,18474,18321,18314,18311]
18736. 2 | 9 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17351,18335,18332,18500,18328,18474,18321,18340,18314]
18738. 2 | 10 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17349,18335,18332,18500,18328,18474,18321,18346,18314]
18740. 2 | 11 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17347,18335,18332,18500,18328,18474,18321,18352,18314]
18742. 2 | 12 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17345,18335,18332,18500,18328,18474,18321,18358,18314]
18744. 2 | 13 | 4 | 18 | 6 | 19 | 7 | ~8 [avatar split clause 17343,18335,18332,18500,18328,18474,18321,18364,18314]
18746. 20 <=> m2_finseq_1(sK4,k5_numbers) [avatar definition]
18870. 1 | ~2 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17198,18335,18332,18746,18328,18321,18318,18314,18311]
18872. 2 | 9 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17195,18335,18332,18746,18328,18321,18318,18340,18314]
18874. 2 | 10 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17193,18335,18332,18746,18328,18321,18318,18346,18314]
18876. 2 | 11 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17191,18335,18332,18746,18328,18321,18318,18352,18314]
18878. 2 | 12 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17189,18335,18332,18746,18328,18321,18318,18358,18314]
18880. 2 | 13 | 3 | ~4 | 6 | 20 | 7 | ~8 [avatar split clause 17187,18335,18332,18746,18328,18321,18318,18364,18314]
18892. 1 | ~2 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17162,18335,18332,18746,18328,18370,18321,18314,18311]
18894. 2 | 9 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17159,18335,18332,18746,18328,18370,18321,18340,18314]
18896. 2 | 10 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17157,18335,18332,18746,18328,18370,18321,18346,18314]
18898. 2 | 11 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17155,18335,18332,18746,18328,18370,18321,18352,18314]
18900. 2 | 12 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17153,18335,18332,18746,18328,18370,18321,18358,18314]
18902. 2 | 13 | 4 | 14 | 6 | 20 | 7 | ~8 [avatar split clause 17151,18335,18332,18746,18328,18370,18321,18364,18314]
18914. 1 | ~2 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17138,18335,18332,18746,18328,18396,18321,18314,18311]
18916. 2 | 9 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17135,18335,18332,18746,18328,18396,18321,18340,18314]
18918. 2 | 10 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17133,18335,18332,18746,18328,18396,18321,18346,18314]
18920. 2 | 11 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17131,18335,18332,18746,18328,18396,18321,18352,18314]
18922. 2 | 12 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17129,18335,18332,18746,18328,18396,18321,18358,18314]
18924. 2 | 13 | 4 | 15 | 6 | 20 | 7 | ~8 [avatar split clause 17127,18335,18332,18746,18328,18396,18321,18364,18314]
18936. 1 | ~2 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17114,18335,18332,18746,18328,18422,18321,18314,18311]
18938. 2 | 9 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17111,18335,18332,18746,18328,18422,18321,18340,18314]
18940. 2 | 10 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17109,18335,18332,18746,18328,18422,18321,18346,18314]
18942. 2 | 11 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17107,18335,18332,18746,18328,18422,18321,18352,18314]
18944. 2 | 12 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17105,18335,18332,18746,18328,18422,18321,18358,18314]
18946. 2 | 13 | 4 | 16 | 6 | 20 | 7 | ~8 [avatar split clause 17103,18335,18332,18746,18328,18422,18321,18364,18314]
18958. 1 | ~2 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17090,18335,18332,18746,18328,18448,18321,18314,18311]
18960. 2 | 9 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17087,18335,18332,18746,18328,18448,18321,18340,18314]
18962. 2 | 10 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17085,18335,18332,18746,18328,18448,18321,18346,18314]
18964. 2 | 11 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17083,18335,18332,18746,18328,18448,18321,18352,18314]
18966. 2 | 12 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17081,18335,18332,18746,18328,18448,18321,18358,18314]
18968. 2 | 13 | 4 | 17 | 6 | 20 | 7 | ~8 [avatar split clause 17079,18335,18332,18746,18328,18448,18321,18364,18314]
18980. 1 | ~2 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17066,18335,18332,18746,18328,18474,18321,18314,18311]
18982. 2 | 9 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17063,18335,18332,18746,18328,18474,18321,18340,18314]
18984. 2 | 10 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17061,18335,18332,18746,18328,18474,18321,18346,18314]
18986. 2 | 11 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17059,18335,18332,18746,18328,18474,18321,18352,18314]
18988. 2 | 12 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17057,18335,18332,18746,18328,18474,18321,18358,18314]
18990. 2 | 13 | 4 | 18 | 6 | 20 | 7 | ~8 [avatar split clause 17055,18335,18332,18746,18328,18474,18321,18364,18314]
18992. 21 <=> v1_zf_lang(sK4) [avatar definition]
18994. v1_zf_lang(sK4) <- (21) [avatar component clause 18992] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
19116. 1 | ~2 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16910,18335,18332,18992,18328,18321,18318,18314,18311]
19118. 2 | 9 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16907,18335,18332,18992,18328,18321,18318,18340,18314]
19120. 2 | 10 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16905,18335,18332,18992,18328,18321,18318,18346,18314]
19122. 2 | 11 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16903,18335,18332,18992,18328,18321,18318,18352,18314]
19124. 2 | 12 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16901,18335,18332,18992,18328,18321,18318,18358,18314]
19126. 2 | 13 | 3 | ~4 | 6 | 21 | 7 | ~8 [avatar split clause 16899,18335,18332,18992,18328,18321,18318,18364,18314]
19138. 1 | ~2 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16874,18335,18332,18992,18328,18370,18321,18314,18311]
19140. 2 | 9 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16871,18335,18332,18992,18328,18370,18321,18340,18314]
19142. 2 | 10 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16869,18335,18332,18992,18328,18370,18321,18346,18314]
19144. 2 | 11 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16867,18335,18332,18992,18328,18370,18321,18352,18314]
19146. 2 | 12 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16865,18335,18332,18992,18328,18370,18321,18358,18314]
19148. 2 | 13 | 4 | 14 | 6 | 21 | 7 | ~8 [avatar split clause 16863,18335,18332,18992,18328,18370,18321,18364,18314]
19160. 1 | ~2 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16850,18335,18332,18992,18328,18396,18321,18314,18311]
19162. 2 | 9 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16847,18335,18332,18992,18328,18396,18321,18340,18314]
19164. 2 | 10 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16845,18335,18332,18992,18328,18396,18321,18346,18314]
19166. 2 | 11 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16843,18335,18332,18992,18328,18396,18321,18352,18314]
19168. 2 | 12 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16841,18335,18332,18992,18328,18396,18321,18358,18314]
19170. 2 | 13 | 4 | 15 | 6 | 21 | 7 | ~8 [avatar split clause 16839,18335,18332,18992,18328,18396,18321,18364,18314]
19182. 1 | ~2 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16826,18335,18332,18992,18328,18422,18321,18314,18311]
19184. 2 | 9 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16823,18335,18332,18992,18328,18422,18321,18340,18314]
19186. 2 | 10 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16821,18335,18332,18992,18328,18422,18321,18346,18314]
19188. 2 | 11 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16819,18335,18332,18992,18328,18422,18321,18352,18314]
19190. 2 | 12 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16817,18335,18332,18992,18328,18422,18321,18358,18314]
19192. 2 | 13 | 4 | 16 | 6 | 21 | 7 | ~8 [avatar split clause 16815,18335,18332,18992,18328,18422,18321,18364,18314]
19204. 1 | ~2 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16802,18335,18332,18992,18328,18448,18321,18314,18311]
19206. 2 | 9 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16799,18335,18332,18992,18328,18448,18321,18340,18314]
19208. 2 | 10 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16797,18335,18332,18992,18328,18448,18321,18346,18314]
19210. 2 | 11 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16795,18335,18332,18992,18328,18448,18321,18352,18314]
19212. 2 | 12 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16793,18335,18332,18992,18328,18448,18321,18358,18314]
19214. 2 | 13 | 4 | 17 | 6 | 21 | 7 | ~8 [avatar split clause 16791,18335,18332,18992,18328,18448,18321,18364,18314]
19226. 1 | ~2 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16778,18335,18332,18992,18328,18474,18321,18314,18311]
19228. 2 | 9 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16775,18335,18332,18992,18328,18474,18321,18340,18314]
19230. 2 | 10 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16773,18335,18332,18992,18328,18474,18321,18346,18314]
19232. 2 | 11 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16771,18335,18332,18992,18328,18474,18321,18352,18314]
19234. 2 | 12 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16769,18335,18332,18992,18328,18474,18321,18358,18314]
19236. 2 | 13 | 4 | 18 | 6 | 21 | 7 | ~8 [avatar split clause 16767,18335,18332,18992,18328,18474,18321,18364,18314]
19238. 22 <=> m2_finseq_1(sK3,k5_numbers) [avatar definition]
19241. 1 | ~2 | 3 | ~4 | 5 | 22 | 7 | ~8 [avatar split clause 16766,18335,18332,19238,18325,18321,18318,18314,18311]
19249. 2 | 12 | 3 | ~4 | 5 | 22 | 7 | ~8 [avatar split clause 16757,18335,18332,19238,18325,18321,18318,18358,18314]
19320. 2 | 9 | 3 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16667,18335,18332,19238,18325,18448,18318,18340,18314]
19322. 2 | 10 | 3 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16665,18335,18332,19238,18325,18448,18318,18346,18314]
19324. 2 | 11 | 3 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16663,18335,18332,19238,18325,18448,18318,18352,18314]
19328. 2 | 13 | 3 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16659,18335,18332,19238,18325,18448,18318,18364,18314]
19331. 2 | 9 | 4 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16655,18335,18332,19238,18325,18448,18321,18340,18314]
19333. 2 | 10 | 4 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16653,18335,18332,19238,18325,18448,18321,18346,18314]
19335. 2 | 11 | 4 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16651,18335,18332,19238,18325,18448,18321,18352,18314]
19339. 2 | 13 | 4 | 17 | 5 | 22 | 7 | ~8 [avatar split clause 16647,18335,18332,19238,18325,18448,18321,18364,18314]
19353. 2 | 9 | 4 | 18 | 5 | 22 | 7 | ~8 [avatar split clause 16631,18335,18332,19238,18325,18474,18321,18340,18314]
19355. 2 | 10 | 4 | 18 | 5 | 22 | 7 | ~8 [avatar split clause 16629,18335,18332,19238,18325,18474,18321,18346,18314]
19357. 2 | 11 | 4 | 18 | 5 | 22 | 7 | ~8 [avatar split clause 16627,18335,18332,19238,18325,18474,18321,18352,18314]
19361. 2 | 13 | 4 | 18 | 5 | 22 | 7 | ~8 [avatar split clause 16623,18335,18332,19238,18325,18474,18321,18364,18314]
19362. 1 | ~2 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16622,18335,18332,19238,18328,18321,18318,18314,18311]
19364. 2 | 9 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16619,18335,18332,19238,18328,18321,18318,18340,18314]
19366. 2 | 10 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16617,18335,18332,19238,18328,18321,18318,18346,18314]
19368. 2 | 11 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16615,18335,18332,19238,18328,18321,18318,18352,18314]
19370. 2 | 12 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16613,18335,18332,19238,18328,18321,18318,18358,18314]
19372. 2 | 13 | 3 | ~4 | 6 | 22 | 7 | ~8 [avatar split clause 16611,18335,18332,19238,18328,18321,18318,18364,18314]
19384. 1 | ~2 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16586,18335,18332,19238,18328,18370,18321,18314,18311]
19386. 2 | 9 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16583,18335,18332,19238,18328,18370,18321,18340,18314]
19388. 2 | 10 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16581,18335,18332,19238,18328,18370,18321,18346,18314]
19390. 2 | 11 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16579,18335,18332,19238,18328,18370,18321,18352,18314]
19392. 2 | 12 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16577,18335,18332,19238,18328,18370,18321,18358,18314]
19394. 2 | 13 | 4 | 14 | 6 | 22 | 7 | ~8 [avatar split clause 16575,18335,18332,19238,18328,18370,18321,18364,18314]
19406. 1 | ~2 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16562,18335,18332,19238,18328,18396,18321,18314,18311]
19408. 2 | 9 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16559,18335,18332,19238,18328,18396,18321,18340,18314]
19410. 2 | 10 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16557,18335,18332,19238,18328,18396,18321,18346,18314]
19412. 2 | 11 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16555,18335,18332,19238,18328,18396,18321,18352,18314]
19414. 2 | 12 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16553,18335,18332,19238,18328,18396,18321,18358,18314]
19416. 2 | 13 | 4 | 15 | 6 | 22 | 7 | ~8 [avatar split clause 16551,18335,18332,19238,18328,18396,18321,18364,18314]
19428. 1 | ~2 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16538,18335,18332,19238,18328,18422,18321,18314,18311]
19430. 2 | 9 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16535,18335,18332,19238,18328,18422,18321,18340,18314]
19432. 2 | 10 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16533,18335,18332,19238,18328,18422,18321,18346,18314]
19434. 2 | 11 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16531,18335,18332,19238,18328,18422,18321,18352,18314]
19436. 2 | 12 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16529,18335,18332,19238,18328,18422,18321,18358,18314]
19438. 2 | 13 | 4 | 16 | 6 | 22 | 7 | ~8 [avatar split clause 16527,18335,18332,19238,18328,18422,18321,18364,18314]
19439. 1 | ~2 | 3 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16526,18335,18332,19238,18328,18448,18318,18314,18311]
19441. 2 | 9 | 3 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16523,18335,18332,19238,18328,18448,18318,18340,18314]
19443. 2 | 10 | 3 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16521,18335,18332,19238,18328,18448,18318,18346,18314]
19445. 2 | 11 | 3 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16519,18335,18332,19238,18328,18448,18318,18352,18314]
19449. 2 | 13 | 3 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16515,18335,18332,19238,18328,18448,18318,18364,18314]
19450. 1 | ~2 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16514,18335,18332,19238,18328,18448,18321,18314,18311]
19452. 2 | 9 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16511,18335,18332,19238,18328,18448,18321,18340,18314]
19454. 2 | 10 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16509,18335,18332,19238,18328,18448,18321,18346,18314]
19456. 2 | 11 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16507,18335,18332,19238,18328,18448,18321,18352,18314]
19458. 2 | 12 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16505,18335,18332,19238,18328,18448,18321,18358,18314]
19460. 2 | 13 | 4 | 17 | 6 | 22 | 7 | ~8 [avatar split clause 16503,18335,18332,19238,18328,18448,18321,18364,18314]
19472. 1 | ~2 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16490,18335,18332,19238,18328,18474,18321,18314,18311]
19474. 2 | 9 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16487,18335,18332,19238,18328,18474,18321,18340,18314]
19476. 2 | 10 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16485,18335,18332,19238,18328,18474,18321,18346,18314]
19478. 2 | 11 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16483,18335,18332,19238,18328,18474,18321,18352,18314]
19480. 2 | 12 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16481,18335,18332,19238,18328,18474,18321,18358,18314]
19482. 2 | 13 | 4 | 18 | 6 | 22 | 7 | ~8 [avatar split clause 16479,18335,18332,19238,18328,18474,18321,18364,18314]
19484. 23 <=> v1_zf_lang(sK3) [avatar definition]
19486. v1_zf_lang(sK3) <- (23) [avatar component clause 19484] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
19608. 1 | ~2 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16334,18335,18332,19484,18328,18321,18318,18314,18311]
19610. 2 | 9 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16331,18335,18332,19484,18328,18321,18318,18340,18314]
19612. 2 | 10 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16329,18335,18332,19484,18328,18321,18318,18346,18314]
19614. 2 | 11 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16327,18335,18332,19484,18328,18321,18318,18352,18314]
19616. 2 | 12 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16325,18335,18332,19484,18328,18321,18318,18358,18314]
19618. 2 | 13 | 3 | ~4 | 6 | 23 | 7 | ~8 [avatar split clause 16323,18335,18332,19484,18328,18321,18318,18364,18314]
19630. 1 | ~2 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16298,18335,18332,19484,18328,18370,18321,18314,18311]
19632. 2 | 9 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16295,18335,18332,19484,18328,18370,18321,18340,18314]
19634. 2 | 10 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16293,18335,18332,19484,18328,18370,18321,18346,18314]
19636. 2 | 11 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16291,18335,18332,19484,18328,18370,18321,18352,18314]
19638. 2 | 12 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16289,18335,18332,19484,18328,18370,18321,18358,18314]
19640. 2 | 13 | 4 | 14 | 6 | 23 | 7 | ~8 [avatar split clause 16287,18335,18332,19484,18328,18370,18321,18364,18314]
19652. 1 | ~2 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16274,18335,18332,19484,18328,18396,18321,18314,18311]
19654. 2 | 9 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16271,18335,18332,19484,18328,18396,18321,18340,18314]
19656. 2 | 10 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16269,18335,18332,19484,18328,18396,18321,18346,18314]
19658. 2 | 11 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16267,18335,18332,19484,18328,18396,18321,18352,18314]
19660. 2 | 12 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16265,18335,18332,19484,18328,18396,18321,18358,18314]
19662. 2 | 13 | 4 | 15 | 6 | 23 | 7 | ~8 [avatar split clause 16263,18335,18332,19484,18328,18396,18321,18364,18314]
19674. 1 | ~2 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16250,18335,18332,19484,18328,18422,18321,18314,18311]
19676. 2 | 9 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16247,18335,18332,19484,18328,18422,18321,18340,18314]
19678. 2 | 10 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16245,18335,18332,19484,18328,18422,18321,18346,18314]
19680. 2 | 11 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16243,18335,18332,19484,18328,18422,18321,18352,18314]
19682. 2 | 12 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16241,18335,18332,19484,18328,18422,18321,18358,18314]
19684. 2 | 13 | 4 | 16 | 6 | 23 | 7 | ~8 [avatar split clause 16239,18335,18332,19484,18328,18422,18321,18364,18314]
19696. 1 | ~2 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16226,18335,18332,19484,18328,18448,18321,18314,18311]
19698. 2 | 9 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16223,18335,18332,19484,18328,18448,18321,18340,18314]
19700. 2 | 10 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16221,18335,18332,19484,18328,18448,18321,18346,18314]
19702. 2 | 11 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16219,18335,18332,19484,18328,18448,18321,18352,18314]
19704. 2 | 12 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16217,18335,18332,19484,18328,18448,18321,18358,18314]
19706. 2 | 13 | 4 | 17 | 6 | 23 | 7 | ~8 [avatar split clause 16215,18335,18332,19484,18328,18448,18321,18364,18314]
19718. 1 | ~2 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16202,18335,18332,19484,18328,18474,18321,18314,18311]
19720. 2 | 9 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16199,18335,18332,19484,18328,18474,18321,18340,18314]
19722. 2 | 10 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16197,18335,18332,19484,18328,18474,18321,18346,18314]
19724. 2 | 11 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16195,18335,18332,19484,18328,18474,18321,18352,18314]
19726. 2 | 12 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16193,18335,18332,19484,18328,18474,18321,18358,18314]
19728. 2 | 13 | 4 | 18 | 6 | 23 | 7 | ~8 [avatar split clause 16191,18335,18332,19484,18328,18474,18321,18364,18314]
19730. 24 <=> sK0 = k13_zf_lang(sK1,sK2) [avatar definition]
19732. sK0 = k13_zf_lang(sK1,sK2) <- (24) [avatar component clause 19730] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
21064. 1 | ~2 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12734,19730,18335,18328,18325,18321,18318,18314,18311]
21066. 2 | 9 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12731,19730,18335,18328,18325,18321,18318,18340,18314]
21068. 2 | 10 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12729,19730,18335,18328,18325,18321,18318,18346,18314]
21070. 2 | 11 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12727,19730,18335,18328,18325,18321,18318,18352,18314]
21072. 2 | 12 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12725,19730,18335,18328,18325,18321,18318,18358,18314]
21074. 2 | 13 | 3 | ~4 | 5 | ~6 | 8 | 24 [avatar split clause 12723,19730,18335,18328,18325,18321,18318,18364,18314]
21086. 1 | ~2 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12698,19730,18335,18328,18325,18370,18321,18314,18311]
21087. 1 | 9 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12696,19730,18335,18328,18325,18370,18321,18340,18311]
21088. 2 | 9 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12695,19730,18335,18328,18325,18370,18321,18340,18314]
21090. 2 | 10 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12693,19730,18335,18328,18325,18370,18321,18346,18314]
21091. 1 | 11 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12692,19730,18335,18328,18325,18370,18321,18352,18311]
21092. 2 | 11 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12691,19730,18335,18328,18325,18370,18321,18352,18314]
21094. 2 | 12 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12689,19730,18335,18328,18325,18370,18321,18358,18314]
21095. 1 | 13 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12688,19730,18335,18328,18325,18370,18321,18364,18311]
21096. 2 | 13 | 4 | 14 | 5 | ~6 | 8 | 24 [avatar split clause 12687,19730,18335,18328,18325,18370,18321,18364,18314]
21108. 1 | ~2 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12674,19730,18335,18328,18325,18396,18321,18314,18311]
21110. 2 | 9 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12671,19730,18335,18328,18325,18396,18321,18340,18314]
21111. 1 | 10 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12670,19730,18335,18328,18325,18396,18321,18346,18311]
21112. 2 | 10 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12669,19730,18335,18328,18325,18396,18321,18346,18314]
21114. 2 | 11 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12667,19730,18335,18328,18325,18396,18321,18352,18314]
21115. 1 | 12 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12666,19730,18335,18328,18325,18396,18321,18358,18311]
21116. 2 | 12 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12665,19730,18335,18328,18325,18396,18321,18358,18314]
21118. 2 | 13 | 4 | 15 | 5 | ~6 | 8 | 24 [avatar split clause 12663,19730,18335,18328,18325,18396,18321,18364,18314]
21130. 1 | ~2 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12650,19730,18335,18328,18325,18422,18321,18314,18311]
21131. 1 | 9 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12648,19730,18335,18328,18325,18422,18321,18340,18311]
21132. 2 | 9 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12647,19730,18335,18328,18325,18422,18321,18340,18314]
21133. 1 | 10 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12646,19730,18335,18328,18325,18422,18321,18346,18311]
21134. 2 | 10 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12645,19730,18335,18328,18325,18422,18321,18346,18314]
21135. 1 | 11 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12644,19730,18335,18328,18325,18422,18321,18352,18311]
21136. 2 | 11 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12643,19730,18335,18328,18325,18422,18321,18352,18314]
21137. 1 | 12 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12642,19730,18335,18328,18325,18422,18321,18358,18311]
21138. 2 | 12 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12641,19730,18335,18328,18325,18422,18321,18358,18314]
21139. 1 | 13 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12640,19730,18335,18328,18325,18422,18321,18364,18311]
21140. 2 | 13 | 4 | 16 | 5 | ~6 | 8 | 24 [avatar split clause 12639,19730,18335,18328,18325,18422,18321,18364,18314]
21152. 1 | ~2 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12626,19730,18335,18328,18325,18448,18321,18314,18311]
21153. 1 | 9 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12624,19730,18335,18328,18325,18448,18321,18340,18311]
21154. 2 | 9 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12623,19730,18335,18328,18325,18448,18321,18340,18314]
21155. 1 | 10 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12622,19730,18335,18328,18325,18448,18321,18346,18311]
21156. 2 | 10 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12621,19730,18335,18328,18325,18448,18321,18346,18314]
21157. 1 | 11 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12620,19730,18335,18328,18325,18448,18321,18352,18311]
21158. 2 | 11 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12619,19730,18335,18328,18325,18448,18321,18352,18314]
21160. 2 | 12 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12617,19730,18335,18328,18325,18448,18321,18358,18314]
21161. 1 | 13 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12616,19730,18335,18328,18325,18448,18321,18364,18311]
21162. 2 | 13 | 4 | 17 | 5 | ~6 | 8 | 24 [avatar split clause 12615,19730,18335,18328,18325,18448,18321,18364,18314]
21174. 1 | ~2 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12602,19730,18335,18328,18325,18474,18321,18314,18311]
21175. 1 | 9 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12600,19730,18335,18328,18325,18474,18321,18340,18311]
21176. 2 | 9 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12599,19730,18335,18328,18325,18474,18321,18340,18314]
21177. 1 | 10 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12598,19730,18335,18328,18325,18474,18321,18346,18311]
21178. 2 | 10 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12597,19730,18335,18328,18325,18474,18321,18346,18314]
21179. 1 | 11 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12596,19730,18335,18328,18325,18474,18321,18352,18311]
21180. 2 | 11 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12595,19730,18335,18328,18325,18474,18321,18352,18314]
21181. 1 | 12 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12594,19730,18335,18328,18325,18474,18321,18358,18311]
21182. 2 | 12 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12593,19730,18335,18328,18325,18474,18321,18358,18314]
21183. 1 | 13 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12592,19730,18335,18328,18325,18474,18321,18364,18311]
21184. 2 | 13 | 4 | 18 | 5 | ~6 | 8 | 24 [avatar split clause 12591,19730,18335,18328,18325,18474,18321,18364,18314]
21306. 1 | ~2 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12302,19730,18335,18500,18328,18321,18318,18314,18311]
21308. 2 | 9 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12299,19730,18335,18500,18328,18321,18318,18340,18314]
21310. 2 | 10 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12297,19730,18335,18500,18328,18321,18318,18346,18314]
21312. 2 | 11 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12295,19730,18335,18500,18328,18321,18318,18352,18314]
21314. 2 | 12 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12293,19730,18335,18500,18328,18321,18318,18358,18314]
21316. 2 | 13 | 3 | ~4 | 6 | 19 | 8 | 24 [avatar split clause 12291,19730,18335,18500,18328,18321,18318,18364,18314]
21317. 1 | ~2 | 3 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12278,19730,18335,18500,18328,18370,18318,18314,18311]
21321. 2 | 10 | 3 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12273,19730,18335,18500,18328,18370,18318,18346,18314]
21325. 2 | 12 | 3 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12269,19730,18335,18500,18328,18370,18318,18358,18314]
21327. 2 | 13 | 3 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12267,19730,18335,18500,18328,18370,18318,18364,18314]
21328. 1 | ~2 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12266,19730,18335,18500,18328,18370,18321,18314,18311]
21330. 2 | 9 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12263,19730,18335,18500,18328,18370,18321,18340,18314]
21332. 2 | 10 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12261,19730,18335,18500,18328,18370,18321,18346,18314]
21334. 2 | 11 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12259,19730,18335,18500,18328,18370,18321,18352,18314]
21336. 2 | 12 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12257,19730,18335,18500,18328,18370,18321,18358,18314]
21338. 2 | 13 | 4 | 14 | 6 | 19 | 8 | 24 [avatar split clause 12255,19730,18335,18500,18328,18370,18321,18364,18314]
21339. 1 | ~2 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12254,19730,18335,18500,18328,18396,18318,18314,18311]
21341. 2 | 9 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12251,19730,18335,18500,18328,18396,18318,18340,18314]
21343. 2 | 10 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12249,19730,18335,18500,18328,18396,18318,18346,18314]
21345. 2 | 11 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12247,19730,18335,18500,18328,18396,18318,18352,18314]
21347. 2 | 12 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12245,19730,18335,18500,18328,18396,18318,18358,18314]
21349. 2 | 13 | 3 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12243,19730,18335,18500,18328,18396,18318,18364,18314]
21350. 1 | ~2 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12242,19730,18335,18500,18328,18396,18321,18314,18311]
21352. 2 | 9 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12239,19730,18335,18500,18328,18396,18321,18340,18314]
21354. 2 | 10 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12237,19730,18335,18500,18328,18396,18321,18346,18314]
21356. 2 | 11 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12235,19730,18335,18500,18328,18396,18321,18352,18314]
21358. 2 | 12 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12233,19730,18335,18500,18328,18396,18321,18358,18314]
21360. 2 | 13 | 4 | 15 | 6 | 19 | 8 | 24 [avatar split clause 12231,19730,18335,18500,18328,18396,18321,18364,18314]
21361. 1 | ~2 | 3 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12230,19730,18335,18500,18328,18422,18318,18314,18311]
21363. 2 | 9 | 3 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12227,19730,18335,18500,18328,18422,18318,18340,18314]
21367. 2 | 11 | 3 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12223,19730,18335,18500,18328,18422,18318,18352,18314]
21371. 2 | 13 | 3 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12219,19730,18335,18500,18328,18422,18318,18364,18314]
21372. 1 | ~2 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12218,19730,18335,18500,18328,18422,18321,18314,18311]
21374. 2 | 9 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12215,19730,18335,18500,18328,18422,18321,18340,18314]
21376. 2 | 10 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12213,19730,18335,18500,18328,18422,18321,18346,18314]
21378. 2 | 11 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12211,19730,18335,18500,18328,18422,18321,18352,18314]
21380. 2 | 12 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12209,19730,18335,18500,18328,18422,18321,18358,18314]
21382. 2 | 13 | 4 | 16 | 6 | 19 | 8 | 24 [avatar split clause 12207,19730,18335,18500,18328,18422,18321,18364,18314]
21383. 1 | ~2 | 3 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12206,19730,18335,18500,18328,18448,18318,18314,18311]
21387. 2 | 10 | 3 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12201,19730,18335,18500,18328,18448,18318,18346,18314]
21391. 2 | 12 | 3 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12197,19730,18335,18500,18328,18448,18318,18358,18314]
21394. 1 | ~2 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12194,19730,18335,18500,18328,18448,18321,18314,18311]
21396. 2 | 9 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12191,19730,18335,18500,18328,18448,18321,18340,18314]
21398. 2 | 10 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12189,19730,18335,18500,18328,18448,18321,18346,18314]
21400. 2 | 11 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12187,19730,18335,18500,18328,18448,18321,18352,18314]
21402. 2 | 12 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12185,19730,18335,18500,18328,18448,18321,18358,18314]
21404. 2 | 13 | 4 | 17 | 6 | 19 | 8 | 24 [avatar split clause 12183,19730,18335,18500,18328,18448,18321,18364,18314]
21405. 1 | ~2 | 3 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12182,19730,18335,18500,18328,18474,18318,18314,18311]
21415. 2 | 13 | 3 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12171,19730,18335,18500,18328,18474,18318,18364,18314]
21416. 1 | ~2 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12170,19730,18335,18500,18328,18474,18321,18314,18311]
21418. 2 | 9 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12167,19730,18335,18500,18328,18474,18321,18340,18314]
21420. 2 | 10 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12165,19730,18335,18500,18328,18474,18321,18346,18314]
21422. 2 | 11 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12163,19730,18335,18500,18328,18474,18321,18352,18314]
21424. 2 | 12 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12161,19730,18335,18500,18328,18474,18321,18358,18314]
21426. 2 | 13 | 4 | 18 | 6 | 19 | 8 | 24 [avatar split clause 12159,19730,18335,18500,18328,18474,18321,18364,18314]
21435. 2 | 12 | 3 | ~4 | 5 | 20 | 8 | 24 [avatar split clause 12149,19730,18335,18746,18325,18321,18318,18358,18314]
21548. 1 | ~2 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12014,19730,18335,18746,18328,18321,18318,18314,18311]
21550. 2 | 9 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12011,19730,18335,18746,18328,18321,18318,18340,18314]
21552. 2 | 10 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12009,19730,18335,18746,18328,18321,18318,18346,18314]
21554. 2 | 11 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12007,19730,18335,18746,18328,18321,18318,18352,18314]
21556. 2 | 12 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12005,19730,18335,18746,18328,18321,18318,18358,18314]
21558. 2 | 13 | 3 | ~4 | 6 | 20 | 8 | 24 [avatar split clause 12003,19730,18335,18746,18328,18321,18318,18364,18314]
21561. 2 | 9 | 3 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11987,19730,18335,18746,18328,18370,18318,18340,18314]
21563. 2 | 10 | 3 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11985,19730,18335,18746,18328,18370,18318,18346,18314]
21565. 2 | 11 | 3 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11983,19730,18335,18746,18328,18370,18318,18352,18314]
21570. 1 | ~2 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11978,19730,18335,18746,18328,18370,18321,18314,18311]
21572. 2 | 9 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11975,19730,18335,18746,18328,18370,18321,18340,18314]
21574. 2 | 10 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11973,19730,18335,18746,18328,18370,18321,18346,18314]
21576. 2 | 11 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11971,19730,18335,18746,18328,18370,18321,18352,18314]
21578. 2 | 12 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11969,19730,18335,18746,18328,18370,18321,18358,18314]
21580. 2 | 13 | 4 | 14 | 6 | 20 | 8 | 24 [avatar split clause 11967,19730,18335,18746,18328,18370,18321,18364,18314]
21583. 2 | 9 | 3 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11963,19730,18335,18746,18328,18396,18318,18340,18314]
21585. 2 | 10 | 3 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11961,19730,18335,18746,18328,18396,18318,18346,18314]
21587. 2 | 11 | 3 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11959,19730,18335,18746,18328,18396,18318,18352,18314]
21591. 2 | 13 | 3 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11955,19730,18335,18746,18328,18396,18318,18364,18314]
21592. 1 | ~2 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11954,19730,18335,18746,18328,18396,18321,18314,18311]
21594. 2 | 9 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11951,19730,18335,18746,18328,18396,18321,18340,18314]
21596. 2 | 10 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11949,19730,18335,18746,18328,18396,18321,18346,18314]
21598. 2 | 11 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11947,19730,18335,18746,18328,18396,18321,18352,18314]
21600. 2 | 12 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11945,19730,18335,18746,18328,18396,18321,18358,18314]
21602. 2 | 13 | 4 | 15 | 6 | 20 | 8 | 24 [avatar split clause 11943,19730,18335,18746,18328,18396,18321,18364,18314]
21605. 2 | 9 | 3 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11939,19730,18335,18746,18328,18422,18318,18340,18314]
21607. 2 | 10 | 3 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11937,19730,18335,18746,18328,18422,18318,18346,18314]
21609. 2 | 11 | 3 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11935,19730,18335,18746,18328,18422,18318,18352,18314]
21613. 2 | 13 | 3 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11931,19730,18335,18746,18328,18422,18318,18364,18314]
21614. 1 | ~2 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11930,19730,18335,18746,18328,18422,18321,18314,18311]
21616. 2 | 9 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11927,19730,18335,18746,18328,18422,18321,18340,18314]
21618. 2 | 10 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11925,19730,18335,18746,18328,18422,18321,18346,18314]
21620. 2 | 11 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11923,19730,18335,18746,18328,18422,18321,18352,18314]
21622. 2 | 12 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11921,19730,18335,18746,18328,18422,18321,18358,18314]
21624. 2 | 13 | 4 | 16 | 6 | 20 | 8 | 24 [avatar split clause 11919,19730,18335,18746,18328,18422,18321,18364,18314]
21627. 2 | 9 | 3 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11915,19730,18335,18746,18328,18448,18318,18340,18314]
21629. 2 | 10 | 3 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11913,19730,18335,18746,18328,18448,18318,18346,18314]
21631. 2 | 11 | 3 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11911,19730,18335,18746,18328,18448,18318,18352,18314]
21633. 2 | 12 | 3 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11909,19730,18335,18746,18328,18448,18318,18358,18314]
21635. 2 | 13 | 3 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11907,19730,18335,18746,18328,18448,18318,18364,18314]
21636. 1 | ~2 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11906,19730,18335,18746,18328,18448,18321,18314,18311]
21638. 2 | 9 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11903,19730,18335,18746,18328,18448,18321,18340,18314]
21640. 2 | 10 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11901,19730,18335,18746,18328,18448,18321,18346,18314]
21642. 2 | 11 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11899,19730,18335,18746,18328,18448,18321,18352,18314]
21644. 2 | 12 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11897,19730,18335,18746,18328,18448,18321,18358,18314]
21646. 2 | 13 | 4 | 17 | 6 | 20 | 8 | 24 [avatar split clause 11895,19730,18335,18746,18328,18448,18321,18364,18314]
21651. 2 | 10 | 3 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11889,19730,18335,18746,18328,18474,18318,18346,18314]
21653. 2 | 11 | 3 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11887,19730,18335,18746,18328,18474,18318,18352,18314]
21655. 2 | 12 | 3 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11885,19730,18335,18746,18328,18474,18318,18358,18314]
21658. 1 | ~2 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11882,19730,18335,18746,18328,18474,18321,18314,18311]
21660. 2 | 9 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11879,19730,18335,18746,18328,18474,18321,18340,18314]
21662. 2 | 10 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11877,19730,18335,18746,18328,18474,18321,18346,18314]
21664. 2 | 11 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11875,19730,18335,18746,18328,18474,18321,18352,18314]
21666. 2 | 12 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11873,19730,18335,18746,18328,18474,18321,18358,18314]
21668. 2 | 13 | 4 | 18 | 6 | 20 | 8 | 24 [avatar split clause 11871,19730,18335,18746,18328,18474,18321,18364,18314]
21790. 1 | ~2 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11726,19730,18335,18992,18328,18321,18318,18314,18311]
21792. 2 | 9 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11723,19730,18335,18992,18328,18321,18318,18340,18314]
21794. 2 | 10 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11721,19730,18335,18992,18328,18321,18318,18346,18314]
21796. 2 | 11 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11719,19730,18335,18992,18328,18321,18318,18352,18314]
21798. 2 | 12 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11717,19730,18335,18992,18328,18321,18318,18358,18314]
21800. 2 | 13 | 3 | ~4 | 6 | 21 | 8 | 24 [avatar split clause 11715,19730,18335,18992,18328,18321,18318,18364,18314]
21801. 1 | ~2 | 3 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11702,19730,18335,18992,18328,18370,18318,18314,18311]
21805. 2 | 10 | 3 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11697,19730,18335,18992,18328,18370,18318,18346,18314]
21809. 2 | 12 | 3 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11693,19730,18335,18992,18328,18370,18318,18358,18314]
21811. 2 | 13 | 3 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11691,19730,18335,18992,18328,18370,18318,18364,18314]
21812. 1 | ~2 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11690,19730,18335,18992,18328,18370,18321,18314,18311]
21814. 2 | 9 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11687,19730,18335,18992,18328,18370,18321,18340,18314]
21816. 2 | 10 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11685,19730,18335,18992,18328,18370,18321,18346,18314]
21818. 2 | 11 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11683,19730,18335,18992,18328,18370,18321,18352,18314]
21820. 2 | 12 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11681,19730,18335,18992,18328,18370,18321,18358,18314]
21822. 2 | 13 | 4 | 14 | 6 | 21 | 8 | 24 [avatar split clause 11679,19730,18335,18992,18328,18370,18321,18364,18314]
21827. 2 | 10 | 3 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11673,19730,18335,18992,18328,18396,18318,18346,18314]
21834. 1 | ~2 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11666,19730,18335,18992,18328,18396,18321,18314,18311]
21836. 2 | 9 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11663,19730,18335,18992,18328,18396,18321,18340,18314]
21838. 2 | 10 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11661,19730,18335,18992,18328,18396,18321,18346,18314]
21840. 2 | 11 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11659,19730,18335,18992,18328,18396,18321,18352,18314]
21842. 2 | 12 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11657,19730,18335,18992,18328,18396,18321,18358,18314]
21844. 2 | 13 | 4 | 15 | 6 | 21 | 8 | 24 [avatar split clause 11655,19730,18335,18992,18328,18396,18321,18364,18314]
21845. 1 | ~2 | 3 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11654,19730,18335,18992,18328,18422,18318,18314,18311]
21847. 2 | 9 | 3 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11651,19730,18335,18992,18328,18422,18318,18340,18314]
21851. 2 | 11 | 3 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11647,19730,18335,18992,18328,18422,18318,18352,18314]
21855. 2 | 13 | 3 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11643,19730,18335,18992,18328,18422,18318,18364,18314]
21856. 1 | ~2 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11642,19730,18335,18992,18328,18422,18321,18314,18311]
21858. 2 | 9 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11639,19730,18335,18992,18328,18422,18321,18340,18314]
21860. 2 | 10 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11637,19730,18335,18992,18328,18422,18321,18346,18314]
21862. 2 | 11 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11635,19730,18335,18992,18328,18422,18321,18352,18314]
21864. 2 | 12 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11633,19730,18335,18992,18328,18422,18321,18358,18314]
21866. 2 | 13 | 4 | 16 | 6 | 21 | 8 | 24 [avatar split clause 11631,19730,18335,18992,18328,18422,18321,18364,18314]
21867. 1 | ~2 | 3 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11630,19730,18335,18992,18328,18448,18318,18314,18311]
21878. 1 | ~2 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11618,19730,18335,18992,18328,18448,18321,18314,18311]
21880. 2 | 9 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11615,19730,18335,18992,18328,18448,18321,18340,18314]
21882. 2 | 10 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11613,19730,18335,18992,18328,18448,18321,18346,18314]
21884. 2 | 11 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11611,19730,18335,18992,18328,18448,18321,18352,18314]
21886. 2 | 12 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11609,19730,18335,18992,18328,18448,18321,18358,18314]
21888. 2 | 13 | 4 | 17 | 6 | 21 | 8 | 24 [avatar split clause 11607,19730,18335,18992,18328,18448,18321,18364,18314]
21889. 1 | ~2 | 3 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11606,19730,18335,18992,18328,18474,18318,18314,18311]
21897. 2 | 12 | 3 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11597,19730,18335,18992,18328,18474,18318,18358,18314]
21900. 1 | ~2 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11594,19730,18335,18992,18328,18474,18321,18314,18311]
21902. 2 | 9 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11591,19730,18335,18992,18328,18474,18321,18340,18314]
21904. 2 | 10 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11589,19730,18335,18992,18328,18474,18321,18346,18314]
21906. 2 | 11 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11587,19730,18335,18992,18328,18474,18321,18352,18314]
21908. 2 | 12 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11585,19730,18335,18992,18328,18474,18321,18358,18314]
21910. 2 | 13 | 4 | 18 | 6 | 21 | 8 | 24 [avatar split clause 11583,19730,18335,18992,18328,18474,18321,18364,18314]
21919. 2 | 12 | 3 | ~4 | 5 | 22 | 8 | 24 [avatar split clause 11573,19730,18335,19238,18325,18321,18318,18358,18314]
22032. 1 | ~2 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11438,19730,18335,19238,18328,18321,18318,18314,18311]
22034. 2 | 9 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11435,19730,18335,19238,18328,18321,18318,18340,18314]
22036. 2 | 10 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11433,19730,18335,19238,18328,18321,18318,18346,18314]
22038. 2 | 11 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11431,19730,18335,19238,18328,18321,18318,18352,18314]
22040. 2 | 12 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11429,19730,18335,19238,18328,18321,18318,18358,18314]
22042. 2 | 13 | 3 | ~4 | 6 | 22 | 8 | 24 [avatar split clause 11427,19730,18335,19238,18328,18321,18318,18364,18314]
22047. 2 | 10 | 3 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11409,19730,18335,19238,18328,18370,18318,18346,18314]
22053. 2 | 13 | 3 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11403,19730,18335,19238,18328,18370,18318,18364,18314]
22054. 1 | ~2 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11402,19730,18335,19238,18328,18370,18321,18314,18311]
22056. 2 | 9 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11399,19730,18335,19238,18328,18370,18321,18340,18314]
22058. 2 | 10 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11397,19730,18335,19238,18328,18370,18321,18346,18314]
22060. 2 | 11 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11395,19730,18335,19238,18328,18370,18321,18352,18314]
22062. 2 | 12 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11393,19730,18335,19238,18328,18370,18321,18358,18314]
22064. 2 | 13 | 4 | 14 | 6 | 22 | 8 | 24 [avatar split clause 11391,19730,18335,19238,18328,18370,18321,18364,18314]
22067. 2 | 9 | 3 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11387,19730,18335,19238,18328,18396,18318,18340,18314]
22069. 2 | 10 | 3 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11385,19730,18335,19238,18328,18396,18318,18346,18314]
22071. 2 | 11 | 3 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11383,19730,18335,19238,18328,18396,18318,18352,18314]
22073. 2 | 12 | 3 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11381,19730,18335,19238,18328,18396,18318,18358,18314]
22075. 2 | 13 | 3 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11379,19730,18335,19238,18328,18396,18318,18364,18314]
22076. 1 | ~2 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11378,19730,18335,19238,18328,18396,18321,18314,18311]
22078. 2 | 9 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11375,19730,18335,19238,18328,18396,18321,18340,18314]
22080. 2 | 10 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11373,19730,18335,19238,18328,18396,18321,18346,18314]
22082. 2 | 11 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11371,19730,18335,19238,18328,18396,18321,18352,18314]
22084. 2 | 12 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11369,19730,18335,19238,18328,18396,18321,18358,18314]
22086. 2 | 13 | 4 | 15 | 6 | 22 | 8 | 24 [avatar split clause 11367,19730,18335,19238,18328,18396,18321,18364,18314]
22089. 2 | 9 | 3 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11363,19730,18335,19238,18328,18422,18318,18340,18314]
22097. 2 | 13 | 3 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11355,19730,18335,19238,18328,18422,18318,18364,18314]
22098. 1 | ~2 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11354,19730,18335,19238,18328,18422,18321,18314,18311]
22100. 2 | 9 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11351,19730,18335,19238,18328,18422,18321,18340,18314]
22102. 2 | 10 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11349,19730,18335,19238,18328,18422,18321,18346,18314]
22104. 2 | 11 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11347,19730,18335,19238,18328,18422,18321,18352,18314]
22106. 2 | 12 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11345,19730,18335,19238,18328,18422,18321,18358,18314]
22108. 2 | 13 | 4 | 16 | 6 | 22 | 8 | 24 [avatar split clause 11343,19730,18335,19238,18328,18422,18321,18364,18314]
22113. 2 | 10 | 3 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11337,19730,18335,19238,18328,18448,18318,18346,18314]
22117. 2 | 12 | 3 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11333,19730,18335,19238,18328,18448,18318,18358,18314]
22120. 1 | ~2 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11330,19730,18335,19238,18328,18448,18321,18314,18311]
22122. 2 | 9 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11327,19730,18335,19238,18328,18448,18321,18340,18314]
22124. 2 | 10 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11325,19730,18335,19238,18328,18448,18321,18346,18314]
22126. 2 | 11 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11323,19730,18335,19238,18328,18448,18321,18352,18314]
22128. 2 | 12 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11321,19730,18335,19238,18328,18448,18321,18358,18314]
22130. 2 | 13 | 4 | 17 | 6 | 22 | 8 | 24 [avatar split clause 11319,19730,18335,19238,18328,18448,18321,18364,18314]
22133. 2 | 9 | 3 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11315,19730,18335,19238,18328,18474,18318,18340,18314]
22135. 2 | 10 | 3 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11313,19730,18335,19238,18328,18474,18318,18346,18314]
22139. 2 | 12 | 3 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11309,19730,18335,19238,18328,18474,18318,18358,18314]
22141. 2 | 13 | 3 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11307,19730,18335,19238,18328,18474,18318,18364,18314]
22142. 1 | ~2 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11306,19730,18335,19238,18328,18474,18321,18314,18311]
22144. 2 | 9 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11303,19730,18335,19238,18328,18474,18321,18340,18314]
22146. 2 | 10 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11301,19730,18335,19238,18328,18474,18321,18346,18314]
22148. 2 | 11 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11299,19730,18335,19238,18328,18474,18321,18352,18314]
22150. 2 | 12 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11297,19730,18335,19238,18328,18474,18321,18358,18314]
22152. 2 | 13 | 4 | 18 | 6 | 22 | 8 | 24 [avatar split clause 11295,19730,18335,19238,18328,18474,18321,18364,18314]
22274. 1 | ~2 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11150,19730,18335,19484,18328,18321,18318,18314,18311]
22276. 2 | 9 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11147,19730,18335,19484,18328,18321,18318,18340,18314]
22278. 2 | 10 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11145,19730,18335,19484,18328,18321,18318,18346,18314]
22280. 2 | 11 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11143,19730,18335,19484,18328,18321,18318,18352,18314]
22282. 2 | 12 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11141,19730,18335,19484,18328,18321,18318,18358,18314]
22284. 2 | 13 | 3 | ~4 | 6 | 23 | 8 | 24 [avatar split clause 11139,19730,18335,19484,18328,18321,18318,18364,18314]
22285. 1 | ~2 | 3 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11126,19730,18335,19484,18328,18370,18318,18314,18311]
22295. 2 | 13 | 3 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11115,19730,18335,19484,18328,18370,18318,18364,18314]
22296. 1 | ~2 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11114,19730,18335,19484,18328,18370,18321,18314,18311]
22298. 2 | 9 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11111,19730,18335,19484,18328,18370,18321,18340,18314]
22300. 2 | 10 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11109,19730,18335,19484,18328,18370,18321,18346,18314]
22302. 2 | 11 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11107,19730,18335,19484,18328,18370,18321,18352,18314]
22304. 2 | 12 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11105,19730,18335,19484,18328,18370,18321,18358,18314]
22306. 2 | 13 | 4 | 14 | 6 | 23 | 8 | 24 [avatar split clause 11103,19730,18335,19484,18328,18370,18321,18364,18314]
22307. 1 | ~2 | 3 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11102,19730,18335,19484,18328,18396,18318,18314,18311]
22309. 2 | 9 | 3 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11099,19730,18335,19484,18328,18396,18318,18340,18314]
22318. 1 | ~2 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11090,19730,18335,19484,18328,18396,18321,18314,18311]
22320. 2 | 9 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11087,19730,18335,19484,18328,18396,18321,18340,18314]
22322. 2 | 10 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11085,19730,18335,19484,18328,18396,18321,18346,18314]
22324. 2 | 11 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11083,19730,18335,19484,18328,18396,18321,18352,18314]
22326. 2 | 12 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11081,19730,18335,19484,18328,18396,18321,18358,18314]
22328. 2 | 13 | 4 | 15 | 6 | 23 | 8 | 24 [avatar split clause 11079,19730,18335,19484,18328,18396,18321,18364,18314]
22331. 2 | 9 | 3 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11075,19730,18335,19484,18328,18422,18318,18340,18314]
22335. 2 | 11 | 3 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11071,19730,18335,19484,18328,18422,18318,18352,18314]
22339. 2 | 13 | 3 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11067,19730,18335,19484,18328,18422,18318,18364,18314]
22340. 1 | ~2 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11066,19730,18335,19484,18328,18422,18321,18314,18311]
22342. 2 | 9 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11063,19730,18335,19484,18328,18422,18321,18340,18314]
22344. 2 | 10 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11061,19730,18335,19484,18328,18422,18321,18346,18314]
22346. 2 | 11 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11059,19730,18335,19484,18328,18422,18321,18352,18314]
22348. 2 | 12 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11057,19730,18335,19484,18328,18422,18321,18358,18314]
22350. 2 | 13 | 4 | 16 | 6 | 23 | 8 | 24 [avatar split clause 11055,19730,18335,19484,18328,18422,18321,18364,18314]
22351. 1 | ~2 | 3 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11054,19730,18335,19484,18328,18448,18318,18314,18311]
22353. 2 | 9 | 3 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11051,19730,18335,19484,18328,18448,18318,18340,18314]
22357. 2 | 11 | 3 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11047,19730,18335,19484,18328,18448,18318,18352,18314]
22359. 2 | 12 | 3 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11045,19730,18335,19484,18328,18448,18318,18358,18314]
22361. 2 | 13 | 3 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11043,19730,18335,19484,18328,18448,18318,18364,18314]
22362. 1 | ~2 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11042,19730,18335,19484,18328,18448,18321,18314,18311]
22364. 2 | 9 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11039,19730,18335,19484,18328,18448,18321,18340,18314]
22366. 2 | 10 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11037,19730,18335,19484,18328,18448,18321,18346,18314]
22368. 2 | 11 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11035,19730,18335,19484,18328,18448,18321,18352,18314]
22370. 2 | 12 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11033,19730,18335,19484,18328,18448,18321,18358,18314]
22372. 2 | 13 | 4 | 17 | 6 | 23 | 8 | 24 [avatar split clause 11031,19730,18335,19484,18328,18448,18321,18364,18314]
22375. 2 | 9 | 3 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11027,19730,18335,19484,18328,18474,18318,18340,18314]
22379. 2 | 11 | 3 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11023,19730,18335,19484,18328,18474,18318,18352,18314]
22383. 2 | 13 | 3 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11019,19730,18335,19484,18328,18474,18318,18364,18314]
22384. 1 | ~2 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11018,19730,18335,19484,18328,18474,18321,18314,18311]
22386. 2 | 9 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11015,19730,18335,19484,18328,18474,18321,18340,18314]
22388. 2 | 10 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11013,19730,18335,19484,18328,18474,18321,18346,18314]
22390. 2 | 11 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11011,19730,18335,19484,18328,18474,18321,18352,18314]
22392. 2 | 12 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11009,19730,18335,19484,18328,18474,18321,18358,18314]
22394. 2 | 13 | 4 | 18 | 6 | 23 | 8 | 24 [avatar split clause 11007,19730,18335,19484,18328,18474,18321,18364,18314]
22396. 25 <=> m2_finseq_1(sK2,k5_numbers) [avatar definition]
23730. 1 | ~2 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9278,22396,18335,18328,18325,18321,18318,18314,18311]
23732. 2 | 9 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9275,22396,18335,18328,18325,18321,18318,18340,18314]
23734. 2 | 10 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9273,22396,18335,18328,18325,18321,18318,18346,18314]
23736. 2 | 11 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9271,22396,18335,18328,18325,18321,18318,18352,18314]
23738. 2 | 12 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9269,22396,18335,18328,18325,18321,18318,18358,18314]
23740. 2 | 13 | 3 | ~4 | 5 | ~6 | 8 | 25 [avatar split clause 9267,22396,18335,18328,18325,18321,18318,18364,18314]
23752. 1 | ~2 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9242,22396,18335,18328,18325,18370,18321,18314,18311]
23753. 1 | 9 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9240,22396,18335,18328,18325,18370,18321,18340,18311]
23754. 2 | 9 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9239,22396,18335,18328,18325,18370,18321,18340,18314]
23755. 1 | 10 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9238,22396,18335,18328,18325,18370,18321,18346,18311]
23756. 2 | 10 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9237,22396,18335,18328,18325,18370,18321,18346,18314]
23757. 1 | 11 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9236,22396,18335,18328,18325,18370,18321,18352,18311]
23758. 2 | 11 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9235,22396,18335,18328,18325,18370,18321,18352,18314]
23759. 1 | 12 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9234,22396,18335,18328,18325,18370,18321,18358,18311]
23760. 2 | 12 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9233,22396,18335,18328,18325,18370,18321,18358,18314]
23761. 1 | 13 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9232,22396,18335,18328,18325,18370,18321,18364,18311]
23762. 2 | 13 | 4 | 14 | 5 | ~6 | 8 | 25 [avatar split clause 9231,22396,18335,18328,18325,18370,18321,18364,18314]
23774. 1 | ~2 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9218,22396,18335,18328,18325,18396,18321,18314,18311]
23776. 2 | 9 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9215,22396,18335,18328,18325,18396,18321,18340,18314]
23777. 1 | 10 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9214,22396,18335,18328,18325,18396,18321,18346,18311]
23778. 2 | 10 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9213,22396,18335,18328,18325,18396,18321,18346,18314]
23780. 2 | 11 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9211,22396,18335,18328,18325,18396,18321,18352,18314]
23782. 2 | 12 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9209,22396,18335,18328,18325,18396,18321,18358,18314]
23784. 2 | 13 | 4 | 15 | 5 | ~6 | 8 | 25 [avatar split clause 9207,22396,18335,18328,18325,18396,18321,18364,18314]
23796. 1 | ~2 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9194,22396,18335,18328,18325,18422,18321,18314,18311]
23797. 1 | 9 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9192,22396,18335,18328,18325,18422,18321,18340,18311]
23798. 2 | 9 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9191,22396,18335,18328,18325,18422,18321,18340,18314]
23799. 1 | 10 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9190,22396,18335,18328,18325,18422,18321,18346,18311]
23800. 2 | 10 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9189,22396,18335,18328,18325,18422,18321,18346,18314]
23801. 1 | 11 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9188,22396,18335,18328,18325,18422,18321,18352,18311]
23802. 2 | 11 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9187,22396,18335,18328,18325,18422,18321,18352,18314]
23803. 1 | 12 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9186,22396,18335,18328,18325,18422,18321,18358,18311]
23804. 2 | 12 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9185,22396,18335,18328,18325,18422,18321,18358,18314]
23805. 1 | 13 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9184,22396,18335,18328,18325,18422,18321,18364,18311]
23806. 2 | 13 | 4 | 16 | 5 | ~6 | 8 | 25 [avatar split clause 9183,22396,18335,18328,18325,18422,18321,18364,18314]
23818. 1 | ~2 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9170,22396,18335,18328,18325,18448,18321,18314,18311]
23820. 2 | 9 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9167,22396,18335,18328,18325,18448,18321,18340,18314]
23822. 2 | 10 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9165,22396,18335,18328,18325,18448,18321,18346,18314]
23824. 2 | 11 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9163,22396,18335,18328,18325,18448,18321,18352,18314]
23826. 2 | 12 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9161,22396,18335,18328,18325,18448,18321,18358,18314]
23828. 2 | 13 | 4 | 17 | 5 | ~6 | 8 | 25 [avatar split clause 9159,22396,18335,18328,18325,18448,18321,18364,18314]
23840. 1 | ~2 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9146,22396,18335,18328,18325,18474,18321,18314,18311]
23842. 2 | 9 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9143,22396,18335,18328,18325,18474,18321,18340,18314]
23844. 2 | 10 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9141,22396,18335,18328,18325,18474,18321,18346,18314]
23846. 2 | 11 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9139,22396,18335,18328,18325,18474,18321,18352,18314]
23847. 1 | 12 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9138,22396,18335,18328,18325,18474,18321,18358,18311]
23848. 2 | 12 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9137,22396,18335,18328,18325,18474,18321,18358,18314]
23850. 2 | 13 | 4 | 18 | 5 | ~6 | 8 | 25 [avatar split clause 9135,22396,18335,18328,18325,18474,18321,18364,18314]
23855. 2 | 10 | 3 | ~4 | 5 | 19 | 8 | 25 [avatar split clause 8985,22396,18335,18500,18325,18321,18318,18346,18314]
23857. 2 | 11 | 3 | ~4 | 5 | 19 | 8 | 25 [avatar split clause 8983,22396,18335,18500,18325,18321,18318,18352,18314]
23972. 1 | ~2 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8846,22396,18335,18500,18328,18321,18318,18314,18311]
23974. 2 | 9 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8843,22396,18335,18500,18328,18321,18318,18340,18314]
23976. 2 | 10 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8841,22396,18335,18500,18328,18321,18318,18346,18314]
23978. 2 | 11 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8839,22396,18335,18500,18328,18321,18318,18352,18314]
23980. 2 | 12 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8837,22396,18335,18500,18328,18321,18318,18358,18314]
23982. 2 | 13 | 3 | ~4 | 6 | 19 | 8 | 25 [avatar split clause 8835,22396,18335,18500,18328,18321,18318,18364,18314]
23983. 1 | ~2 | 3 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8822,22396,18335,18500,18328,18370,18318,18314,18311]
23985. 2 | 9 | 3 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8819,22396,18335,18500,18328,18370,18318,18340,18314]
23989. 2 | 11 | 3 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8815,22396,18335,18500,18328,18370,18318,18352,18314]
23993. 2 | 13 | 3 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8811,22396,18335,18500,18328,18370,18318,18364,18314]
23994. 1 | ~2 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8810,22396,18335,18500,18328,18370,18321,18314,18311]
23996. 2 | 9 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8807,22396,18335,18500,18328,18370,18321,18340,18314]
23998. 2 | 10 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8805,22396,18335,18500,18328,18370,18321,18346,18314]
24000. 2 | 11 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8803,22396,18335,18500,18328,18370,18321,18352,18314]
24002. 2 | 12 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8801,22396,18335,18500,18328,18370,18321,18358,18314]
24004. 2 | 13 | 4 | 14 | 6 | 19 | 8 | 25 [avatar split clause 8799,22396,18335,18500,18328,18370,18321,18364,18314]
24005. 1 | ~2 | 3 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8798,22396,18335,18500,18328,18396,18318,18314,18311]
24007. 2 | 9 | 3 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8795,22396,18335,18500,18328,18396,18318,18340,18314]
24015. 2 | 13 | 3 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8787,22396,18335,18500,18328,18396,18318,18364,18314]
24016. 1 | ~2 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8786,22396,18335,18500,18328,18396,18321,18314,18311]
24018. 2 | 9 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8783,22396,18335,18500,18328,18396,18321,18340,18314]
24020. 2 | 10 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8781,22396,18335,18500,18328,18396,18321,18346,18314]
24022. 2 | 11 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8779,22396,18335,18500,18328,18396,18321,18352,18314]
24024. 2 | 12 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8777,22396,18335,18500,18328,18396,18321,18358,18314]
24026. 2 | 13 | 4 | 15 | 6 | 19 | 8 | 25 [avatar split clause 8775,22396,18335,18500,18328,18396,18321,18364,18314]
24027. 1 | ~2 | 3 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8774,22396,18335,18500,18328,18422,18318,18314,18311]
24038. 1 | ~2 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8762,22396,18335,18500,18328,18422,18321,18314,18311]
24040. 2 | 9 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8759,22396,18335,18500,18328,18422,18321,18340,18314]
24042. 2 | 10 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8757,22396,18335,18500,18328,18422,18321,18346,18314]
24044. 2 | 11 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8755,22396,18335,18500,18328,18422,18321,18352,18314]
24046. 2 | 12 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8753,22396,18335,18500,18328,18422,18321,18358,18314]
24048. 2 | 13 | 4 | 16 | 6 | 19 | 8 | 25 [avatar split clause 8751,22396,18335,18500,18328,18422,18321,18364,18314]
24049. 1 | ~2 | 3 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8750,22396,18335,18500,18328,18448,18318,18314,18311]
24051. 2 | 9 | 3 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8747,22396,18335,18500,18328,18448,18318,18340,18314]
24057. 2 | 12 | 3 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8741,22396,18335,18500,18328,18448,18318,18358,18314]
24060. 1 | ~2 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8738,22396,18335,18500,18328,18448,18321,18314,18311]
24062. 2 | 9 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8735,22396,18335,18500,18328,18448,18321,18340,18314]
24064. 2 | 10 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8733,22396,18335,18500,18328,18448,18321,18346,18314]
24066. 2 | 11 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8731,22396,18335,18500,18328,18448,18321,18352,18314]
24068. 2 | 12 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8729,22396,18335,18500,18328,18448,18321,18358,18314]
24070. 2 | 13 | 4 | 17 | 6 | 19 | 8 | 25 [avatar split clause 8727,22396,18335,18500,18328,18448,18321,18364,18314]
24071. 1 | ~2 | 3 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8726,22396,18335,18500,18328,18474,18318,18314,18311]
24075. 2 | 10 | 3 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8721,22396,18335,18500,18328,18474,18318,18346,18314]
24081. 2 | 13 | 3 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8715,22396,18335,18500,18328,18474,18318,18364,18314]
24082. 1 | ~2 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8714,22396,18335,18500,18328,18474,18321,18314,18311]
24084. 2 | 9 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8711,22396,18335,18500,18328,18474,18321,18340,18314]
24086. 2 | 10 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8709,22396,18335,18500,18328,18474,18321,18346,18314]
24088. 2 | 11 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8707,22396,18335,18500,18328,18474,18321,18352,18314]
24090. 2 | 12 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8705,22396,18335,18500,18328,18474,18321,18358,18314]
24092. 2 | 13 | 4 | 18 | 6 | 19 | 8 | 25 [avatar split clause 8703,22396,18335,18500,18328,18474,18321,18364,18314]
24097. 2 | 10 | 3 | ~4 | 5 | 20 | 8 | 25 [avatar split clause 8697,22396,18335,18746,18325,18321,18318,18346,18314]
24099. 2 | 11 | 3 | ~4 | 5 | 20 | 8 | 25 [avatar split clause 8695,22396,18335,18746,18325,18321,18318,18352,18314]
24214. 1 | ~2 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8558,22396,18335,18746,18328,18321,18318,18314,18311]
24216. 2 | 9 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8555,22396,18335,18746,18328,18321,18318,18340,18314]
24218. 2 | 10 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8553,22396,18335,18746,18328,18321,18318,18346,18314]
24220. 2 | 11 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8551,22396,18335,18746,18328,18321,18318,18352,18314]
24222. 2 | 12 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8549,22396,18335,18746,18328,18321,18318,18358,18314]
24224. 2 | 13 | 3 | ~4 | 6 | 20 | 8 | 25 [avatar split clause 8547,22396,18335,18746,18328,18321,18318,18364,18314]
24225. 1 | ~2 | 3 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8534,22396,18335,18746,18328,18370,18318,18314,18311]
24229. 2 | 10 | 3 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8529,22396,18335,18746,18328,18370,18318,18346,18314]
24231. 2 | 11 | 3 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8527,22396,18335,18746,18328,18370,18318,18352,18314]
24233. 2 | 12 | 3 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8525,22396,18335,18746,18328,18370,18318,18358,18314]
24235. 2 | 13 | 3 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8523,22396,18335,18746,18328,18370,18318,18364,18314]
24236. 1 | ~2 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8522,22396,18335,18746,18328,18370,18321,18314,18311]
24238. 2 | 9 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8519,22396,18335,18746,18328,18370,18321,18340,18314]
24240. 2 | 10 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8517,22396,18335,18746,18328,18370,18321,18346,18314]
24242. 2 | 11 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8515,22396,18335,18746,18328,18370,18321,18352,18314]
24244. 2 | 12 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8513,22396,18335,18746,18328,18370,18321,18358,18314]
24246. 2 | 13 | 4 | 14 | 6 | 20 | 8 | 25 [avatar split clause 8511,22396,18335,18746,18328,18370,18321,18364,18314]
24247. 1 | ~2 | 3 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8510,22396,18335,18746,18328,18396,18318,18314,18311]
24249. 2 | 9 | 3 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8507,22396,18335,18746,18328,18396,18318,18340,18314]
24251. 2 | 10 | 3 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8505,22396,18335,18746,18328,18396,18318,18346,18314]
24257. 2 | 13 | 3 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8499,22396,18335,18746,18328,18396,18318,18364,18314]
24258. 1 | ~2 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8498,22396,18335,18746,18328,18396,18321,18314,18311]
24260. 2 | 9 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8495,22396,18335,18746,18328,18396,18321,18340,18314]
24262. 2 | 10 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8493,22396,18335,18746,18328,18396,18321,18346,18314]
24264. 2 | 11 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8491,22396,18335,18746,18328,18396,18321,18352,18314]
24266. 2 | 12 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8489,22396,18335,18746,18328,18396,18321,18358,18314]
24268. 2 | 13 | 4 | 15 | 6 | 20 | 8 | 25 [avatar split clause 8487,22396,18335,18746,18328,18396,18321,18364,18314]
24269. 1 | ~2 | 3 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8486,22396,18335,18746,18328,18422,18318,18314,18311]
24280. 1 | ~2 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8474,22396,18335,18746,18328,18422,18321,18314,18311]
24282. 2 | 9 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8471,22396,18335,18746,18328,18422,18321,18340,18314]
24284. 2 | 10 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8469,22396,18335,18746,18328,18422,18321,18346,18314]
24286. 2 | 11 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8467,22396,18335,18746,18328,18422,18321,18352,18314]
24288. 2 | 12 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8465,22396,18335,18746,18328,18422,18321,18358,18314]
24290. 2 | 13 | 4 | 16 | 6 | 20 | 8 | 25 [avatar split clause 8463,22396,18335,18746,18328,18422,18321,18364,18314]
24295. 2 | 10 | 3 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8457,22396,18335,18746,18328,18448,18318,18346,18314]
24297. 2 | 11 | 3 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8455,22396,18335,18746,18328,18448,18318,18352,18314]
24299. 2 | 12 | 3 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8453,22396,18335,18746,18328,18448,18318,18358,18314]
24302. 1 | ~2 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8450,22396,18335,18746,18328,18448,18321,18314,18311]
24304. 2 | 9 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8447,22396,18335,18746,18328,18448,18321,18340,18314]
24306. 2 | 10 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8445,22396,18335,18746,18328,18448,18321,18346,18314]
24308. 2 | 11 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8443,22396,18335,18746,18328,18448,18321,18352,18314]
24310. 2 | 12 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8441,22396,18335,18746,18328,18448,18321,18358,18314]
24312. 2 | 13 | 4 | 17 | 6 | 20 | 8 | 25 [avatar split clause 8439,22396,18335,18746,18328,18448,18321,18364,18314]
24313. 1 | ~2 | 3 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8438,22396,18335,18746,18328,18474,18318,18314,18311]
24317. 2 | 10 | 3 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8433,22396,18335,18746,18328,18474,18318,18346,18314]
24319. 2 | 11 | 3 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8431,22396,18335,18746,18328,18474,18318,18352,18314]
24324. 1 | ~2 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8426,22396,18335,18746,18328,18474,18321,18314,18311]
24326. 2 | 9 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8423,22396,18335,18746,18328,18474,18321,18340,18314]
24328. 2 | 10 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8421,22396,18335,18746,18328,18474,18321,18346,18314]
24330. 2 | 11 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8419,22396,18335,18746,18328,18474,18321,18352,18314]
24332. 2 | 12 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8417,22396,18335,18746,18328,18474,18321,18358,18314]
24334. 2 | 13 | 4 | 18 | 6 | 20 | 8 | 25 [avatar split clause 8415,22396,18335,18746,18328,18474,18321,18364,18314]
24339. 2 | 10 | 3 | ~4 | 5 | 21 | 8 | 25 [avatar split clause 8409,22396,18335,18992,18325,18321,18318,18346,18314]
24341. 2 | 11 | 3 | ~4 | 5 | 21 | 8 | 25 [avatar split clause 8407,22396,18335,18992,18325,18321,18318,18352,18314]
24456. 1 | ~2 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8270,22396,18335,18992,18328,18321,18318,18314,18311]
24458. 2 | 9 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8267,22396,18335,18992,18328,18321,18318,18340,18314]
24460. 2 | 10 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8265,22396,18335,18992,18328,18321,18318,18346,18314]
24462. 2 | 11 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8263,22396,18335,18992,18328,18321,18318,18352,18314]
24464. 2 | 12 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8261,22396,18335,18992,18328,18321,18318,18358,18314]
24466. 2 | 13 | 3 | ~4 | 6 | 21 | 8 | 25 [avatar split clause 8259,22396,18335,18992,18328,18321,18318,18364,18314]
24467. 1 | ~2 | 3 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8246,22396,18335,18992,18328,18370,18318,18314,18311]
24469. 2 | 9 | 3 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8243,22396,18335,18992,18328,18370,18318,18340,18314]
24471. 2 | 10 | 3 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8241,22396,18335,18992,18328,18370,18318,18346,18314]
24473. 2 | 11 | 3 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8239,22396,18335,18992,18328,18370,18318,18352,18314]
24475. 2 | 12 | 3 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8237,22396,18335,18992,18328,18370,18318,18358,18314]
24478. 1 | ~2 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8234,22396,18335,18992,18328,18370,18321,18314,18311]
24480. 2 | 9 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8231,22396,18335,18992,18328,18370,18321,18340,18314]
24482. 2 | 10 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8229,22396,18335,18992,18328,18370,18321,18346,18314]
24484. 2 | 11 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8227,22396,18335,18992,18328,18370,18321,18352,18314]
24486. 2 | 12 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8225,22396,18335,18992,18328,18370,18321,18358,18314]
24488. 2 | 13 | 4 | 14 | 6 | 21 | 8 | 25 [avatar split clause 8223,22396,18335,18992,18328,18370,18321,18364,18314]
24489. 1 | ~2 | 3 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8222,22396,18335,18992,18328,18396,18318,18314,18311]
24491. 2 | 9 | 3 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8219,22396,18335,18992,18328,18396,18318,18340,18314]
24495. 2 | 11 | 3 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8215,22396,18335,18992,18328,18396,18318,18352,18314]
24497. 2 | 12 | 3 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8213,22396,18335,18992,18328,18396,18318,18358,18314]
24499. 2 | 13 | 3 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8211,22396,18335,18992,18328,18396,18318,18364,18314]
24500. 1 | ~2 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8210,22396,18335,18992,18328,18396,18321,18314,18311]
24502. 2 | 9 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8207,22396,18335,18992,18328,18396,18321,18340,18314]
24504. 2 | 10 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8205,22396,18335,18992,18328,18396,18321,18346,18314]
24506. 2 | 11 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8203,22396,18335,18992,18328,18396,18321,18352,18314]
24508. 2 | 12 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8201,22396,18335,18992,18328,18396,18321,18358,18314]
24510. 2 | 13 | 4 | 15 | 6 | 21 | 8 | 25 [avatar split clause 8199,22396,18335,18992,18328,18396,18321,18364,18314]
24522. 1 | ~2 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8186,22396,18335,18992,18328,18422,18321,18314,18311]
24524. 2 | 9 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8183,22396,18335,18992,18328,18422,18321,18340,18314]
24526. 2 | 10 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8181,22396,18335,18992,18328,18422,18321,18346,18314]
24528. 2 | 11 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8179,22396,18335,18992,18328,18422,18321,18352,18314]
24530. 2 | 12 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8177,22396,18335,18992,18328,18422,18321,18358,18314]
24532. 2 | 13 | 4 | 16 | 6 | 21 | 8 | 25 [avatar split clause 8175,22396,18335,18992,18328,18422,18321,18364,18314]
24533. 1 | ~2 | 3 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8174,22396,18335,18992,18328,18448,18318,18314,18311]
24535. 2 | 9 | 3 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8171,22396,18335,18992,18328,18448,18318,18340,18314]
24541. 2 | 12 | 3 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8165,22396,18335,18992,18328,18448,18318,18358,18314]
24544. 1 | ~2 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8162,22396,18335,18992,18328,18448,18321,18314,18311]
24546. 2 | 9 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8159,22396,18335,18992,18328,18448,18321,18340,18314]
24548. 2 | 10 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8157,22396,18335,18992,18328,18448,18321,18346,18314]
24550. 2 | 11 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8155,22396,18335,18992,18328,18448,18321,18352,18314]
24552. 2 | 12 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8153,22396,18335,18992,18328,18448,18321,18358,18314]
24554. 2 | 13 | 4 | 17 | 6 | 21 | 8 | 25 [avatar split clause 8151,22396,18335,18992,18328,18448,18321,18364,18314]
24555. 1 | ~2 | 3 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8150,22396,18335,18992,18328,18474,18318,18314,18311]
24557. 2 | 9 | 3 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8147,22396,18335,18992,18328,18474,18318,18340,18314]
24561. 2 | 11 | 3 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8143,22396,18335,18992,18328,18474,18318,18352,18314]
24563. 2 | 12 | 3 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8141,22396,18335,18992,18328,18474,18318,18358,18314]
24566. 1 | ~2 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8138,22396,18335,18992,18328,18474,18321,18314,18311]
24568. 2 | 9 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8135,22396,18335,18992,18328,18474,18321,18340,18314]
24570. 2 | 10 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8133,22396,18335,18992,18328,18474,18321,18346,18314]
24572. 2 | 11 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8131,22396,18335,18992,18328,18474,18321,18352,18314]
24574. 2 | 12 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8129,22396,18335,18992,18328,18474,18321,18358,18314]
24576. 2 | 13 | 4 | 18 | 6 | 21 | 8 | 25 [avatar split clause 8127,22396,18335,18992,18328,18474,18321,18364,18314]
24579. 2 | 9 | 3 | ~4 | 5 | 22 | 8 | 25 [avatar split clause 8123,22396,18335,19238,18325,18321,18318,18340,18314]
24581. 2 | 10 | 3 | ~4 | 5 | 22 | 8 | 25 [avatar split clause 8121,22396,18335,19238,18325,18321,18318,18346,18314]
24583. 2 | 11 | 3 | ~4 | 5 | 22 | 8 | 25 [avatar split clause 8119,22396,18335,19238,18325,18321,18318,18352,18314]
24585. 2 | 12 | 3 | ~4 | 5 | 22 | 8 | 25 [avatar split clause 8117,22396,18335,19238,18325,18321,18318,18358,18314]
24698. 1 | ~2 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7982,22396,18335,19238,18328,18321,18318,18314,18311]
24700. 2 | 9 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7979,22396,18335,19238,18328,18321,18318,18340,18314]
24702. 2 | 10 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7977,22396,18335,19238,18328,18321,18318,18346,18314]
24704. 2 | 11 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7975,22396,18335,19238,18328,18321,18318,18352,18314]
24706. 2 | 12 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7973,22396,18335,19238,18328,18321,18318,18358,18314]
24708. 2 | 13 | 3 | ~4 | 6 | 22 | 8 | 25 [avatar split clause 7971,22396,18335,19238,18328,18321,18318,18364,18314]
24713. 2 | 10 | 3 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7953,22396,18335,19238,18328,18370,18318,18346,18314]
24715. 2 | 11 | 3 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7951,22396,18335,19238,18328,18370,18318,18352,18314]
24717. 2 | 12 | 3 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7949,22396,18335,19238,18328,18370,18318,18358,18314]
24720. 1 | ~2 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7946,22396,18335,19238,18328,18370,18321,18314,18311]
24722. 2 | 9 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7943,22396,18335,19238,18328,18370,18321,18340,18314]
24724. 2 | 10 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7941,22396,18335,19238,18328,18370,18321,18346,18314]
24726. 2 | 11 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7939,22396,18335,19238,18328,18370,18321,18352,18314]
24728. 2 | 12 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7937,22396,18335,19238,18328,18370,18321,18358,18314]
24730. 2 | 13 | 4 | 14 | 6 | 22 | 8 | 25 [avatar split clause 7935,22396,18335,19238,18328,18370,18321,18364,18314]
24733. 2 | 9 | 3 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7931,22396,18335,19238,18328,18396,18318,18340,18314]
24735. 2 | 10 | 3 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7929,22396,18335,19238,18328,18396,18318,18346,18314]
24737. 2 | 11 | 3 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7927,22396,18335,19238,18328,18396,18318,18352,18314]
24739. 2 | 12 | 3 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7925,22396,18335,19238,18328,18396,18318,18358,18314]
24742. 1 | ~2 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7922,22396,18335,19238,18328,18396,18321,18314,18311]
24744. 2 | 9 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7919,22396,18335,19238,18328,18396,18321,18340,18314]
24746. 2 | 10 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7917,22396,18335,19238,18328,18396,18321,18346,18314]
24748. 2 | 11 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7915,22396,18335,19238,18328,18396,18321,18352,18314]
24750. 2 | 12 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7913,22396,18335,19238,18328,18396,18321,18358,18314]
24752. 2 | 13 | 4 | 15 | 6 | 22 | 8 | 25 [avatar split clause 7911,22396,18335,19238,18328,18396,18321,18364,18314]
24764. 1 | ~2 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7898,22396,18335,19238,18328,18422,18321,18314,18311]
24766. 2 | 9 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7895,22396,18335,19238,18328,18422,18321,18340,18314]
24768. 2 | 10 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7893,22396,18335,19238,18328,18422,18321,18346,18314]
24770. 2 | 11 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7891,22396,18335,19238,18328,18422,18321,18352,18314]
24772. 2 | 12 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7889,22396,18335,19238,18328,18422,18321,18358,18314]
24774. 2 | 13 | 4 | 16 | 6 | 22 | 8 | 25 [avatar split clause 7887,22396,18335,19238,18328,18422,18321,18364,18314]
24777. 2 | 9 | 3 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7883,22396,18335,19238,18328,18448,18318,18340,18314]
24779. 2 | 10 | 3 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7881,22396,18335,19238,18328,18448,18318,18346,18314]
24781. 2 | 11 | 3 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7879,22396,18335,19238,18328,18448,18318,18352,18314]
24783. 2 | 12 | 3 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7877,22396,18335,19238,18328,18448,18318,18358,18314]
24785. 2 | 13 | 3 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7875,22396,18335,19238,18328,18448,18318,18364,18314]
24786. 1 | ~2 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7874,22396,18335,19238,18328,18448,18321,18314,18311]
24788. 2 | 9 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7871,22396,18335,19238,18328,18448,18321,18340,18314]
24790. 2 | 10 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7869,22396,18335,19238,18328,18448,18321,18346,18314]
24792. 2 | 11 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7867,22396,18335,19238,18328,18448,18321,18352,18314]
24794. 2 | 12 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7865,22396,18335,19238,18328,18448,18321,18358,18314]
24796. 2 | 13 | 4 | 17 | 6 | 22 | 8 | 25 [avatar split clause 7863,22396,18335,19238,18328,18448,18321,18364,18314]
24797. 1 | ~2 | 3 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7862,22396,18335,19238,18328,18474,18318,18314,18311]
24808. 1 | ~2 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7850,22396,18335,19238,18328,18474,18321,18314,18311]
24810. 2 | 9 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7847,22396,18335,19238,18328,18474,18321,18340,18314]
24812. 2 | 10 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7845,22396,18335,19238,18328,18474,18321,18346,18314]
24814. 2 | 11 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7843,22396,18335,19238,18328,18474,18321,18352,18314]
24816. 2 | 12 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7841,22396,18335,19238,18328,18474,18321,18358,18314]
24818. 2 | 13 | 4 | 18 | 6 | 22 | 8 | 25 [avatar split clause 7839,22396,18335,19238,18328,18474,18321,18364,18314]
24821. 2 | 9 | 3 | ~4 | 5 | 23 | 8 | 25 [avatar split clause 7835,22396,18335,19484,18325,18321,18318,18340,18314]
24823. 2 | 10 | 3 | ~4 | 5 | 23 | 8 | 25 [avatar split clause 7833,22396,18335,19484,18325,18321,18318,18346,18314]
24825. 2 | 11 | 3 | ~4 | 5 | 23 | 8 | 25 [avatar split clause 7831,22396,18335,19484,18325,18321,18318,18352,18314]
24940. 1 | ~2 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7694,22396,18335,19484,18328,18321,18318,18314,18311]
24942. 2 | 9 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7691,22396,18335,19484,18328,18321,18318,18340,18314]
24944. 2 | 10 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7689,22396,18335,19484,18328,18321,18318,18346,18314]
24946. 2 | 11 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7687,22396,18335,19484,18328,18321,18318,18352,18314]
24948. 2 | 12 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7685,22396,18335,19484,18328,18321,18318,18358,18314]
24950. 2 | 13 | 3 | ~4 | 6 | 23 | 8 | 25 [avatar split clause 7683,22396,18335,19484,18328,18321,18318,18364,18314]
24951. 1 | ~2 | 3 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7670,22396,18335,19484,18328,18370,18318,18314,18311]
24961. 2 | 13 | 3 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7659,22396,18335,19484,18328,18370,18318,18364,18314]
24962. 1 | ~2 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7658,22396,18335,19484,18328,18370,18321,18314,18311]
24964. 2 | 9 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7655,22396,18335,19484,18328,18370,18321,18340,18314]
24966. 2 | 10 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7653,22396,18335,19484,18328,18370,18321,18346,18314]
24968. 2 | 11 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7651,22396,18335,19484,18328,18370,18321,18352,18314]
24970. 2 | 12 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7649,22396,18335,19484,18328,18370,18321,18358,18314]
24972. 2 | 13 | 4 | 14 | 6 | 23 | 8 | 25 [avatar split clause 7647,22396,18335,19484,18328,18370,18321,18364,18314]
24973. 1 | ~2 | 3 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7646,22396,18335,19484,18328,18396,18318,18314,18311]
24975. 2 | 9 | 3 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7643,22396,18335,19484,18328,18396,18318,18340,18314]
24984. 1 | ~2 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7634,22396,18335,19484,18328,18396,18321,18314,18311]
24986. 2 | 9 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7631,22396,18335,19484,18328,18396,18321,18340,18314]
24988. 2 | 10 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7629,22396,18335,19484,18328,18396,18321,18346,18314]
24990. 2 | 11 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7627,22396,18335,19484,18328,18396,18321,18352,18314]
24992. 2 | 12 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7625,22396,18335,19484,18328,18396,18321,18358,18314]
24994. 2 | 13 | 4 | 15 | 6 | 23 | 8 | 25 [avatar split clause 7623,22396,18335,19484,18328,18396,18321,18364,18314]
24995. 1 | ~2 | 3 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7622,22396,18335,19484,18328,18422,18318,18314,18311]
25006. 1 | ~2 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7610,22396,18335,19484,18328,18422,18321,18314,18311]
25008. 2 | 9 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7607,22396,18335,19484,18328,18422,18321,18340,18314]
25010. 2 | 10 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7605,22396,18335,19484,18328,18422,18321,18346,18314]
25012. 2 | 11 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7603,22396,18335,19484,18328,18422,18321,18352,18314]
25014. 2 | 12 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7601,22396,18335,19484,18328,18422,18321,18358,18314]
25016. 2 | 13 | 4 | 16 | 6 | 23 | 8 | 25 [avatar split clause 7599,22396,18335,19484,18328,18422,18321,18364,18314]
25017. 1 | ~2 | 3 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7598,22396,18335,19484,18328,18448,18318,18314,18311]
25023. 2 | 11 | 3 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7591,22396,18335,19484,18328,18448,18318,18352,18314]
25025. 2 | 12 | 3 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7589,22396,18335,19484,18328,18448,18318,18358,18314]
25028. 1 | ~2 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7586,22396,18335,19484,18328,18448,18321,18314,18311]
25030. 2 | 9 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7583,22396,18335,19484,18328,18448,18321,18340,18314]
25032. 2 | 10 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7581,22396,18335,19484,18328,18448,18321,18346,18314]
25034. 2 | 11 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7579,22396,18335,19484,18328,18448,18321,18352,18314]
25036. 2 | 12 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7577,22396,18335,19484,18328,18448,18321,18358,18314]
25038. 2 | 13 | 4 | 17 | 6 | 23 | 8 | 25 [avatar split clause 7575,22396,18335,19484,18328,18448,18321,18364,18314]
25039. 1 | ~2 | 3 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7574,22396,18335,19484,18328,18474,18318,18314,18311]
25045. 2 | 11 | 3 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7567,22396,18335,19484,18328,18474,18318,18352,18314]
25047. 2 | 12 | 3 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7565,22396,18335,19484,18328,18474,18318,18358,18314]
25050. 1 | ~2 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7562,22396,18335,19484,18328,18474,18321,18314,18311]
25052. 2 | 9 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7559,22396,18335,19484,18328,18474,18321,18340,18314]
25054. 2 | 10 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7557,22396,18335,19484,18328,18474,18321,18346,18314]
25056. 2 | 11 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7555,22396,18335,19484,18328,18474,18321,18352,18314]
25058. 2 | 12 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7553,22396,18335,19484,18328,18474,18321,18358,18314]
25060. 2 | 13 | 4 | 18 | 6 | 23 | 8 | 25 [avatar split clause 7551,22396,18335,19484,18328,18474,18321,18364,18314]
25062. 26 <=> v1_zf_lang(sK2) [avatar definition]
26396. 1 | ~2 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5822,25062,18335,18328,18325,18321,18318,18314,18311]
26398. 2 | 9 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5819,25062,18335,18328,18325,18321,18318,18340,18314]
26400. 2 | 10 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5817,25062,18335,18328,18325,18321,18318,18346,18314]
26402. 2 | 11 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5815,25062,18335,18328,18325,18321,18318,18352,18314]
26404. 2 | 12 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5813,25062,18335,18328,18325,18321,18318,18358,18314]
26406. 2 | 13 | 3 | ~4 | 5 | ~6 | 8 | 26 [avatar split clause 5811,25062,18335,18328,18325,18321,18318,18364,18314]
26418. 1 | ~2 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5786,25062,18335,18328,18325,18370,18321,18314,18311]
26420. 2 | 9 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5783,25062,18335,18328,18325,18370,18321,18340,18314]
26422. 2 | 10 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5781,25062,18335,18328,18325,18370,18321,18346,18314]
26424. 2 | 11 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5779,25062,18335,18328,18325,18370,18321,18352,18314]
26425. 1 | 12 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5778,25062,18335,18328,18325,18370,18321,18358,18311]
26426. 2 | 12 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5777,25062,18335,18328,18325,18370,18321,18358,18314]
26428. 2 | 13 | 4 | 14 | 5 | ~6 | 8 | 26 [avatar split clause 5775,25062,18335,18328,18325,18370,18321,18364,18314]
26440. 1 | ~2 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5762,25062,18335,18328,18325,18396,18321,18314,18311]
26442. 2 | 9 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5759,25062,18335,18328,18325,18396,18321,18340,18314]
26443. 1 | 10 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5758,25062,18335,18328,18325,18396,18321,18346,18311]
26444. 2 | 10 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5757,25062,18335,18328,18325,18396,18321,18346,18314]
26446. 2 | 11 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5755,25062,18335,18328,18325,18396,18321,18352,18314]
26448. 2 | 12 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5753,25062,18335,18328,18325,18396,18321,18358,18314]
26450. 2 | 13 | 4 | 15 | 5 | ~6 | 8 | 26 [avatar split clause 5751,25062,18335,18328,18325,18396,18321,18364,18314]
26462. 1 | ~2 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5738,25062,18335,18328,18325,18422,18321,18314,18311]
26464. 2 | 9 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5735,25062,18335,18328,18325,18422,18321,18340,18314]
26465. 1 | 10 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5734,25062,18335,18328,18325,18422,18321,18346,18311]
26466. 2 | 10 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5733,25062,18335,18328,18325,18422,18321,18346,18314]
26468. 2 | 11 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5731,25062,18335,18328,18325,18422,18321,18352,18314]
26469. 1 | 12 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5730,25062,18335,18328,18325,18422,18321,18358,18311]
26470. 2 | 12 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5729,25062,18335,18328,18325,18422,18321,18358,18314]
26472. 2 | 13 | 4 | 16 | 5 | ~6 | 8 | 26 [avatar split clause 5727,25062,18335,18328,18325,18422,18321,18364,18314]
26484. 1 | ~2 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5714,25062,18335,18328,18325,18448,18321,18314,18311]
26485. 1 | 9 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5712,25062,18335,18328,18325,18448,18321,18340,18311]
26486. 2 | 9 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5711,25062,18335,18328,18325,18448,18321,18340,18314]
26488. 2 | 10 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5709,25062,18335,18328,18325,18448,18321,18346,18314]
26489. 1 | 11 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5708,25062,18335,18328,18325,18448,18321,18352,18311]
26490. 2 | 11 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5707,25062,18335,18328,18325,18448,18321,18352,18314]
26492. 2 | 12 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5705,25062,18335,18328,18325,18448,18321,18358,18314]
26493. 1 | 13 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5704,25062,18335,18328,18325,18448,18321,18364,18311]
26494. 2 | 13 | 4 | 17 | 5 | ~6 | 8 | 26 [avatar split clause 5703,25062,18335,18328,18325,18448,18321,18364,18314]
26506. 1 | ~2 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5690,25062,18335,18328,18325,18474,18321,18314,18311]
26508. 2 | 9 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5687,25062,18335,18328,18325,18474,18321,18340,18314]
26510. 2 | 10 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5685,25062,18335,18328,18325,18474,18321,18346,18314]
26512. 2 | 11 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5683,25062,18335,18328,18325,18474,18321,18352,18314]
26513. 1 | 12 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5682,25062,18335,18328,18325,18474,18321,18358,18311]
26514. 2 | 12 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5681,25062,18335,18328,18325,18474,18321,18358,18314]
26515. 1 | 13 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5680,25062,18335,18328,18325,18474,18321,18364,18311]
26516. 2 | 13 | 4 | 18 | 5 | ~6 | 8 | 26 [avatar split clause 5679,25062,18335,18328,18325,18474,18321,18364,18314]
26521. 2 | 10 | 3 | ~4 | 5 | 19 | 8 | 26 [avatar split clause 5529,25062,18335,18500,18325,18321,18318,18346,18314]
26523. 2 | 11 | 3 | ~4 | 5 | 19 | 8 | 26 [avatar split clause 5527,25062,18335,18500,18325,18321,18318,18352,18314]
26638. 1 | ~2 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5390,25062,18335,18500,18328,18321,18318,18314,18311]
26640. 2 | 9 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5387,25062,18335,18500,18328,18321,18318,18340,18314]
26642. 2 | 10 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5385,25062,18335,18500,18328,18321,18318,18346,18314]
26644. 2 | 11 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5383,25062,18335,18500,18328,18321,18318,18352,18314]
26646. 2 | 12 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5381,25062,18335,18500,18328,18321,18318,18358,18314]
26648. 2 | 13 | 3 | ~4 | 6 | 19 | 8 | 26 [avatar split clause 5379,25062,18335,18500,18328,18321,18318,18364,18314]
26649. 1 | ~2 | 3 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5366,25062,18335,18500,18328,18370,18318,18314,18311]
26653. 2 | 10 | 3 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5361,25062,18335,18500,18328,18370,18318,18346,18314]
26659. 2 | 13 | 3 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5355,25062,18335,18500,18328,18370,18318,18364,18314]
26660. 1 | ~2 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5354,25062,18335,18500,18328,18370,18321,18314,18311]
26662. 2 | 9 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5351,25062,18335,18500,18328,18370,18321,18340,18314]
26664. 2 | 10 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5349,25062,18335,18500,18328,18370,18321,18346,18314]
26666. 2 | 11 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5347,25062,18335,18500,18328,18370,18321,18352,18314]
26668. 2 | 12 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5345,25062,18335,18500,18328,18370,18321,18358,18314]
26670. 2 | 13 | 4 | 14 | 6 | 19 | 8 | 26 [avatar split clause 5343,25062,18335,18500,18328,18370,18321,18364,18314]
26675. 2 | 10 | 3 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5337,25062,18335,18500,18328,18396,18318,18346,18314]
26677. 2 | 11 | 3 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5335,25062,18335,18500,18328,18396,18318,18352,18314]
26681. 2 | 13 | 3 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5331,25062,18335,18500,18328,18396,18318,18364,18314]
26682. 1 | ~2 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5330,25062,18335,18500,18328,18396,18321,18314,18311]
26684. 2 | 9 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5327,25062,18335,18500,18328,18396,18321,18340,18314]
26686. 2 | 10 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5325,25062,18335,18500,18328,18396,18321,18346,18314]
26688. 2 | 11 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5323,25062,18335,18500,18328,18396,18321,18352,18314]
26690. 2 | 12 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5321,25062,18335,18500,18328,18396,18321,18358,18314]
26692. 2 | 13 | 4 | 15 | 6 | 19 | 8 | 26 [avatar split clause 5319,25062,18335,18500,18328,18396,18321,18364,18314]
26693. 1 | ~2 | 3 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5318,25062,18335,18500,18328,18422,18318,18314,18311]
26695. 2 | 9 | 3 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5315,25062,18335,18500,18328,18422,18318,18340,18314]
26697. 2 | 10 | 3 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5313,25062,18335,18500,18328,18422,18318,18346,18314]
26699. 2 | 11 | 3 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5311,25062,18335,18500,18328,18422,18318,18352,18314]
26704. 1 | ~2 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5306,25062,18335,18500,18328,18422,18321,18314,18311]
26706. 2 | 9 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5303,25062,18335,18500,18328,18422,18321,18340,18314]
26708. 2 | 10 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5301,25062,18335,18500,18328,18422,18321,18346,18314]
26710. 2 | 11 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5299,25062,18335,18500,18328,18422,18321,18352,18314]
26712. 2 | 12 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5297,25062,18335,18500,18328,18422,18321,18358,18314]
26714. 2 | 13 | 4 | 16 | 6 | 19 | 8 | 26 [avatar split clause 5295,25062,18335,18500,18328,18422,18321,18364,18314]
26717. 2 | 9 | 3 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5291,25062,18335,18500,18328,18448,18318,18340,18314]
26719. 2 | 10 | 3 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5289,25062,18335,18500,18328,18448,18318,18346,18314]
26721. 2 | 11 | 3 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5287,25062,18335,18500,18328,18448,18318,18352,18314]
26726. 1 | ~2 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5282,25062,18335,18500,18328,18448,18321,18314,18311]
26728. 2 | 9 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5279,25062,18335,18500,18328,18448,18321,18340,18314]
26730. 2 | 10 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5277,25062,18335,18500,18328,18448,18321,18346,18314]
26732. 2 | 11 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5275,25062,18335,18500,18328,18448,18321,18352,18314]
26734. 2 | 12 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5273,25062,18335,18500,18328,18448,18321,18358,18314]
26736. 2 | 13 | 4 | 17 | 6 | 19 | 8 | 26 [avatar split clause 5271,25062,18335,18500,18328,18448,18321,18364,18314]
26737. 1 | ~2 | 3 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5270,25062,18335,18500,18328,18474,18318,18314,18311]
26739. 2 | 9 | 3 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5267,25062,18335,18500,18328,18474,18318,18340,18314]
26743. 2 | 11 | 3 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5263,25062,18335,18500,18328,18474,18318,18352,18314]
26747. 2 | 13 | 3 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5259,25062,18335,18500,18328,18474,18318,18364,18314]
26748. 1 | ~2 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5258,25062,18335,18500,18328,18474,18321,18314,18311]
26750. 2 | 9 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5255,25062,18335,18500,18328,18474,18321,18340,18314]
26752. 2 | 10 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5253,25062,18335,18500,18328,18474,18321,18346,18314]
26754. 2 | 11 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5251,25062,18335,18500,18328,18474,18321,18352,18314]
26756. 2 | 12 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5249,25062,18335,18500,18328,18474,18321,18358,18314]
26758. 2 | 13 | 4 | 18 | 6 | 19 | 8 | 26 [avatar split clause 5247,25062,18335,18500,18328,18474,18321,18364,18314]
26763. 2 | 10 | 3 | ~4 | 5 | 20 | 8 | 26 [avatar split clause 5241,25062,18335,18746,18325,18321,18318,18346,18314]
26765. 2 | 11 | 3 | ~4 | 5 | 20 | 8 | 26 [avatar split clause 5239,25062,18335,18746,18325,18321,18318,18352,18314]
26880. 1 | ~2 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5102,25062,18335,18746,18328,18321,18318,18314,18311]
26882. 2 | 9 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5099,25062,18335,18746,18328,18321,18318,18340,18314]
26884. 2 | 10 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5097,25062,18335,18746,18328,18321,18318,18346,18314]
26886. 2 | 11 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5095,25062,18335,18746,18328,18321,18318,18352,18314]
26888. 2 | 12 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5093,25062,18335,18746,18328,18321,18318,18358,18314]
26890. 2 | 13 | 3 | ~4 | 6 | 20 | 8 | 26 [avatar split clause 5091,25062,18335,18746,18328,18321,18318,18364,18314]
26891. 1 | ~2 | 3 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5078,25062,18335,18746,18328,18370,18318,18314,18311]
26899. 2 | 12 | 3 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5069,25062,18335,18746,18328,18370,18318,18358,18314]
26902. 1 | ~2 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5066,25062,18335,18746,18328,18370,18321,18314,18311]
26904. 2 | 9 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5063,25062,18335,18746,18328,18370,18321,18340,18314]
26906. 2 | 10 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5061,25062,18335,18746,18328,18370,18321,18346,18314]
26908. 2 | 11 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5059,25062,18335,18746,18328,18370,18321,18352,18314]
26910. 2 | 12 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5057,25062,18335,18746,18328,18370,18321,18358,18314]
26912. 2 | 13 | 4 | 14 | 6 | 20 | 8 | 26 [avatar split clause 5055,25062,18335,18746,18328,18370,18321,18364,18314]
26913. 1 | ~2 | 3 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5054,25062,18335,18746,18328,18396,18318,18314,18311]
26919. 2 | 11 | 3 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5047,25062,18335,18746,18328,18396,18318,18352,18314]
26921. 2 | 12 | 3 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5045,25062,18335,18746,18328,18396,18318,18358,18314]
26924. 1 | ~2 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5042,25062,18335,18746,18328,18396,18321,18314,18311]
26926. 2 | 9 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5039,25062,18335,18746,18328,18396,18321,18340,18314]
26928. 2 | 10 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5037,25062,18335,18746,18328,18396,18321,18346,18314]
26930. 2 | 11 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5035,25062,18335,18746,18328,18396,18321,18352,18314]
26932. 2 | 12 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5033,25062,18335,18746,18328,18396,18321,18358,18314]
26934. 2 | 13 | 4 | 15 | 6 | 20 | 8 | 26 [avatar split clause 5031,25062,18335,18746,18328,18396,18321,18364,18314]
26935. 1 | ~2 | 3 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5030,25062,18335,18746,18328,18422,18318,18314,18311]
26941. 2 | 11 | 3 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5023,25062,18335,18746,18328,18422,18318,18352,18314]
26946. 1 | ~2 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5018,25062,18335,18746,18328,18422,18321,18314,18311]
26948. 2 | 9 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5015,25062,18335,18746,18328,18422,18321,18340,18314]
26950. 2 | 10 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5013,25062,18335,18746,18328,18422,18321,18346,18314]
26952. 2 | 11 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5011,25062,18335,18746,18328,18422,18321,18352,18314]
26954. 2 | 12 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5009,25062,18335,18746,18328,18422,18321,18358,18314]
26956. 2 | 13 | 4 | 16 | 6 | 20 | 8 | 26 [avatar split clause 5007,25062,18335,18746,18328,18422,18321,18364,18314]
26957. 1 | ~2 | 3 | 17 | 6 | 20 | 8 | 26 [avatar split clause 5006,25062,18335,18746,18328,18448,18318,18314,18311]
26967. 2 | 13 | 3 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4995,25062,18335,18746,18328,18448,18318,18364,18314]
26968. 1 | ~2 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4994,25062,18335,18746,18328,18448,18321,18314,18311]
26970. 2 | 9 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4991,25062,18335,18746,18328,18448,18321,18340,18314]
26972. 2 | 10 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4989,25062,18335,18746,18328,18448,18321,18346,18314]
26974. 2 | 11 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4987,25062,18335,18746,18328,18448,18321,18352,18314]
26976. 2 | 12 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4985,25062,18335,18746,18328,18448,18321,18358,18314]
26978. 2 | 13 | 4 | 17 | 6 | 20 | 8 | 26 [avatar split clause 4983,25062,18335,18746,18328,18448,18321,18364,18314]
26979. 1 | ~2 | 3 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4982,25062,18335,18746,18328,18474,18318,18314,18311]
26981. 2 | 9 | 3 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4979,25062,18335,18746,18328,18474,18318,18340,18314]
26983. 2 | 10 | 3 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4977,25062,18335,18746,18328,18474,18318,18346,18314]
26990. 1 | ~2 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4970,25062,18335,18746,18328,18474,18321,18314,18311]
26992. 2 | 9 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4967,25062,18335,18746,18328,18474,18321,18340,18314]
26994. 2 | 10 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4965,25062,18335,18746,18328,18474,18321,18346,18314]
26996. 2 | 11 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4963,25062,18335,18746,18328,18474,18321,18352,18314]
26998. 2 | 12 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4961,25062,18335,18746,18328,18474,18321,18358,18314]
27000. 2 | 13 | 4 | 18 | 6 | 20 | 8 | 26 [avatar split clause 4959,25062,18335,18746,18328,18474,18321,18364,18314]
27007. 2 | 11 | 3 | ~4 | 5 | 21 | 8 | 26 [avatar split clause 4951,25062,18335,18992,18325,18321,18318,18352,18314]
27122. 1 | ~2 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4814,25062,18335,18992,18328,18321,18318,18314,18311]
27124. 2 | 9 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4811,25062,18335,18992,18328,18321,18318,18340,18314]
27126. 2 | 10 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4809,25062,18335,18992,18328,18321,18318,18346,18314]
27128. 2 | 11 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4807,25062,18335,18992,18328,18321,18318,18352,18314]
27130. 2 | 12 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4805,25062,18335,18992,18328,18321,18318,18358,18314]
27132. 2 | 13 | 3 | ~4 | 6 | 21 | 8 | 26 [avatar split clause 4803,25062,18335,18992,18328,18321,18318,18364,18314]
27133. 1 | ~2 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4790,25062,18335,18992,18328,18370,18318,18314,18311]
27135. 2 | 9 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4787,25062,18335,18992,18328,18370,18318,18340,18314]
27137. 2 | 10 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4785,25062,18335,18992,18328,18370,18318,18346,18314]
27139. 2 | 11 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4783,25062,18335,18992,18328,18370,18318,18352,18314]
27141. 2 | 12 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4781,25062,18335,18992,18328,18370,18318,18358,18314]
27143. 2 | 13 | 3 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4779,25062,18335,18992,18328,18370,18318,18364,18314]
27144. 1 | ~2 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4778,25062,18335,18992,18328,18370,18321,18314,18311]
27146. 2 | 9 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4775,25062,18335,18992,18328,18370,18321,18340,18314]
27148. 2 | 10 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4773,25062,18335,18992,18328,18370,18321,18346,18314]
27150. 2 | 11 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4771,25062,18335,18992,18328,18370,18321,18352,18314]
27152. 2 | 12 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4769,25062,18335,18992,18328,18370,18321,18358,18314]
27154. 2 | 13 | 4 | 14 | 6 | 21 | 8 | 26 [avatar split clause 4767,25062,18335,18992,18328,18370,18321,18364,18314]
27155. 1 | ~2 | 3 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4766,25062,18335,18992,18328,18396,18318,18314,18311]
27157. 2 | 9 | 3 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4763,25062,18335,18992,18328,18396,18318,18340,18314]
27161. 2 | 11 | 3 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4759,25062,18335,18992,18328,18396,18318,18352,18314]
27163. 2 | 12 | 3 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4757,25062,18335,18992,18328,18396,18318,18358,18314]
27165. 2 | 13 | 3 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4755,25062,18335,18992,18328,18396,18318,18364,18314]
27166. 1 | ~2 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4754,25062,18335,18992,18328,18396,18321,18314,18311]
27168. 2 | 9 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4751,25062,18335,18992,18328,18396,18321,18340,18314]
27170. 2 | 10 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4749,25062,18335,18992,18328,18396,18321,18346,18314]
27172. 2 | 11 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4747,25062,18335,18992,18328,18396,18321,18352,18314]
27174. 2 | 12 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4745,25062,18335,18992,18328,18396,18321,18358,18314]
27176. 2 | 13 | 4 | 15 | 6 | 21 | 8 | 26 [avatar split clause 4743,25062,18335,18992,18328,18396,18321,18364,18314]
27177. 1 | ~2 | 3 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4742,25062,18335,18992,18328,18422,18318,18314,18311]
27179. 2 | 9 | 3 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4739,25062,18335,18992,18328,18422,18318,18340,18314]
27183. 2 | 11 | 3 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4735,25062,18335,18992,18328,18422,18318,18352,18314]
27188. 1 | ~2 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4730,25062,18335,18992,18328,18422,18321,18314,18311]
27190. 2 | 9 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4727,25062,18335,18992,18328,18422,18321,18340,18314]
27192. 2 | 10 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4725,25062,18335,18992,18328,18422,18321,18346,18314]
27194. 2 | 11 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4723,25062,18335,18992,18328,18422,18321,18352,18314]
27196. 2 | 12 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4721,25062,18335,18992,18328,18422,18321,18358,18314]
27198. 2 | 13 | 4 | 16 | 6 | 21 | 8 | 26 [avatar split clause 4719,25062,18335,18992,18328,18422,18321,18364,18314]
27199. 1 | ~2 | 3 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4718,25062,18335,18992,18328,18448,18318,18314,18311]
27201. 2 | 9 | 3 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4715,25062,18335,18992,18328,18448,18318,18340,18314]
27207. 2 | 12 | 3 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4709,25062,18335,18992,18328,18448,18318,18358,18314]
27209. 2 | 13 | 3 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4707,25062,18335,18992,18328,18448,18318,18364,18314]
27210. 1 | ~2 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4706,25062,18335,18992,18328,18448,18321,18314,18311]
27212. 2 | 9 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4703,25062,18335,18992,18328,18448,18321,18340,18314]
27214. 2 | 10 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4701,25062,18335,18992,18328,18448,18321,18346,18314]
27216. 2 | 11 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4699,25062,18335,18992,18328,18448,18321,18352,18314]
27218. 2 | 12 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4697,25062,18335,18992,18328,18448,18321,18358,18314]
27220. 2 | 13 | 4 | 17 | 6 | 21 | 8 | 26 [avatar split clause 4695,25062,18335,18992,18328,18448,18321,18364,18314]
27221. 1 | ~2 | 3 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4694,25062,18335,18992,18328,18474,18318,18314,18311]
27223. 2 | 9 | 3 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4691,25062,18335,18992,18328,18474,18318,18340,18314]
27227. 2 | 11 | 3 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4687,25062,18335,18992,18328,18474,18318,18352,18314]
27231. 2 | 13 | 3 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4683,25062,18335,18992,18328,18474,18318,18364,18314]
27232. 1 | ~2 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4682,25062,18335,18992,18328,18474,18321,18314,18311]
27234. 2 | 9 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4679,25062,18335,18992,18328,18474,18321,18340,18314]
27236. 2 | 10 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4677,25062,18335,18992,18328,18474,18321,18346,18314]
27238. 2 | 11 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4675,25062,18335,18992,18328,18474,18321,18352,18314]
27240. 2 | 12 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4673,25062,18335,18992,18328,18474,18321,18358,18314]
27242. 2 | 13 | 4 | 18 | 6 | 21 | 8 | 26 [avatar split clause 4671,25062,18335,18992,18328,18474,18321,18364,18314]
27249. 2 | 11 | 3 | ~4 | 5 | 22 | 8 | 26 [avatar split clause 4663,25062,18335,19238,18325,18321,18318,18352,18314]
27251. 2 | 12 | 3 | ~4 | 5 | 22 | 8 | 26 [avatar split clause 4661,25062,18335,19238,18325,18321,18318,18358,18314]
27364. 1 | ~2 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4526,25062,18335,19238,18328,18321,18318,18314,18311]
27366. 2 | 9 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4523,25062,18335,19238,18328,18321,18318,18340,18314]
27368. 2 | 10 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4521,25062,18335,19238,18328,18321,18318,18346,18314]
27370. 2 | 11 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4519,25062,18335,19238,18328,18321,18318,18352,18314]
27372. 2 | 12 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4517,25062,18335,19238,18328,18321,18318,18358,18314]
27374. 2 | 13 | 3 | ~4 | 6 | 22 | 8 | 26 [avatar split clause 4515,25062,18335,19238,18328,18321,18318,18364,18314]
27375. 1 | ~2 | 3 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4502,25062,18335,19238,18328,18370,18318,18314,18311]
27377. 2 | 9 | 3 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4499,25062,18335,19238,18328,18370,18318,18340,18314]
27379. 2 | 10 | 3 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4497,25062,18335,19238,18328,18370,18318,18346,18314]
27383. 2 | 12 | 3 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4493,25062,18335,19238,18328,18370,18318,18358,18314]
27385. 2 | 13 | 3 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4491,25062,18335,19238,18328,18370,18318,18364,18314]
27386. 1 | ~2 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4490,25062,18335,19238,18328,18370,18321,18314,18311]
27388. 2 | 9 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4487,25062,18335,19238,18328,18370,18321,18340,18314]
27390. 2 | 10 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4485,25062,18335,19238,18328,18370,18321,18346,18314]
27392. 2 | 11 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4483,25062,18335,19238,18328,18370,18321,18352,18314]
27394. 2 | 12 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4481,25062,18335,19238,18328,18370,18321,18358,18314]
27396. 2 | 13 | 4 | 14 | 6 | 22 | 8 | 26 [avatar split clause 4479,25062,18335,19238,18328,18370,18321,18364,18314]
27397. 1 | ~2 | 3 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4478,25062,18335,19238,18328,18396,18318,18314,18311]
27405. 2 | 12 | 3 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4469,25062,18335,19238,18328,18396,18318,18358,18314]
27407. 2 | 13 | 3 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4467,25062,18335,19238,18328,18396,18318,18364,18314]
27408. 1 | ~2 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4466,25062,18335,19238,18328,18396,18321,18314,18311]
27410. 2 | 9 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4463,25062,18335,19238,18328,18396,18321,18340,18314]
27412. 2 | 10 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4461,25062,18335,19238,18328,18396,18321,18346,18314]
27414. 2 | 11 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4459,25062,18335,19238,18328,18396,18321,18352,18314]
27416. 2 | 12 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4457,25062,18335,19238,18328,18396,18321,18358,18314]
27418. 2 | 13 | 4 | 15 | 6 | 22 | 8 | 26 [avatar split clause 4455,25062,18335,19238,18328,18396,18321,18364,18314]
27419. 1 | ~2 | 3 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4454,25062,18335,19238,18328,18422,18318,18314,18311]
27421. 2 | 9 | 3 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4451,25062,18335,19238,18328,18422,18318,18340,18314]
27423. 2 | 10 | 3 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4449,25062,18335,19238,18328,18422,18318,18346,18314]
27430. 1 | ~2 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4442,25062,18335,19238,18328,18422,18321,18314,18311]
27432. 2 | 9 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4439,25062,18335,19238,18328,18422,18321,18340,18314]
27434. 2 | 10 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4437,25062,18335,19238,18328,18422,18321,18346,18314]
27436. 2 | 11 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4435,25062,18335,19238,18328,18422,18321,18352,18314]
27438. 2 | 12 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4433,25062,18335,19238,18328,18422,18321,18358,18314]
27440. 2 | 13 | 4 | 16 | 6 | 22 | 8 | 26 [avatar split clause 4431,25062,18335,19238,18328,18422,18321,18364,18314]
27441. 1 | ~2 | 3 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4430,25062,18335,19238,18328,18448,18318,18314,18311]
27443. 2 | 9 | 3 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4427,25062,18335,19238,18328,18448,18318,18340,18314]
27449. 2 | 12 | 3 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4421,25062,18335,19238,18328,18448,18318,18358,18314]
27452. 1 | ~2 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4418,25062,18335,19238,18328,18448,18321,18314,18311]
27454. 2 | 9 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4415,25062,18335,19238,18328,18448,18321,18340,18314]
27456. 2 | 10 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4413,25062,18335,19238,18328,18448,18321,18346,18314]
27458. 2 | 11 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4411,25062,18335,19238,18328,18448,18321,18352,18314]
27459. 1 | 12 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4410,25062,18335,19238,18328,18448,18321,18358,18311]
27460. 2 | 12 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4409,25062,18335,19238,18328,18448,18321,18358,18314]
27462. 2 | 13 | 4 | 17 | 6 | 22 | 8 | 26 [avatar split clause 4407,25062,18335,19238,18328,18448,18321,18364,18314]
27463. 1 | ~2 | 3 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4406,25062,18335,19238,18328,18474,18318,18314,18311]
27467. 2 | 10 | 3 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4401,25062,18335,19238,18328,18474,18318,18346,18314]
27471. 2 | 12 | 3 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4397,25062,18335,19238,18328,18474,18318,18358,18314]
27474. 1 | ~2 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4394,25062,18335,19238,18328,18474,18321,18314,18311]
27476. 2 | 9 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4391,25062,18335,19238,18328,18474,18321,18340,18314]
27478. 2 | 10 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4389,25062,18335,19238,18328,18474,18321,18346,18314]
27480. 2 | 11 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4387,25062,18335,19238,18328,18474,18321,18352,18314]
27482. 2 | 12 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4385,25062,18335,19238,18328,18474,18321,18358,18314]
27484. 2 | 13 | 4 | 18 | 6 | 22 | 8 | 26 [avatar split clause 4383,25062,18335,19238,18328,18474,18321,18364,18314]
27489. 2 | 10 | 3 | ~4 | 5 | 23 | 8 | 26 [avatar split clause 4377,25062,18335,19484,18325,18321,18318,18346,18314]
27606. 1 | ~2 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4238,25062,18335,19484,18328,18321,18318,18314,18311]
27608. 2 | 9 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4235,25062,18335,19484,18328,18321,18318,18340,18314]
27610. 2 | 10 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4233,25062,18335,19484,18328,18321,18318,18346,18314]
27612. 2 | 11 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4231,25062,18335,19484,18328,18321,18318,18352,18314]
27614. 2 | 12 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4229,25062,18335,19484,18328,18321,18318,18358,18314]
27616. 2 | 13 | 3 | ~4 | 6 | 23 | 8 | 26 [avatar split clause 4227,25062,18335,19484,18328,18321,18318,18364,18314]
27617. 1 | ~2 | 3 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4214,25062,18335,19484,18328,18370,18318,18314,18311]
27621. 2 | 10 | 3 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4209,25062,18335,19484,18328,18370,18318,18346,18314]
27627. 2 | 13 | 3 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4203,25062,18335,19484,18328,18370,18318,18364,18314]
27628. 1 | ~2 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4202,25062,18335,19484,18328,18370,18321,18314,18311]
27630. 2 | 9 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4199,25062,18335,19484,18328,18370,18321,18340,18314]
27632. 2 | 10 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4197,25062,18335,19484,18328,18370,18321,18346,18314]
27634. 2 | 11 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4195,25062,18335,19484,18328,18370,18321,18352,18314]
27636. 2 | 12 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4193,25062,18335,19484,18328,18370,18321,18358,18314]
27638. 2 | 13 | 4 | 14 | 6 | 23 | 8 | 26 [avatar split clause 4191,25062,18335,19484,18328,18370,18321,18364,18314]
27639. 1 | ~2 | 3 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4190,25062,18335,19484,18328,18396,18318,18314,18311]
27643. 2 | 10 | 3 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4185,25062,18335,19484,18328,18396,18318,18346,18314]
27645. 2 | 11 | 3 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4183,25062,18335,19484,18328,18396,18318,18352,18314]
27650. 1 | ~2 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4178,25062,18335,19484,18328,18396,18321,18314,18311]
27652. 2 | 9 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4175,25062,18335,19484,18328,18396,18321,18340,18314]
27654. 2 | 10 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4173,25062,18335,19484,18328,18396,18321,18346,18314]
27656. 2 | 11 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4171,25062,18335,19484,18328,18396,18321,18352,18314]
27658. 2 | 12 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4169,25062,18335,19484,18328,18396,18321,18358,18314]
27660. 2 | 13 | 4 | 15 | 6 | 23 | 8 | 26 [avatar split clause 4167,25062,18335,19484,18328,18396,18321,18364,18314]
27661. 1 | ~2 | 3 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4166,25062,18335,19484,18328,18422,18318,18314,18311]
27665. 2 | 10 | 3 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4161,25062,18335,19484,18328,18422,18318,18346,18314]
27672. 1 | ~2 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4154,25062,18335,19484,18328,18422,18321,18314,18311]
27674. 2 | 9 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4151,25062,18335,19484,18328,18422,18321,18340,18314]
27676. 2 | 10 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4149,25062,18335,19484,18328,18422,18321,18346,18314]
27678. 2 | 11 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4147,25062,18335,19484,18328,18422,18321,18352,18314]
27680. 2 | 12 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4145,25062,18335,19484,18328,18422,18321,18358,18314]
27682. 2 | 13 | 4 | 16 | 6 | 23 | 8 | 26 [avatar split clause 4143,25062,18335,19484,18328,18422,18321,18364,18314]
27685. 2 | 9 | 3 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4139,25062,18335,19484,18328,18448,18318,18340,18314]
27687. 2 | 10 | 3 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4137,25062,18335,19484,18328,18448,18318,18346,18314]
27691. 2 | 12 | 3 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4133,25062,18335,19484,18328,18448,18318,18358,18314]
27693. 2 | 13 | 3 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4131,25062,18335,19484,18328,18448,18318,18364,18314]
27694. 1 | ~2 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4130,25062,18335,19484,18328,18448,18321,18314,18311]
27696. 2 | 9 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4127,25062,18335,19484,18328,18448,18321,18340,18314]
27698. 2 | 10 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4125,25062,18335,19484,18328,18448,18321,18346,18314]
27700. 2 | 11 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4123,25062,18335,19484,18328,18448,18321,18352,18314]
27702. 2 | 12 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4121,25062,18335,19484,18328,18448,18321,18358,18314]
27704. 2 | 13 | 4 | 17 | 6 | 23 | 8 | 26 [avatar split clause 4119,25062,18335,19484,18328,18448,18321,18364,18314]
27705. 1 | ~2 | 3 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4118,25062,18335,19484,18328,18474,18318,18314,18311]
27716. 1 | ~2 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4106,25062,18335,19484,18328,18474,18321,18314,18311]
27718. 2 | 9 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4103,25062,18335,19484,18328,18474,18321,18340,18314]
27720. 2 | 10 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4101,25062,18335,19484,18328,18474,18321,18346,18314]
27722. 2 | 11 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4099,25062,18335,19484,18328,18474,18321,18352,18314]
27724. 2 | 12 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4097,25062,18335,19484,18328,18474,18321,18358,18314]
27726. 2 | 13 | 4 | 18 | 6 | 23 | 8 | 26 [avatar split clause 4095,25062,18335,19484,18328,18474,18321,18364,18314]
27728. 27 <=> m2_subset_1(sK1,k5_numbers,k1_zf_lang) [avatar definition]
27730. m2_subset_1(sK1,k5_numbers,k1_zf_lang) <- (27) [avatar component clause 27728] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
29062. 1 | ~2 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2366,27728,18335,18328,18325,18321,18318,18314,18311]
29064. 2 | 9 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2363,27728,18335,18328,18325,18321,18318,18340,18314]
29066. 2 | 10 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2361,27728,18335,18328,18325,18321,18318,18346,18314]
29068. 2 | 11 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2359,27728,18335,18328,18325,18321,18318,18352,18314]
29070. 2 | 12 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2357,27728,18335,18328,18325,18321,18318,18358,18314]
29072. 2 | 13 | 3 | ~4 | 5 | ~6 | 8 | 27 [avatar split clause 2355,27728,18335,18328,18325,18321,18318,18364,18314]
29084. 1 | ~2 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2330,27728,18335,18328,18325,18370,18321,18314,18311]
29085. 1 | 9 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2328,27728,18335,18328,18325,18370,18321,18340,18311]
29086. 2 | 9 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2327,27728,18335,18328,18325,18370,18321,18340,18314]
29088. 2 | 10 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2325,27728,18335,18328,18325,18370,18321,18346,18314]
29089. 1 | 11 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2324,27728,18335,18328,18325,18370,18321,18352,18311]
29090. 2 | 11 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2323,27728,18335,18328,18325,18370,18321,18352,18314]
29091. 1 | 12 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2322,27728,18335,18328,18325,18370,18321,18358,18311]
29092. 2 | 12 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2321,27728,18335,18328,18325,18370,18321,18358,18314]
29093. 1 | 13 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2320,27728,18335,18328,18325,18370,18321,18364,18311]
29094. 2 | 13 | 4 | 14 | 5 | ~6 | 8 | 27 [avatar split clause 2319,27728,18335,18328,18325,18370,18321,18364,18314]
29106. 1 | ~2 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2306,27728,18335,18328,18325,18396,18321,18314,18311]
29107. 1 | 9 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2304,27728,18335,18328,18325,18396,18321,18340,18311]
29108. 2 | 9 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2303,27728,18335,18328,18325,18396,18321,18340,18314]
29109. 1 | 10 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2302,27728,18335,18328,18325,18396,18321,18346,18311]
29110. 2 | 10 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2301,27728,18335,18328,18325,18396,18321,18346,18314]
29111. 1 | 11 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2300,27728,18335,18328,18325,18396,18321,18352,18311]
29112. 2 | 11 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2299,27728,18335,18328,18325,18396,18321,18352,18314]
29114. 2 | 12 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2297,27728,18335,18328,18325,18396,18321,18358,18314]
29115. 1 | 13 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2296,27728,18335,18328,18325,18396,18321,18364,18311]
29116. 2 | 13 | 4 | 15 | 5 | ~6 | 8 | 27 [avatar split clause 2295,27728,18335,18328,18325,18396,18321,18364,18314]
29128. 1 | ~2 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2282,27728,18335,18328,18325,18422,18321,18314,18311]
29129. 1 | 9 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2280,27728,18335,18328,18325,18422,18321,18340,18311]
29130. 2 | 9 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2279,27728,18335,18328,18325,18422,18321,18340,18314]
29131. 1 | 10 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2278,27728,18335,18328,18325,18422,18321,18346,18311]
29132. 2 | 10 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2277,27728,18335,18328,18325,18422,18321,18346,18314]
29133. 1 | 11 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2276,27728,18335,18328,18325,18422,18321,18352,18311]
29134. 2 | 11 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2275,27728,18335,18328,18325,18422,18321,18352,18314]
29135. 1 | 12 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2274,27728,18335,18328,18325,18422,18321,18358,18311]
29136. 2 | 12 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2273,27728,18335,18328,18325,18422,18321,18358,18314]
29137. 1 | 13 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2272,27728,18335,18328,18325,18422,18321,18364,18311]
29138. 2 | 13 | 4 | 16 | 5 | ~6 | 8 | 27 [avatar split clause 2271,27728,18335,18328,18325,18422,18321,18364,18314]
29150. 1 | ~2 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2258,27728,18335,18328,18325,18448,18321,18314,18311]
29151. 1 | 9 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2256,27728,18335,18328,18325,18448,18321,18340,18311]
29152. 2 | 9 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2255,27728,18335,18328,18325,18448,18321,18340,18314]
29154. 2 | 10 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2253,27728,18335,18328,18325,18448,18321,18346,18314]
29155. 1 | 11 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2252,27728,18335,18328,18325,18448,18321,18352,18311]
29156. 2 | 11 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2251,27728,18335,18328,18325,18448,18321,18352,18314]
29158. 2 | 12 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2249,27728,18335,18328,18325,18448,18321,18358,18314]
29159. 1 | 13 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2248,27728,18335,18328,18325,18448,18321,18364,18311]
29160. 2 | 13 | 4 | 17 | 5 | ~6 | 8 | 27 [avatar split clause 2247,27728,18335,18328,18325,18448,18321,18364,18314]
29172. 1 | ~2 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2234,27728,18335,18328,18325,18474,18321,18314,18311]
29173. 1 | 9 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2232,27728,18335,18328,18325,18474,18321,18340,18311]
29174. 2 | 9 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2231,27728,18335,18328,18325,18474,18321,18340,18314]
29175. 1 | 10 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2230,27728,18335,18328,18325,18474,18321,18346,18311]
29176. 2 | 10 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2229,27728,18335,18328,18325,18474,18321,18346,18314]
29177. 1 | 11 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2228,27728,18335,18328,18325,18474,18321,18352,18311]
29178. 2 | 11 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2227,27728,18335,18328,18325,18474,18321,18352,18314]
29179. 1 | 12 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2226,27728,18335,18328,18325,18474,18321,18358,18311]
29180. 2 | 12 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2225,27728,18335,18328,18325,18474,18321,18358,18314]
29181. 1 | 13 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2224,27728,18335,18328,18325,18474,18321,18364,18311]
29182. 2 | 13 | 4 | 18 | 5 | ~6 | 8 | 27 [avatar split clause 2223,27728,18335,18328,18325,18474,18321,18364,18314]
29187. 2 | 10 | 3 | ~4 | 5 | 19 | 8 | 27 [avatar split clause 2073,27728,18335,18500,18325,18321,18318,18346,18314]
29189. 2 | 11 | 3 | ~4 | 5 | 19 | 8 | 27 [avatar split clause 2071,27728,18335,18500,18325,18321,18318,18352,18314]
29304. 1 | ~2 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1934,27728,18335,18500,18328,18321,18318,18314,18311]
29306. 2 | 9 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1931,27728,18335,18500,18328,18321,18318,18340,18314]
29308. 2 | 10 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1929,27728,18335,18500,18328,18321,18318,18346,18314]
29310. 2 | 11 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1927,27728,18335,18500,18328,18321,18318,18352,18314]
29312. 2 | 12 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1925,27728,18335,18500,18328,18321,18318,18358,18314]
29314. 2 | 13 | 3 | ~4 | 6 | 19 | 8 | 27 [avatar split clause 1923,27728,18335,18500,18328,18321,18318,18364,18314]
29317. 2 | 9 | 3 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1907,27728,18335,18500,18328,18370,18318,18340,18314]
29319. 2 | 10 | 3 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1905,27728,18335,18500,18328,18370,18318,18346,18314]
29321. 2 | 11 | 3 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1903,27728,18335,18500,18328,18370,18318,18352,18314]
29323. 2 | 12 | 3 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1901,27728,18335,18500,18328,18370,18318,18358,18314]
29326. 1 | ~2 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1898,27728,18335,18500,18328,18370,18321,18314,18311]
29328. 2 | 9 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1895,27728,18335,18500,18328,18370,18321,18340,18314]
29330. 2 | 10 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1893,27728,18335,18500,18328,18370,18321,18346,18314]
29332. 2 | 11 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1891,27728,18335,18500,18328,18370,18321,18352,18314]
29334. 2 | 12 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1889,27728,18335,18500,18328,18370,18321,18358,18314]
29336. 2 | 13 | 4 | 14 | 6 | 19 | 8 | 27 [avatar split clause 1887,27728,18335,18500,18328,18370,18321,18364,18314]
29341. 2 | 10 | 3 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1881,27728,18335,18500,18328,18396,18318,18346,18314]
29348. 1 | ~2 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1874,27728,18335,18500,18328,18396,18321,18314,18311]
29350. 2 | 9 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1871,27728,18335,18500,18328,18396,18321,18340,18314]
29352. 2 | 10 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1869,27728,18335,18500,18328,18396,18321,18346,18314]
29354. 2 | 11 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1867,27728,18335,18500,18328,18396,18321,18352,18314]
29356. 2 | 12 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1865,27728,18335,18500,18328,18396,18321,18358,18314]
29358. 2 | 13 | 4 | 15 | 6 | 19 | 8 | 27 [avatar split clause 1863,27728,18335,18500,18328,18396,18321,18364,18314]
29361. 2 | 9 | 3 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1859,27728,18335,18500,18328,18422,18318,18340,18314]
29363. 2 | 10 | 3 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1857,27728,18335,18500,18328,18422,18318,18346,18314]
29365. 2 | 11 | 3 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1855,27728,18335,18500,18328,18422,18318,18352,18314]
29370. 1 | ~2 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1850,27728,18335,18500,18328,18422,18321,18314,18311]
29372. 2 | 9 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1847,27728,18335,18500,18328,18422,18321,18340,18314]
29374. 2 | 10 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1845,27728,18335,18500,18328,18422,18321,18346,18314]
29376. 2 | 11 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1843,27728,18335,18500,18328,18422,18321,18352,18314]
29378. 2 | 12 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1841,27728,18335,18500,18328,18422,18321,18358,18314]
29380. 2 | 13 | 4 | 16 | 6 | 19 | 8 | 27 [avatar split clause 1839,27728,18335,18500,18328,18422,18321,18364,18314]
29381. 1 | ~2 | 3 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1838,27728,18335,18500,18328,18448,18318,18314,18311]
29385. 2 | 10 | 3 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1833,27728,18335,18500,18328,18448,18318,18346,18314]
29389. 2 | 12 | 3 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1829,27728,18335,18500,18328,18448,18318,18358,18314]
29391. 2 | 13 | 3 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1827,27728,18335,18500,18328,18448,18318,18364,18314]
29392. 1 | ~2 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1826,27728,18335,18500,18328,18448,18321,18314,18311]
29394. 2 | 9 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1823,27728,18335,18500,18328,18448,18321,18340,18314]
29396. 2 | 10 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1821,27728,18335,18500,18328,18448,18321,18346,18314]
29398. 2 | 11 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1819,27728,18335,18500,18328,18448,18321,18352,18314]
29400. 2 | 12 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1817,27728,18335,18500,18328,18448,18321,18358,18314]
29402. 2 | 13 | 4 | 17 | 6 | 19 | 8 | 27 [avatar split clause 1815,27728,18335,18500,18328,18448,18321,18364,18314]
29405. 2 | 9 | 3 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1811,27728,18335,18500,18328,18474,18318,18340,18314]
29409. 2 | 11 | 3 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1807,27728,18335,18500,18328,18474,18318,18352,18314]
29411. 2 | 12 | 3 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1805,27728,18335,18500,18328,18474,18318,18358,18314]
29413. 2 | 13 | 3 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1803,27728,18335,18500,18328,18474,18318,18364,18314]
29414. 1 | ~2 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1802,27728,18335,18500,18328,18474,18321,18314,18311]
29416. 2 | 9 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1799,27728,18335,18500,18328,18474,18321,18340,18314]
29418. 2 | 10 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1797,27728,18335,18500,18328,18474,18321,18346,18314]
29420. 2 | 11 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1795,27728,18335,18500,18328,18474,18321,18352,18314]
29422. 2 | 12 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1793,27728,18335,18500,18328,18474,18321,18358,18314]
29424. 2 | 13 | 4 | 18 | 6 | 19 | 8 | 27 [avatar split clause 1791,27728,18335,18500,18328,18474,18321,18364,18314]
29546. 1 | ~2 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1646,27728,18335,18746,18328,18321,18318,18314,18311]
29548. 2 | 9 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1643,27728,18335,18746,18328,18321,18318,18340,18314]
29550. 2 | 10 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1641,27728,18335,18746,18328,18321,18318,18346,18314]
29552. 2 | 11 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1639,27728,18335,18746,18328,18321,18318,18352,18314]
29554. 2 | 12 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1637,27728,18335,18746,18328,18321,18318,18358,18314]
29556. 2 | 13 | 3 | ~4 | 6 | 20 | 8 | 27 [avatar split clause 1635,27728,18335,18746,18328,18321,18318,18364,18314]
29561. 2 | 10 | 3 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1617,27728,18335,18746,18328,18370,18318,18346,18314]
29568. 1 | ~2 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1610,27728,18335,18746,18328,18370,18321,18314,18311]
29570. 2 | 9 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1607,27728,18335,18746,18328,18370,18321,18340,18314]
29572. 2 | 10 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1605,27728,18335,18746,18328,18370,18321,18346,18314]
29574. 2 | 11 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1603,27728,18335,18746,18328,18370,18321,18352,18314]
29576. 2 | 12 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1601,27728,18335,18746,18328,18370,18321,18358,18314]
29578. 2 | 13 | 4 | 14 | 6 | 20 | 8 | 27 [avatar split clause 1599,27728,18335,18746,18328,18370,18321,18364,18314]
29581. 2 | 9 | 3 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1595,27728,18335,18746,18328,18396,18318,18340,18314]
29583. 2 | 10 | 3 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1593,27728,18335,18746,18328,18396,18318,18346,18314]
29587. 2 | 12 | 3 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1589,27728,18335,18746,18328,18396,18318,18358,18314]
29590. 1 | ~2 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1586,27728,18335,18746,18328,18396,18321,18314,18311]
29592. 2 | 9 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1583,27728,18335,18746,18328,18396,18321,18340,18314]
29594. 2 | 10 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1581,27728,18335,18746,18328,18396,18321,18346,18314]
29596. 2 | 11 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1579,27728,18335,18746,18328,18396,18321,18352,18314]
29598. 2 | 12 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1577,27728,18335,18746,18328,18396,18321,18358,18314]
29600. 2 | 13 | 4 | 15 | 6 | 20 | 8 | 27 [avatar split clause 1575,27728,18335,18746,18328,18396,18321,18364,18314]
29601. 1 | ~2 | 3 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1574,27728,18335,18746,18328,18422,18318,18314,18311]
29603. 2 | 9 | 3 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1571,27728,18335,18746,18328,18422,18318,18340,18314]
29607. 2 | 11 | 3 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1567,27728,18335,18746,18328,18422,18318,18352,18314]
29612. 1 | ~2 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1562,27728,18335,18746,18328,18422,18321,18314,18311]
29614. 2 | 9 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1559,27728,18335,18746,18328,18422,18321,18340,18314]
29616. 2 | 10 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1557,27728,18335,18746,18328,18422,18321,18346,18314]
29618. 2 | 11 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1555,27728,18335,18746,18328,18422,18321,18352,18314]
29620. 2 | 12 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1553,27728,18335,18746,18328,18422,18321,18358,18314]
29622. 2 | 13 | 4 | 16 | 6 | 20 | 8 | 27 [avatar split clause 1551,27728,18335,18746,18328,18422,18321,18364,18314]
29625. 2 | 9 | 3 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1547,27728,18335,18746,18328,18448,18318,18340,18314]
29627. 2 | 10 | 3 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1545,27728,18335,18746,18328,18448,18318,18346,18314]
29631. 2 | 12 | 3 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1541,27728,18335,18746,18328,18448,18318,18358,18314]
29633. 2 | 13 | 3 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1539,27728,18335,18746,18328,18448,18318,18364,18314]
29634. 1 | ~2 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1538,27728,18335,18746,18328,18448,18321,18314,18311]
29636. 2 | 9 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1535,27728,18335,18746,18328,18448,18321,18340,18314]
29638. 2 | 10 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1533,27728,18335,18746,18328,18448,18321,18346,18314]
29640. 2 | 11 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1531,27728,18335,18746,18328,18448,18321,18352,18314]
29642. 2 | 12 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1529,27728,18335,18746,18328,18448,18321,18358,18314]
29644. 2 | 13 | 4 | 17 | 6 | 20 | 8 | 27 [avatar split clause 1527,27728,18335,18746,18328,18448,18321,18364,18314]
29653. 2 | 12 | 3 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1517,27728,18335,18746,18328,18474,18318,18358,18314]
29655. 2 | 13 | 3 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1515,27728,18335,18746,18328,18474,18318,18364,18314]
29656. 1 | ~2 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1514,27728,18335,18746,18328,18474,18321,18314,18311]
29658. 2 | 9 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1511,27728,18335,18746,18328,18474,18321,18340,18314]
29660. 2 | 10 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1509,27728,18335,18746,18328,18474,18321,18346,18314]
29662. 2 | 11 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1507,27728,18335,18746,18328,18474,18321,18352,18314]
29664. 2 | 12 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1505,27728,18335,18746,18328,18474,18321,18358,18314]
29666. 2 | 13 | 4 | 18 | 6 | 20 | 8 | 27 [avatar split clause 1503,27728,18335,18746,18328,18474,18321,18364,18314]
29673. 2 | 11 | 3 | ~4 | 5 | 21 | 8 | 27 [avatar split clause 1495,27728,18335,18992,18325,18321,18318,18352,18314]
29788. 1 | ~2 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1358,27728,18335,18992,18328,18321,18318,18314,18311]
29790. 2 | 9 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1355,27728,18335,18992,18328,18321,18318,18340,18314]
29792. 2 | 10 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1353,27728,18335,18992,18328,18321,18318,18346,18314]
29794. 2 | 11 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1351,27728,18335,18992,18328,18321,18318,18352,18314]
29796. 2 | 12 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1349,27728,18335,18992,18328,18321,18318,18358,18314]
29798. 2 | 13 | 3 | ~4 | 6 | 21 | 8 | 27 [avatar split clause 1347,27728,18335,18992,18328,18321,18318,18364,18314]
29801. 2 | 9 | 3 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1331,27728,18335,18992,18328,18370,18318,18340,18314]
29803. 2 | 10 | 3 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1329,27728,18335,18992,18328,18370,18318,18346,18314]
29805. 2 | 11 | 3 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1327,27728,18335,18992,18328,18370,18318,18352,18314]
29807. 2 | 12 | 3 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1325,27728,18335,18992,18328,18370,18318,18358,18314]
29809. 2 | 13 | 3 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1323,27728,18335,18992,18328,18370,18318,18364,18314]
29810. 1 | ~2 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1322,27728,18335,18992,18328,18370,18321,18314,18311]
29812. 2 | 9 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1319,27728,18335,18992,18328,18370,18321,18340,18314]
29814. 2 | 10 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1317,27728,18335,18992,18328,18370,18321,18346,18314]
29816. 2 | 11 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1315,27728,18335,18992,18328,18370,18321,18352,18314]
29818. 2 | 12 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1313,27728,18335,18992,18328,18370,18321,18358,18314]
29820. 2 | 13 | 4 | 14 | 6 | 21 | 8 | 27 [avatar split clause 1311,27728,18335,18992,18328,18370,18321,18364,18314]
29823. 2 | 9 | 3 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1307,27728,18335,18992,18328,18396,18318,18340,18314]
29825. 2 | 10 | 3 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1305,27728,18335,18992,18328,18396,18318,18346,18314]
29827. 2 | 11 | 3 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1303,27728,18335,18992,18328,18396,18318,18352,18314]
29829. 2 | 12 | 3 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1301,27728,18335,18992,18328,18396,18318,18358,18314]
29831. 2 | 13 | 3 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1299,27728,18335,18992,18328,18396,18318,18364,18314]
29832. 1 | ~2 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1298,27728,18335,18992,18328,18396,18321,18314,18311]
29834. 2 | 9 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1295,27728,18335,18992,18328,18396,18321,18340,18314]
29836. 2 | 10 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1293,27728,18335,18992,18328,18396,18321,18346,18314]
29838. 2 | 11 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1291,27728,18335,18992,18328,18396,18321,18352,18314]
29840. 2 | 12 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1289,27728,18335,18992,18328,18396,18321,18358,18314]
29842. 2 | 13 | 4 | 15 | 6 | 21 | 8 | 27 [avatar split clause 1287,27728,18335,18992,18328,18396,18321,18364,18314]
29843. 1 | ~2 | 3 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1286,27728,18335,18992,18328,18422,18318,18314,18311]
29845. 2 | 9 | 3 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1283,27728,18335,18992,18328,18422,18318,18340,18314]
29849. 2 | 11 | 3 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1279,27728,18335,18992,18328,18422,18318,18352,18314]
29853. 2 | 13 | 3 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1275,27728,18335,18992,18328,18422,18318,18364,18314]
29854. 1 | ~2 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1274,27728,18335,18992,18328,18422,18321,18314,18311]
29856. 2 | 9 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1271,27728,18335,18992,18328,18422,18321,18340,18314]
29858. 2 | 10 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1269,27728,18335,18992,18328,18422,18321,18346,18314]
29860. 2 | 11 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1267,27728,18335,18992,18328,18422,18321,18352,18314]
29862. 2 | 12 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1265,27728,18335,18992,18328,18422,18321,18358,18314]
29864. 2 | 13 | 4 | 16 | 6 | 21 | 8 | 27 [avatar split clause 1263,27728,18335,18992,18328,18422,18321,18364,18314]
29865. 1 | ~2 | 3 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1262,27728,18335,18992,18328,18448,18318,18314,18311]
29867. 2 | 9 | 3 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1259,27728,18335,18992,18328,18448,18318,18340,18314]
29873. 2 | 12 | 3 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1253,27728,18335,18992,18328,18448,18318,18358,18314]
29875. 2 | 13 | 3 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1251,27728,18335,18992,18328,18448,18318,18364,18314]
29876. 1 | ~2 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1250,27728,18335,18992,18328,18448,18321,18314,18311]
29878. 2 | 9 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1247,27728,18335,18992,18328,18448,18321,18340,18314]
29880. 2 | 10 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1245,27728,18335,18992,18328,18448,18321,18346,18314]
29882. 2 | 11 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1243,27728,18335,18992,18328,18448,18321,18352,18314]
29884. 2 | 12 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1241,27728,18335,18992,18328,18448,18321,18358,18314]
29886. 2 | 13 | 4 | 17 | 6 | 21 | 8 | 27 [avatar split clause 1239,27728,18335,18992,18328,18448,18321,18364,18314]
29887. 1 | ~2 | 3 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1238,27728,18335,18992,18328,18474,18318,18314,18311]
29889. 2 | 9 | 3 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1235,27728,18335,18992,18328,18474,18318,18340,18314]
29893. 2 | 11 | 3 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1231,27728,18335,18992,18328,18474,18318,18352,18314]
29897. 2 | 13 | 3 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1227,27728,18335,18992,18328,18474,18318,18364,18314]
29898. 1 | ~2 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1226,27728,18335,18992,18328,18474,18321,18314,18311]
29900. 2 | 9 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1223,27728,18335,18992,18328,18474,18321,18340,18314]
29902. 2 | 10 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1221,27728,18335,18992,18328,18474,18321,18346,18314]
29904. 2 | 11 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1219,27728,18335,18992,18328,18474,18321,18352,18314]
29906. 2 | 12 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1217,27728,18335,18992,18328,18474,18321,18358,18314]
29908. 2 | 13 | 4 | 18 | 6 | 21 | 8 | 27 [avatar split clause 1215,27728,18335,18992,18328,18474,18321,18364,18314]
29917. 2 | 12 | 3 | ~4 | 5 | 22 | 8 | 27 [avatar split clause 1205,27728,18335,19238,18325,18321,18318,18358,18314]
30030. 1 | ~2 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1070,27728,18335,19238,18328,18321,18318,18314,18311]
30032. 2 | 9 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1067,27728,18335,19238,18328,18321,18318,18340,18314]
30034. 2 | 10 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1065,27728,18335,19238,18328,18321,18318,18346,18314]
30036. 2 | 11 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1063,27728,18335,19238,18328,18321,18318,18352,18314]
30038. 2 | 12 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1061,27728,18335,19238,18328,18321,18318,18358,18314]
30040. 2 | 13 | 3 | ~4 | 6 | 22 | 8 | 27 [avatar split clause 1059,27728,18335,19238,18328,18321,18318,18364,18314]
30043. 2 | 9 | 3 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1043,27728,18335,19238,18328,18370,18318,18340,18314]
30045. 2 | 10 | 3 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1041,27728,18335,19238,18328,18370,18318,18346,18314]
30047. 2 | 11 | 3 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1039,27728,18335,19238,18328,18370,18318,18352,18314]
30051. 2 | 13 | 3 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1035,27728,18335,19238,18328,18370,18318,18364,18314]
30052. 1 | ~2 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1034,27728,18335,19238,18328,18370,18321,18314,18311]
30054. 2 | 9 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1031,27728,18335,19238,18328,18370,18321,18340,18314]
30056. 2 | 10 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1029,27728,18335,19238,18328,18370,18321,18346,18314]
30058. 2 | 11 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1027,27728,18335,19238,18328,18370,18321,18352,18314]
30060. 2 | 12 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1025,27728,18335,19238,18328,18370,18321,18358,18314]
30062. 2 | 13 | 4 | 14 | 6 | 22 | 8 | 27 [avatar split clause 1023,27728,18335,19238,18328,18370,18321,18364,18314]
30065. 2 | 9 | 3 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1019,27728,18335,19238,18328,18396,18318,18340,18314]
30067. 2 | 10 | 3 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1017,27728,18335,19238,18328,18396,18318,18346,18314]
30071. 2 | 12 | 3 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1013,27728,18335,19238,18328,18396,18318,18358,18314]
30073. 2 | 13 | 3 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1011,27728,18335,19238,18328,18396,18318,18364,18314]
30074. 1 | ~2 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1010,27728,18335,19238,18328,18396,18321,18314,18311]
30076. 2 | 9 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1007,27728,18335,19238,18328,18396,18321,18340,18314]
30078. 2 | 10 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1005,27728,18335,19238,18328,18396,18321,18346,18314]
30080. 2 | 11 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1003,27728,18335,19238,18328,18396,18321,18352,18314]
30082. 2 | 12 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 1001,27728,18335,19238,18328,18396,18321,18358,18314]
30084. 2 | 13 | 4 | 15 | 6 | 22 | 8 | 27 [avatar split clause 999,27728,18335,19238,18328,18396,18321,18364,18314]
30091. 2 | 11 | 3 | 16 | 6 | 22 | 8 | 27 [avatar split clause 991,27728,18335,19238,18328,18422,18318,18352,18314]
30096. 1 | ~2 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 986,27728,18335,19238,18328,18422,18321,18314,18311]
30098. 2 | 9 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 983,27728,18335,19238,18328,18422,18321,18340,18314]
30100. 2 | 10 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 981,27728,18335,19238,18328,18422,18321,18346,18314]
30102. 2 | 11 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 979,27728,18335,19238,18328,18422,18321,18352,18314]
30104. 2 | 12 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 977,27728,18335,19238,18328,18422,18321,18358,18314]
30106. 2 | 13 | 4 | 16 | 6 | 22 | 8 | 27 [avatar split clause 975,27728,18335,19238,18328,18422,18321,18364,18314]
30109. 2 | 9 | 3 | 17 | 6 | 22 | 8 | 27 [avatar split clause 971,27728,18335,19238,18328,18448,18318,18340,18314]
30111. 2 | 10 | 3 | 17 | 6 | 22 | 8 | 27 [avatar split clause 969,27728,18335,19238,18328,18448,18318,18346,18314]
30115. 2 | 12 | 3 | 17 | 6 | 22 | 8 | 27 [avatar split clause 965,27728,18335,19238,18328,18448,18318,18358,18314]
30118. 1 | ~2 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 962,27728,18335,19238,18328,18448,18321,18314,18311]
30120. 2 | 9 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 959,27728,18335,19238,18328,18448,18321,18340,18314]
30122. 2 | 10 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 957,27728,18335,19238,18328,18448,18321,18346,18314]
30124. 2 | 11 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 955,27728,18335,19238,18328,18448,18321,18352,18314]
30126. 2 | 12 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 953,27728,18335,19238,18328,18448,18321,18358,18314]
30128. 2 | 13 | 4 | 17 | 6 | 22 | 8 | 27 [avatar split clause 951,27728,18335,19238,18328,18448,18321,18364,18314]
30135. 2 | 11 | 3 | 18 | 6 | 22 | 8 | 27 [avatar split clause 943,27728,18335,19238,18328,18474,18318,18352,18314]
30137. 2 | 12 | 3 | 18 | 6 | 22 | 8 | 27 [avatar split clause 941,27728,18335,19238,18328,18474,18318,18358,18314]
30140. 1 | ~2 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 938,27728,18335,19238,18328,18474,18321,18314,18311]
30142. 2 | 9 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 935,27728,18335,19238,18328,18474,18321,18340,18314]
30144. 2 | 10 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 933,27728,18335,19238,18328,18474,18321,18346,18314]
30146. 2 | 11 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 931,27728,18335,19238,18328,18474,18321,18352,18314]
30148. 2 | 12 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 929,27728,18335,19238,18328,18474,18321,18358,18314]
30150. 2 | 13 | 4 | 18 | 6 | 22 | 8 | 27 [avatar split clause 927,27728,18335,19238,18328,18474,18321,18364,18314]
30153. 2 | 9 | 3 | ~4 | 5 | 23 | 8 | 27 [avatar split clause 923,27728,18335,19484,18325,18321,18318,18340,18314]
30157. 2 | 11 | 3 | ~4 | 5 | 23 | 8 | 27 [avatar split clause 919,27728,18335,19484,18325,18321,18318,18352,18314]
30272. 1 | ~2 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 782,27728,18335,19484,18328,18321,18318,18314,18311]
30274. 2 | 9 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 779,27728,18335,19484,18328,18321,18318,18340,18314]
30276. 2 | 10 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 777,27728,18335,19484,18328,18321,18318,18346,18314]
30278. 2 | 11 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 775,27728,18335,19484,18328,18321,18318,18352,18314]
30280. 2 | 12 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 773,27728,18335,19484,18328,18321,18318,18358,18314]
30282. 2 | 13 | 3 | ~4 | 6 | 23 | 8 | 27 [avatar split clause 771,27728,18335,19484,18328,18321,18318,18364,18314]
30283. 1 | ~2 | 3 | 14 | 6 | 23 | 8 | 27 [avatar split clause 758,27728,18335,19484,18328,18370,18318,18314,18311]
30287. 2 | 10 | 3 | 14 | 6 | 23 | 8 | 27 [avatar split clause 753,27728,18335,19484,18328,18370,18318,18346,18314]
30294. 1 | ~2 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 746,27728,18335,19484,18328,18370,18321,18314,18311]
30296. 2 | 9 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 743,27728,18335,19484,18328,18370,18321,18340,18314]
30298. 2 | 10 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 741,27728,18335,19484,18328,18370,18321,18346,18314]
30300. 2 | 11 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 739,27728,18335,19484,18328,18370,18321,18352,18314]
30302. 2 | 12 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 737,27728,18335,19484,18328,18370,18321,18358,18314]
30304. 2 | 13 | 4 | 14 | 6 | 23 | 8 | 27 [avatar split clause 735,27728,18335,19484,18328,18370,18321,18364,18314]
30305. 1 | ~2 | 3 | 15 | 6 | 23 | 8 | 27 [avatar split clause 734,27728,18335,19484,18328,18396,18318,18314,18311]
30309. 2 | 10 | 3 | 15 | 6 | 23 | 8 | 27 [avatar split clause 729,27728,18335,19484,18328,18396,18318,18346,18314]
30311. 2 | 11 | 3 | 15 | 6 | 23 | 8 | 27 [avatar split clause 727,27728,18335,19484,18328,18396,18318,18352,18314]
30316. 1 | ~2 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 722,27728,18335,19484,18328,18396,18321,18314,18311]
30318. 2 | 9 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 719,27728,18335,19484,18328,18396,18321,18340,18314]
30320. 2 | 10 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 717,27728,18335,19484,18328,18396,18321,18346,18314]
30322. 2 | 11 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 715,27728,18335,19484,18328,18396,18321,18352,18314]
30324. 2 | 12 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 713,27728,18335,19484,18328,18396,18321,18358,18314]
30326. 2 | 13 | 4 | 15 | 6 | 23 | 8 | 27 [avatar split clause 711,27728,18335,19484,18328,18396,18321,18364,18314]
30331. 2 | 10 | 3 | 16 | 6 | 23 | 8 | 27 [avatar split clause 705,27728,18335,19484,18328,18422,18318,18346,18314]
30338. 1 | ~2 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 698,27728,18335,19484,18328,18422,18321,18314,18311]
30340. 2 | 9 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 695,27728,18335,19484,18328,18422,18321,18340,18314]
30342. 2 | 10 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 693,27728,18335,19484,18328,18422,18321,18346,18314]
30344. 2 | 11 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 691,27728,18335,19484,18328,18422,18321,18352,18314]
30346. 2 | 12 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 689,27728,18335,19484,18328,18422,18321,18358,18314]
30348. 2 | 13 | 4 | 16 | 6 | 23 | 8 | 27 [avatar split clause 687,27728,18335,19484,18328,18422,18321,18364,18314]
30349. 1 | ~2 | 3 | 17 | 6 | 23 | 8 | 27 [avatar split clause 686,27728,18335,19484,18328,18448,18318,18314,18311]
30351. 2 | 9 | 3 | 17 | 6 | 23 | 8 | 27 [avatar split clause 683,27728,18335,19484,18328,18448,18318,18340,18314]
30355. 2 | 11 | 3 | 17 | 6 | 23 | 8 | 27 [avatar split clause 679,27728,18335,19484,18328,18448,18318,18352,18314]
30360. 1 | ~2 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 674,27728,18335,19484,18328,18448,18321,18314,18311]
30362. 2 | 9 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 671,27728,18335,19484,18328,18448,18321,18340,18314]
30364. 2 | 10 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 669,27728,18335,19484,18328,18448,18321,18346,18314]
30366. 2 | 11 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 667,27728,18335,19484,18328,18448,18321,18352,18314]
30368. 2 | 12 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 665,27728,18335,19484,18328,18448,18321,18358,18314]
30370. 2 | 13 | 4 | 17 | 6 | 23 | 8 | 27 [avatar split clause 663,27728,18335,19484,18328,18448,18321,18364,18314]
30373. 2 | 9 | 3 | 18 | 6 | 23 | 8 | 27 [avatar split clause 659,27728,18335,19484,18328,18474,18318,18340,18314]
30382. 1 | ~2 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 650,27728,18335,19484,18328,18474,18321,18314,18311]
30384. 2 | 9 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 647,27728,18335,19484,18328,18474,18321,18340,18314]
30386. 2 | 10 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 645,27728,18335,19484,18328,18474,18321,18346,18314]
30388. 2 | 11 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 643,27728,18335,19484,18328,18474,18321,18352,18314]
30390. 2 | 12 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 641,27728,18335,19484,18328,18474,18321,18358,18314]
30392. 2 | 13 | 4 | 18 | 6 | 23 | 8 | 27 [avatar split clause 639,27728,18335,19484,18328,18474,18321,18364,18314]
30394. 28 <=> m2_finseq_1(sK0,k5_numbers) [avatar definition]
30396. m2_finseq_1(sK0,k5_numbers) <- (28) [avatar component clause 30394] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
30397. 28 [avatar split clause 638,30394]
30399. 29 <=> v1_zf_lang(sK0) [avatar definition]
30401. v1_zf_lang(sK0) <- (29) [avatar component clause 30399] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
30402. 29 [avatar split clause 637,30399]
33139. sK0 != sK0 | ~v1_zf_lang(sK7) | ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK8) | ~m2_finseq_1(sK8,k5_numbers) <- (1, 9) [superposition 18312,18342] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
33140. ~v1_zf_lang(sK7) | ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK8) | ~m2_finseq_1(sK8,k5_numbers) <- (1, 9) [trivial inequality removal 33139] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
33141. sK0 != sK0 | ~v1_zf_lang(sK5) | ~m2_finseq_1(sK5,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK6,k5_numbers) <- (3, 14) [superposition 18319,18372] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
33142. ~v1_zf_lang(sK5) | ~m2_finseq_1(sK5,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK6,k5_numbers) <- (3, 14) [trivial inequality removal 33141] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
33408. sK0 != sK0 | ~v1_zf_lang(sK3) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK4) | ~m2_finseq_1(sK4,k5_numbers) <- (5, 19) [superposition 18326,18502] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
33409. ~v1_zf_lang(sK3) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK4) | ~m2_finseq_1(sK4,k5_numbers) <- (5, 19) [trivial inequality removal 33408] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34188. sK0 != X0 | ~v1_zf_lang(sK21(X0)) | ~m2_finseq_1(sK21(X0),k5_numbers) | ~v1_zf_lang(sK22(X0)) | ~m2_finseq_1(sK22(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (5) [superposition 18326,18102] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34189. sK0 != X0 | ~m2_finseq_1(sK21(X0),k5_numbers) | ~v1_zf_lang(sK22(X0)) | ~m2_finseq_1(sK22(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (5) [subsumption resolution 34188,18098] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34190. sK0 != X0 | ~v1_zf_lang(sK22(X0)) | ~m2_finseq_1(sK22(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (5) [subsumption resolution 34189,18099] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34191. sK0 != X0 | ~m2_finseq_1(sK22(X0),k5_numbers) | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (5) [subsumption resolution 34190,18100] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34192. sK0 != X0 | ~v10_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (5) [subsumption resolution 34191,18101] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34193. ~v10_zf_lang(sK0) | ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (5) [equality resolution 34192] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34201. sK0 != X2 | ~m2_subset_1(sK25(X2),k5_numbers,k1_zf_lang) | ~v1_zf_lang(sK26(X2)) | ~m2_finseq_1(sK26(X2),k5_numbers) | ~v11_zf_lang(X2) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) <- (7) [superposition 18333,18113] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34202. sK0 != X2 | ~v1_zf_lang(sK26(X2)) | ~m2_finseq_1(sK26(X2),k5_numbers) | ~v11_zf_lang(X2) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) <- (7) [subsumption resolution 34201,18110] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34203. sK0 != X2 | ~m2_finseq_1(sK26(X2),k5_numbers) | ~v11_zf_lang(X2) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) <- (7) [subsumption resolution 34202,18111] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34204. sK0 != X2 | ~v11_zf_lang(X2) | ~m2_finseq_1(X2,k5_numbers) | ~v1_zf_lang(X2) <- (7) [subsumption resolution 34203,18112] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34207. ~v11_zf_lang(sK0) | ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (7) [equality resolution 34204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34212. ~v11_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (7, 28) [subsumption resolution 34207,30396] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34213. ~v11_zf_lang(sK0) <- (7, 28, 29) [subsumption resolution 34212,30401] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34214. ~8 | ~7 | ~28 | ~29 [avatar split clause 34213,30399,30394,18332,18335]
34215. ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK8) | ~m2_finseq_1(sK8,k5_numbers) <- (1, 9, 13) [subsumption resolution 33140,18366] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34216. ~m2_finseq_1(sK7,k5_numbers) | ~m2_finseq_1(sK8,k5_numbers) <- (1, 9, 11, 13) [subsumption resolution 34215,18354] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34217. ~10 | ~12 | ~1 | ~9 | ~11 | ~13 [avatar split clause 34216,18364,18352,18340,18311,18358,18346]
34785. ~m2_finseq_1(sK0,k5_numbers) | ~m2_finseq_1(sK2,k5_numbers) | ~v1_zf_lang(sK2) | ~m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (24) [superposition 18308,19732] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34793. ~m2_finseq_1(sK0,k5_numbers) | ~m2_finseq_1(sK8,k5_numbers) | ~v1_zf_lang(sK8) | ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK7) | v8_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (9) [superposition 18305,18342] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34800. ~m2_finseq_1(sK8,k5_numbers) | ~v1_zf_lang(sK8) | ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK7) | v8_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (9, 28) [subsumption resolution 34793,30396] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34801. ~v1_zf_lang(sK8) | ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK7) | v8_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (9, 10, 28) [subsumption resolution 34800,18348] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34802. ~m2_finseq_1(sK7,k5_numbers) | ~v1_zf_lang(sK7) | v8_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (9, 10, 11, 28) [subsumption resolution 34801,18354] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34803. ~v1_zf_lang(sK7) | v8_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (9, 10, 11, 12, 28) [subsumption resolution 34802,18360] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34808. ~v1_zf_lang(sK7) | ~v1_zf_lang(sK0) <- (~2, 9, 10, 11, 12, 28) [subsumption resolution 34803,18316] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34809. ~v1_zf_lang(sK7) <- (~2, 9, 10, 11, 12, 28, 29) [subsumption resolution 34808,30401] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34810. ~13 | 2 | ~9 | ~10 | ~11 | ~12 | ~28 | ~29 [avatar split clause 34809,30399,30394,18358,18352,18346,18340,18314,18364]
34825. ~m2_finseq_1(sK5,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK6,k5_numbers) <- (3, 14, 18) [subsumption resolution 33142,18476] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34826. ~m2_finseq_1(sK5,k5_numbers) | ~m2_finseq_1(sK6,k5_numbers) <- (3, 14, 16, 18) [subsumption resolution 34825,18424] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34827. ~15 | ~17 | ~3 | ~14 | ~16 | ~18 [avatar split clause 34826,18474,18422,18370,18318,18448,18396]
34879. ~m2_finseq_1(sK0,k5_numbers) | ~m2_finseq_1(sK4,k5_numbers) | ~v1_zf_lang(sK4) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (19) [superposition 18306,18502] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34887. ~m2_finseq_1(sK0,k5_numbers) | ~m2_finseq_1(sK6,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK5,k5_numbers) | ~v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (14) [superposition 18307,18372] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34894. ~m2_finseq_1(sK6,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK5,k5_numbers) | ~v1_zf_lang(sK5) | v9_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (14, 28) [subsumption resolution 34887,30396] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34915. ~m2_finseq_1(sK2,k5_numbers) | ~v1_zf_lang(sK2) | ~m2_subset_1(sK1,k5_numbers,k1_zf_lang) | v11_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (24, 28) [subsumption resolution 34785,30396] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34929. ~m2_finseq_1(sK4,k5_numbers) | ~v1_zf_lang(sK4) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (19, 28) [subsumption resolution 34879,30396] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34930. ~m2_finseq_1(sK4,k5_numbers) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK3) | v10_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (19, 21, 28) [subsumption resolution 34929,18994] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34931. ~m2_finseq_1(sK4,k5_numbers) | ~m2_finseq_1(sK3,k5_numbers) | v10_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (19, 21, 23, 28) [subsumption resolution 34930,19486] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34932. ~m2_finseq_1(sK4,k5_numbers) | ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK0) <- (~6, 19, 21, 23, 28) [subsumption resolution 34931,18330] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34933. ~m2_finseq_1(sK4,k5_numbers) | ~m2_finseq_1(sK3,k5_numbers) <- (~6, 19, 21, 23, 28, 29) [subsumption resolution 34932,30401] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34934. ~22 | ~20 | 6 | ~19 | ~21 | ~23 | ~28 | ~29 [avatar split clause 34933,30399,30394,19484,18992,18500,18328,18746,19238]
34935. ~m2_finseq_1(sK3,k5_numbers) | ~v1_zf_lang(sK4) | ~m2_finseq_1(sK4,k5_numbers) <- (5, 19, 23) [subsumption resolution 33409,19486] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34936. ~m2_finseq_1(sK3,k5_numbers) | ~m2_finseq_1(sK4,k5_numbers) <- (5, 19, 21, 23) [subsumption resolution 34935,18994] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:20,thDist:-20}
34937. ~20 | ~22 | ~5 | ~19 | ~21 | ~23 [avatar split clause 34936,19484,18992,18500,18325,19238,18746]
34938. ~m2_finseq_1(sK2,k5_numbers) | ~v1_zf_lang(sK2) | v11_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (24, 27, 28) [subsumption resolution 34915,27730] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34939. ~m2_finseq_1(sK2,k5_numbers) | ~v1_zf_lang(sK2) | ~v1_zf_lang(sK0) <- (~8, 24, 27, 28) [subsumption resolution 34938,18337] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34940. ~m2_finseq_1(sK2,k5_numbers) | ~v1_zf_lang(sK2) <- (~8, 24, 27, 28, 29) [subsumption resolution 34939,30401] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34941. ~26 | ~25 | 8 | ~24 | ~27 | ~28 | ~29 [avatar split clause 34940,30399,30394,27728,19730,18335,22396,25062]
34942. ~v10_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (5, 28) [subsumption resolution 34193,30396] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34943. ~v10_zf_lang(sK0) <- (5, 28, 29) [subsumption resolution 34942,30401] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
34944. ~6 | ~5 | ~28 | ~29 [avatar split clause 34943,30399,30394,18325,18328]
34949. ~m2_finseq_1(sK6,k5_numbers) | ~v1_zf_lang(sK6) | ~m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (14, 18, 28) [subsumption resolution 34894,18476] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35044. sK0 != X0 | ~v1_zf_lang(sK19(X0)) | ~m2_finseq_1(sK19(X0),k5_numbers) | ~v1_zf_lang(sK20(X0)) | ~m2_finseq_1(sK20(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (1) [superposition 18312,18096] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35045. sK0 != X0 | ~m2_finseq_1(sK19(X0),k5_numbers) | ~v1_zf_lang(sK20(X0)) | ~m2_finseq_1(sK20(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (1) [subsumption resolution 35044,18092] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35046. sK0 != X0 | ~v1_zf_lang(sK20(X0)) | ~m2_finseq_1(sK20(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (1) [subsumption resolution 35045,18093] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35047. sK0 != X0 | ~m2_finseq_1(sK20(X0),k5_numbers) | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (1) [subsumption resolution 35046,18094] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35048. sK0 != X0 | ~v8_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (1) [subsumption resolution 35047,18095] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35228. ~v8_zf_lang(sK0) | ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (1) [equality resolution 35048] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35229. ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (1, 2) [subsumption resolution 35228,18315] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35230. ~v1_zf_lang(sK0) <- (1, 2, 28) [subsumption resolution 35229,30396] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35231. $false <- (1, 2, 28, 29) [subsumption resolution 35230,30401] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35232. ~1 | ~2 | ~28 | ~29 [avatar contradiction clause 35231]
35233. ~m2_finseq_1(sK6,k5_numbers) | ~m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) | ~v1_zf_lang(sK0) <- (14, 16, 18, 28) [subsumption resolution 34949,18424] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35234. ~m2_finseq_1(sK6,k5_numbers) | ~m2_finseq_1(sK5,k5_numbers) | v9_zf_lang(sK0) <- (14, 16, 18, 28, 29) [subsumption resolution 35233,30401] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35235. 4 | ~17 | ~15 | ~14 | ~16 | ~18 | ~28 | ~29 [avatar split clause 35234,30399,30394,18474,18422,18370,18396,18448,18321]
35248. sK0 != X0 | ~v1_zf_lang(sK23(X0)) | ~m2_finseq_1(sK23(X0),k5_numbers) | ~v1_zf_lang(sK24(X0)) | ~m2_finseq_1(sK24(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (3) [superposition 18319,18108] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35250. sK0 != X0 | ~m2_finseq_1(sK23(X0),k5_numbers) | ~v1_zf_lang(sK24(X0)) | ~m2_finseq_1(sK24(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (3) [subsumption resolution 35248,18104] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35251. sK0 != X0 | ~v1_zf_lang(sK24(X0)) | ~m2_finseq_1(sK24(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (3) [subsumption resolution 35250,18105] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35252. sK0 != X0 | ~m2_finseq_1(sK24(X0),k5_numbers) | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (3) [subsumption resolution 35251,18106] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35253. sK0 != X0 | ~v9_zf_lang(X0) | ~m2_finseq_1(X0,k5_numbers) | ~v1_zf_lang(X0) <- (3) [subsumption resolution 35252,18107] {a:1,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35291. ~v9_zf_lang(sK0) | ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (3) [equality resolution 35253] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35292. ~m2_finseq_1(sK0,k5_numbers) | ~v1_zf_lang(sK0) <- (3, 4) [subsumption resolution 35291,18322] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35293. ~v1_zf_lang(sK0) <- (3, 4, 28) [subsumption resolution 35292,30396] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35294. $false <- (3, 4, 28, 29) [subsumption resolution 35293,30401] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
35295. ~3 | ~4 | ~28 | ~29 [avatar contradiction clause 35294]
35296. $false [avatar sat refutation 18338,18344,18350,18356,18362,18368,18384,18385,18386,18387,18388,18389,18390,18392,18393,18394,18410,18412,18414,18416,18418,18420,18436,18438,18440,18442,18444,18446,18462,18463,18464,18465,18466,18467,18468,18470,18471,18472,18488,18489,18490,18491,18492,18493,18494,18496,18497,18498,18624,18626,18628,18630,18632,18634,18646,18648,18650,18652,18654,18656,18668,18670,18672,18674,18676,18678,18690,18692,18694,18696,18698,18700,18712,18714,18716,18718,18720,18722,18734,18736,18738,18740,18742,18744,18870,18872,18874,18876,18878,18880,18892,18894,18896,18898,18900,18902,18914,18916,18918,18920,18922,18924,18936,18938,18940,18942,18944,18946,18958,18960,18962,18964,18966,18968,18980,18982,18984,18986,18988,18990,19116,19118,19120,19122,19124,19126,19138,19140,19142,19144,19146,19148,19160,19162,19164,19166,19168,19170,19182,19184,19186,19188,19190,19192,19204,19206,19208,19210,19212,19214,19226,19228,19230,19232,19234,19236,19241,19249,19320,19322,19324,19328,19331,19333,19335,19339,19353,19355,19357,19361,19362,19364,19366,19368,19370,19372,19384,19386,19388,19390,19392,19394,19406,19408,19410,19412,19414,19416,19428,19430,19432,19434,19436,19438,19439,19441,19443,19445,19449,19450,19452,19454,19456,19458,19460,19472,19474,19476,19478,19480,19482,19608,19610,19612,19614,19616,19618,19630,19632,19634,19636,19638,19640,19652,19654,19656,19658,19660,19662,19674,19676,19678,19680,19682,19684,19696,19698,19700,19702,19704,19706,19718,19720,19722,19724,19726,19728,21064,21066,21068,21070,21072,21074,21086,21087,21088,21090,21091,21092,21094,21095,21096,21108,21110,21111,21112,21114,21115,21116,21118,21130,21131,21132,21133,21134,21135,21136,21137,21138,21139,21140,21152,21153,21154,21155,21156,21157,21158,21160,21161,21162,21174,21175,21176,21177,21178,21179,21180,21181,21182,21183,21184,21306,21308,21310,21312,21314,21316,21317,21321,21325,21327,21328,21330,21332,21334,21336,21338,21339,21341,21343,21345,21347,21349,21350,21352,21354,21356,21358,21360,21361,21363,21367,21371,21372,21374,21376,21378,21380,21382,21383,21387,21391,21394,21396,21398,21400,21402,21404,21405,21415,21416,21418,21420,21422,21424,21426,21435,21548,21550,21552,21554,21556,21558,21561,21563,21565,21570,21572,21574,21576,21578,21580,21583,21585,21587,21591,21592,21594,21596,21598,21600,21602,21605,21607,21609,21613,21614,21616,21618,21620,21622,21624,21627,21629,21631,21633,21635,21636,21638,21640,21642,21644,21646,21651,21653,21655,21658,21660,21662,21664,21666,21668,21790,21792,21794,21796,21798,21800,21801,21805,21809,21811,21812,21814,21816,21818,21820,21822,21827,21834,21836,21838,21840,21842,21844,21845,21847,21851,21855,21856,21858,21860,21862,21864,21866,21867,21878,21880,21882,21884,21886,21888,21889,21897,21900,21902,21904,21906,21908,21910,21919,22032,22034,22036,22038,22040,22042,22047,22053,22054,22056,22058,22060,22062,22064,22067,22069,22071,22073,22075,22076,22078,22080,22082,22084,22086,22089,22097,22098,22100,22102,22104,22106,22108,22113,22117,22120,22122,22124,22126,22128,22130,22133,22135,22139,22141,22142,22144,22146,22148,22150,22152,22274,22276,22278,22280,22282,22284,22285,22295,22296,22298,22300,22302,22304,22306,22307,22309,22318,22320,22322,22324,22326,22328,22331,22335,22339,22340,22342,22344,22346,22348,22350,22351,22353,22357,22359,22361,22362,22364,22366,22368,22370,22372,22375,22379,22383,22384,22386,22388,22390,22392,22394,23730,23732,23734,23736,23738,23740,23752,23753,23754,23755,23756,23757,23758,23759,23760,23761,23762,23774,23776,23777,23778,23780,23782,23784,23796,23797,23798,23799,23800,23801,23802,23803,23804,23805,23806,23818,23820,23822,23824,23826,23828,23840,23842,23844,23846,23847,23848,23850,23855,23857,23972,23974,23976,23978,23980,23982,23983,23985,23989,23993,23994,23996,23998,24000,24002,24004,24005,24007,24015,24016,24018,24020,24022,24024,24026,24027,24038,24040,24042,24044,24046,24048,24049,24051,24057,24060,24062,24064,24066,24068,24070,24071,24075,24081,24082,24084,24086,24088,24090,24092,24097,24099,24214,24216,24218,24220,24222,24224,24225,24229,24231,24233,24235,24236,24238,24240,24242,24244,24246,24247,24249,24251,24257,24258,24260,24262,24264,24266,24268,24269,24280,24282,24284,24286,24288,24290,24295,24297,24299,24302,24304,24306,24308,24310,24312,24313,24317,24319,24324,24326,24328,24330,24332,24334,24339,24341,24456,24458,24460,24462,24464,24466,24467,24469,24471,24473,24475,24478,24480,24482,24484,24486,24488,24489,24491,24495,24497,24499,24500,24502,24504,24506,24508,24510,24522,24524,24526,24528,24530,24532,24533,24535,24541,24544,24546,24548,24550,24552,24554,24555,24557,24561,24563,24566,24568,24570,24572,24574,24576,24579,24581,24583,24585,24698,24700,24702,24704,24706,24708,24713,24715,24717,24720,24722,24724,24726,24728,24730,24733,24735,24737,24739,24742,24744,24746,24748,24750,24752,24764,24766,24768,24770,24772,24774,24777,24779,24781,24783,24785,24786,24788,24790,24792,24794,24796,24797,24808,24810,24812,24814,24816,24818,24821,24823,24825,24940,24942,24944,24946,24948,24950,24951,24961,24962,24964,24966,24968,24970,24972,24973,24975,24984,24986,24988,24990,24992,24994,24995,25006,25008,25010,25012,25014,25016,25017,25023,25025,25028,25030,25032,25034,25036,25038,25039,25045,25047,25050,25052,25054,25056,25058,25060,26396,26398,26400,26402,26404,26406,26418,26420,26422,26424,26425,26426,26428,26440,26442,26443,26444,26446,26448,26450,26462,26464,26465,26466,26468,26469,26470,26472,26484,26485,26486,26488,26489,26490,26492,26493,26494,26506,26508,26510,26512,26513,26514,26515,26516,26521,26523,26638,26640,26642,26644,26646,26648,26649,26653,26659,26660,26662,26664,26666,26668,26670,26675,26677,26681,26682,26684,26686,26688,26690,26692,26693,26695,26697,26699,26704,26706,26708,26710,26712,26714,26717,26719,26721,26726,26728,26730,26732,26734,26736,26737,26739,26743,26747,26748,26750,26752,26754,26756,26758,26763,26765,26880,26882,26884,26886,26888,26890,26891,26899,26902,26904,26906,26908,26910,26912,26913,26919,26921,26924,26926,26928,26930,26932,26934,26935,26941,26946,26948,26950,26952,26954,26956,26957,26967,26968,26970,26972,26974,26976,26978,26979,26981,26983,26990,26992,26994,26996,26998,27000,27007,27122,27124,27126,27128,27130,27132,27133,27135,27137,27139,27141,27143,27144,27146,27148,27150,27152,27154,27155,27157,27161,27163,27165,27166,27168,27170,27172,27174,27176,27177,27179,27183,27188,27190,27192,27194,27196,27198,27199,27201,27207,27209,27210,27212,27214,27216,27218,27220,27221,27223,27227,27231,27232,27234,27236,27238,27240,27242,27249,27251,27364,27366,27368,27370,27372,27374,27375,27377,27379,27383,27385,27386,27388,27390,27392,27394,27396,27397,27405,27407,27408,27410,27412,27414,27416,27418,27419,27421,27423,27430,27432,27434,27436,27438,27440,27441,27443,27449,27452,27454,27456,27458,27459,27460,27462,27463,27467,27471,27474,27476,27478,27480,27482,27484,27489,27606,27608,27610,27612,27614,27616,27617,27621,27627,27628,27630,27632,27634,27636,27638,27639,27643,27645,27650,27652,27654,27656,27658,27660,27661,27665,27672,27674,27676,27678,27680,27682,27685,27687,27691,27693,27694,27696,27698,27700,27702,27704,27705,27716,27718,27720,27722,27724,27726,29062,29064,29066,29068,29070,29072,29084,29085,29086,29088,29089,29090,29091,29092,29093,29094,29106,29107,29108,29109,29110,29111,29112,29114,29115,29116,29128,29129,29130,29131,29132,29133,29134,29135,29136,29137,29138,29150,29151,29152,29154,29155,29156,29158,29159,29160,29172,29173,29174,29175,29176,29177,29178,29179,29180,29181,29182,29187,29189,29304,29306,29308,29310,29312,29314,29317,29319,29321,29323,29326,29328,29330,29332,29334,29336,29341,29348,29350,29352,29354,29356,29358,29361,29363,29365,29370,29372,29374,29376,29378,29380,29381,29385,29389,29391,29392,29394,29396,29398,29400,29402,29405,29409,29411,29413,29414,29416,29418,29420,29422,29424,29546,29548,29550,29552,29554,29556,29561,29568,29570,29572,29574,29576,29578,29581,29583,29587,29590,29592,29594,29596,29598,29600,29601,29603,29607,29612,29614,29616,29618,29620,29622,29625,29627,29631,29633,29634,29636,29638,29640,29642,29644,29653,29655,29656,29658,29660,29662,29664,29666,29673,29788,29790,29792,29794,29796,29798,29801,29803,29805,29807,29809,29810,29812,29814,29816,29818,29820,29823,29825,29827,29829,29831,29832,29834,29836,29838,29840,29842,29843,29845,29849,29853,29854,29856,29858,29860,29862,29864,29865,29867,29873,29875,29876,29878,29880,29882,29884,29886,29887,29889,29893,29897,29898,29900,29902,29904,29906,29908,29917,30030,30032,30034,30036,30038,30040,30043,30045,30047,30051,30052,30054,30056,30058,30060,30062,30065,30067,30071,30073,30074,30076,30078,30080,30082,30084,30091,30096,30098,30100,30102,30104,30106,30109,30111,30115,30118,30120,30122,30124,30126,30128,30135,30137,30140,30142,30144,30146,30148,30150,30153,30157,30272,30274,30276,30278,30280,30282,30283,30287,30294,30296,30298,30300,30302,30304,30305,30309,30311,30316,30318,30320,30322,30324,30326,30331,30338,30340,30342,30344,30346,30348,30349,30351,30355,30360,30362,30364,30366,30368,30370,30373,30382,30384,30386,30388,30390,30392,30397,30402,34214,34217,34810,34827,34934,34937,34941,34944,35232,35235,35295] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:220904,thDist:-220904}
% SZS output end Proof for zf_lang__t6_zf_lang
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 256
% >>> Preprocessing
% Introduced skolems: 57
% Pure predicates: 5
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 17668
% Generated clauses: 20488
% Active clauses: 651
% Passive clauses: 1347
% Final active clauses: 582
% Final passive clauses: 644
% Main loop iterations started: 690
% 
% >>> Simplifying Inferences
% Duplicate literals: 150
% Trivial inequalities: 6
% Fw subsumption resolutions: 601
% Fw demodulations: 7
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 5391
% Forward subsumptions: 474
% 
% >>> Generating Inferences
% Binary resolution: 1099
% Unit resulting resolution: 5
% Factoring: 1
% Forward superposition: 299
% Backward superposition: 106
% Self superposition: 7
% Equality resolution: 6
% 
% >>> AVATAR
% Split clauses: 12183
% Split components: 96274
% Sat splitting refutations: 16
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12583
% SAT solver unit clauses: 140
% SAT solver binary clauses: 197
% 
% Memory used [KB]: 41065
% Time elapsed: 0.142 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% property evaluation: 0.011 s
% preprocessing: 0.005 s
% SAT solver time: 0.006 s
% minimizing solver time: 0.008 s
% sat proof minimization: 0.013 s
% term sharing: 0.004 s
% congruence closure: 0.010 s ( own 0.002 s ) 
% backward superposition index maintenance: 0.001 s
% splitting component index usage: 0.007 s ( own 0.004 s ) 
% forward subsumption: 0.005 s
% forward subsumption resolution: 0.003 s
% backward subsumption: 0.001 s
% forward demodulation: 0.008 s
% resolution: 0.003 s ( own 0.002 s ) 
% superposition: 0.001 s
% passive container maintenance: 0.001 s
% other: 0.142 s ( own 0.066 s ) 

