i: [1,72,0,2,1,1,0,0] t36_funct_4
a: [3,124,0,2,1,72]
s: 124
i: [1,70,0,2,1,1,0,0] t36_funct_4
a: [3,134,0,2,1,70]
s: 134
i: [1,74,0,3,1,1,0,0] t36_funct_4
a: [3,114,0,3,1,74]
s: 114
i: [1,75,0,5,1,1,0,0] t36_funct_4
a: [3,109,0,5,1,75]
s: 109
i: [1,73,0,6,1,1,0,0] t36_funct_4
a: [3,119,0,6,1,73]
s: 119
i: [1,71,0,6,1,1,0,0] t36_funct_4
a: [3,129,0,6,1,71]
s: 129
i: [1,105,0,2,1,0,0,1] rc1_funct_1
a: [3,146,0,2,1,105]
s: 146
i: [1,104,0,2,1,0,0,1] rc1_funct_1
a: [3,151,0,2,1,104]
s: 151
i: [1,84,0,3,1,0,0,4] reflexivity_r1_tarski
s: 84
i: [1,83,0,4,1,0,0,1] existence_m1_subset_1
s: 83
i: [1,85,0,4,1,0,0,1] fc6_relat_1
s: 85
i: [1,96,0,4,1,0,0,255] rc3_relat_1
s: 96
i: [1,99,0,4,1,0,0,255] rc6_funct_1
s: 99
i: [1,102,0,4,1,0,0,255] rc6_funct_1
s: 102
i: [1,97,0,5,1,0,0,255] rc3_relat_1
s: 97
i: [1,98,0,5,1,0,0,255] rc3_relat_1
s: 98
i: [1,100,0,5,1,0,0,255] rc6_funct_1
s: 100
i: [1,101,0,5,1,0,0,255] rc6_funct_1
s: 101
i: [1,81,0,7,2,0,0,4] t3_subset
s: 81
i: [1,82,0,7,2,0,0,4] t3_subset
s: 82
i: [1,76,0,8,3,0,0,1] cc2_relat_1
s: 76
d: [2,167,1,6,2,1,40,76,119]
d: [2,171,1,2,1,1,21,167,85]
a: [3,175,1,2,1,171]
s: 175
i: [1,93,0,8,2,0,0,1] cc1_relset_1
s: 93
d: [2,168,1,6,2,1,40,76,129]
d: [2,177,1,2,1,1,21,168,85]
a: [3,181,1,2,1,177]
s: 181
i: [1,94,0,9,2,0,0,255] cc2_relset_1
s: 94
i: [1,95,0,9,2,0,0,255] cc2_relset_1
s: 95
d: [2,188,1,3,1,1,40,94,119]
a: [3,195,1,3,1,188]
s: 195
i: [1,86,0,13,5,0,0,1] idempotence_k1_funct_4
a: [3,141,0,9,3,86]
s: 141
d: [2,189,1,3,1,1,40,94,129]
a: [3,200,1,3,1,189]
s: 200
i: [1,77,0,10,4,0,0,1] cc3_funct_1
s: 77
i: [1,90,0,11,3,0,0,2] redefinition_k1_relset_1
s: 90
d: [2,202,1,3,1,1,40,95,119]
a: [3,209,1,3,1,202]
s: 209
i: [1,91,0,11,3,0,0,2] dt_k1_relset_1
s: 91
d: [2,203,1,3,1,1,40,95,129]
a: [3,214,1,3,1,203]
s: 214
i: [1,103,0,11,2,0,0,3] fc8_relset_1
s: 103
i: [1,87,0,12,5,0,0,1] dt_k1_funct_4
s: 87
d: [2,153,1,5,1,1,40,81,119]
a: [3,159,1,5,1,153]
s: 159
i: [1,88,0,12,5,0,0,1] dt_k1_funct_4
s: 88
d: [2,154,1,5,1,1,40,81,129]
a: [3,164,1,5,1,154]
s: 164
i: [1,89,0,12,4,0,0,255] cc6_relat_1
s: 89
i: [1,92,0,12,4,0,0,2] cc5_relat_1
s: 92
d: [2,235,1,12,3,1,40,89,119]
d: [2,240,1,8,2,1,21,235,85]
s: 240
i: [1,79,0,13,4,0,0,1] d2_partfun1
s: 79
d: [2,236,1,12,3,1,40,89,129]
d: [2,241,1,8,2,1,21,236,85]
s: 241
i: [1,80,0,13,4,0,0,1] d2_partfun1
s: 80
i: [1,78,0,18,6,0,0,3] t19_funct_4
s: 78
d: [2,242,1,12,3,1,40,92,119]
d: [2,247,1,8,2,1,21,242,85]
s: 247
d: [2,243,1,12,3,1,40,92,129]
d: [2,248,1,8,2,1,21,243,85]
s: 248
d: [2,155,1,5,1,0,40,81,83]
s: 155
d: [2,170,1,6,2,0,40,76,83]
s: 170
d: [2,186,1,6,1,0,40,93,83]
s: 186
d: [2,169,1,7,3,0,40,76,82]
s: 169
d: [2,185,1,7,2,0,40,93,82]
s: 185
d: [2,191,1,7,1,0,40,94,83]
s: 191
d: [2,205,1,7,1,0,40,95,83]
s: 205
d: [2,190,1,8,2,0,40,94,82]
s: 190
d: [2,269,2,5,1,0,40,190,84]
s: 269
d: [2,204,1,8,2,0,40,95,82]
s: 204
d: [2,275,2,5,1,0,40,204,84]
s: 275
d: [2,219,1,8,3,0,40,77,83]
s: 219
d: [2,218,1,9,4,0,40,77,82]
s: 218
d: [2,230,1,9,1,0,40,103,83]
s: 230
d: [2,222,1,10,3,0,40,91,81]
s: 222
d: [2,291,2,14,5,0,44,222,90]
d: [2,292,2,9,3,0,18,291]
s: 292
d: [2,226,1,15,5,0,44,91,90]
d: [2,227,1,10,3,0,18,226]
s: 227
d: [2,228,1,10,2,0,40,103,82]
s: 228
d: [2,310,2,7,1,0,40,228,84]
s: 310
d: [2,239,1,10,3,0,40,89,83]
s: 239
d: [2,246,1,10,3,0,40,92,83]
s: 246
d: [2,221,1,11,4,0,40,91,76]
s: 221
d: [2,237,1,11,4,0,40,89,82]
s: 237
d: [2,304,2,10,4,0,40,227,76]
s: 304
d: [2,244,1,11,4,0,40,92,82]
s: 244
d: [2,328,3,7,3,1,40,304,195]
d: [2,335,3,5,2,2,21,328,175]
a: [3,339,3,2,1,335]
s: 339
d: [2,251,1,17,6,0,44,90,79]
d: [2,252,1,12,4,0,18,251]
s: 252
d: [2,250,1,17,6,0,44,91,79]
d: [2,253,1,12,4,0,18,250]
s: 253
d: [2,324,3,11,3,0,40,304,269]
d: [2,331,3,7,2,0,21,324,85]
s: 331
d: [2,249,1,17,6,0,44,79,90]
d: [2,254,1,12,4,0,18,249]
s: 254
d: [2,326,3,11,3,0,40,304,97]
d: [2,333,3,7,2,0,21,326,96]
s: 333
d: [2,256,1,17,6,0,44,80,90]
d: [2,257,1,12,4,0,18,256]
s: 257
d: [2,220,1,13,5,0,40,91,77]
s: 220
d: [2,327,3,11,3,0,40,304,100]
d: [2,334,3,7,2,0,21,327,99]
s: 334
d: [2,225,1,13,3,0,40,91,93]
s: 225
d: [2,325,3,15,3,0,40,304,191]
d: [2,332,3,9,2,0,21,325,186]
s: 332
d: [2,223,1,14,3,0,40,91,95]
s: 223
d: [2,224,1,14,3,0,40,91,94]
s: 224
d: [2,308,2,10,3,0,40,227,93]
s: 308
d: [2,238,1,15,5,0,40,89,91]
s: 238
d: [2,413,3,13,2,0,40,308,97]
d: [2,417,3,7,1,0,21,413,96]
s: 417
d: [2,245,1,15,5,0,40,92,91]
s: 245
d: [2,229,1,18,3,0,40,103,91]
s: 229
d: [2,414,3,13,2,0,40,308,100]
d: [2,418,3,7,1,0,21,414,99]
s: 418
d: [2,315,2,15,6,0,44,221,90]
d: [2,316,2,10,4,0,18,315]
s: 316
d: [2,386,2,10,3,0,47,257]
s: 386
d: [2,395,2,17,5,0,44,225,90]
d: [2,396,2,10,3,0,18,395]
s: 396
d: [2,286,2,11,4,0,40,222,169]
s: 286
d: [2,296,3,10,4,0,40,292,169]
s: 296
d: [2,306,2,11,3,0,40,227,95]
s: 306
d: [2,448,3,14,2,0,40,306,269]
d: [2,455,3,8,1,0,21,448,85]
s: 455
d: [2,307,2,11,3,0,40,227,94]
s: 307
d: [2,401,2,18,5,0,44,223,90]
d: [2,402,2,11,3,0,18,401]
s: 402
d: [2,451,3,14,2,0,40,306,97]
d: [2,458,3,8,1,0,21,451,96]
s: 458
d: [2,406,2,18,5,0,44,224,90]
d: [2,407,2,11,3,0,18,406]
s: 407
d: [2,452,3,14,2,0,40,306,100]
d: [2,459,3,8,1,0,21,452,99]
s: 459
d: [2,303,2,12,5,0,40,227,77]
s: 303
d: [2,389,2,17,7,0,44,220,90]
d: [2,390,2,12,5,0,18,389]
s: 390
d: [2,464,3,14,2,0,40,307,269]
d: [2,471,3,8,1,0,21,464,85]
s: 471
d: [2,285,2,13,5,0,40,222,218]
s: 285
d: [2,467,3,14,2,0,40,307,97]
d: [2,474,3,8,1,0,21,467,96]
s: 474
d: [2,289,2,13,3,0,40,222,185]
s: 289
d: [2,309,2,13,3,0,40,227,103]
s: 309
d: [2,468,3,14,2,0,40,307,100]
d: [2,475,3,8,1,0,21,468,99]
s: 475
d: [2,435,2,22,5,0,44,229,90]
d: [2,437,2,13,3,0,18,435]
s: 437
d: [2,482,3,13,4,0,40,303,269]
d: [2,491,3,9,3,0,21,482,85]
s: 491
d: [2,287,2,14,3,0,40,222,204]
s: 287
d: [2,288,2,14,3,0,40,222,190]
s: 288
d: [2,485,3,13,4,0,40,303,97]
d: [2,494,3,9,3,0,21,485,96]
s: 494
d: [2,301,2,14,5,0,40,227,92]
s: 301
d: [2,486,3,13,4,0,40,303,100]
d: [2,495,3,9,3,0,21,486,99]
s: 495
d: [2,302,2,14,5,0,40,227,89]
s: 302
d: [2,421,2,19,7,0,44,238,90]
d: [2,422,2,14,5,0,18,421]
s: 422
d: [2,299,3,10,3,0,40,292,185]
s: 299
d: [2,429,2,19,7,0,44,245,90]
d: [2,430,2,14,5,0,18,429]
s: 430
d: [2,450,3,18,2,0,40,306,191]
d: [2,457,3,10,1,0,21,450,186]
s: 457
d: [2,319,2,15,5,0,40,237,222]
s: 319
d: [2,353,2,15,5,0,40,244,222]
s: 353
d: [2,466,3,18,2,0,40,307,191]
d: [2,473,3,10,1,0,21,466,186]
s: 473
d: [2,359,2,17,6,0,44,227,252]
d: [2,368,2,15,5,0,18,359]
s: 368
d: [2,522,3,18,2,0,40,309,269]
d: [2,531,3,10,1,0,21,522,85]
s: 531
d: [2,358,2,19,7,0,44,252,252]
d: [2,369,2,17,6,0,18,358]
s: 369
d: [2,385,2,19,7,0,44,257,252]
d: [2,387,2,17,6,0,18,385]
s: 387
d: [2,620,3,14,5,1,40,369,200]
d: [2,635,3,12,4,2,21,620,181]
a: [3,639,3,3,1,635]
s: 639
d: [2,311,2,18,3,0,40,228,222]
s: 311
d: [2,619,3,14,5,1,40,369,195]
d: [2,633,3,12,4,2,21,619,175]
d: [2,634,3,9,3,3,21,633,114]
s: 634
d: [2,363,2,28,4,0,44,230,252]
d: [2,370,2,20,3,0,21,363,186]
s: 370
d: [2,362,2,25,9,0,44,78,252]
d: [2,365,2,23,8,0,18,362]
s: 365
d: [2,650,3,21,8,1,44,109,365]
d: [2,655,3,18,7,1,19,650]
d: [2,656,3,16,6,2,21,655,124]
d: [2,657,3,14,5,3,21,656,175]
d: [2,658,3,12,4,4,21,657,134]
d: [2,659,3,10,3,5,21,658,181]
s: 659
d: [2,361,2,25,9,0,44,78,252]
d: [2,366,2,23,8,0,18,361]
s: 366
d: [2,660,4,8,2,5,40,659,84]
a: [3,668,4,4,1,660]
s: 668
d: [2,436,2,37,6,0,44,229,252]
d: [2,439,2,29,5,0,21,436,225]
s: 439
d: [2,684,3,21,8,1,44,109,366]
d: [2,689,3,18,7,1,19,684]
d: [2,690,3,16,6,2,21,689,124]
d: [2,691,3,14,5,3,21,690,175]
d: [2,692,3,12,4,4,21,691,134]
d: [2,693,3,10,3,5,21,692,181]
s: 693
d: [2,706,4,8,2,5,40,693,84]
a: [3,712,4,4,1,706]
s: 712
d: [2,527,3,18,2,0,40,309,97]
d: [2,536,3,10,1,0,21,527,96]
s: 536
d: [2,663,4,12,2,5,40,659,155]
a: [3,681,4,6,1,663]
s: 681
d: [2,528,3,18,2,0,40,309,100]
d: [2,537,3,10,1,0,21,528,99]
s: 537
d: [2,297,3,11,3,0,40,292,204]
s: 297
d: [2,705,4,11,4,5,40,693,292]
d: [2,708,4,9,3,5,21,705,181]
s: 708
d: [2,731,5,0,0,8,52,200,114,195,708]
e: 731
d: [2,732,5,6,2,6,40,708,195]
d: [2,734,5,3,1,7,21,732,114]
d: [2,735,5,0,0,8,21,734,200]
e: 735
d: [2,184,1,2,1,1,40,93,129]
f: 735,189,188,184,171,70,71,72,73,74,75
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_4__t36_funct_4
% SZS output start Proof for funct_4__t36_funct_4
1. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2)) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3)) => (v1_partfun1(X3,X0) => k1_funct_4(X2,X3) = X3))) [input t36_funct_4]
2. ~! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2)) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3)) => (v1_partfun1(X3,X0) => k1_funct_4(X2,X3) = X3))) [negated conjecture 1]
3. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [input cc1_relset_1]
4. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
5. ! [X0,X1] : ! [X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [input cc2_relset_1]
9. ! [X0] : ! [X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0)) [input d2_partfun1]
11. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0))) [input dt_k1_relset_1]
17. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
23. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => k1_relset_1(X0,X1) = k9_xtuple_0(X1)) [input redefinition_k1_relset_1]
25. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : ((v1_funct_1(X1) & v1_relat_1(X1)) => (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) => k1_funct_4(X0,X1) = X1))) [input t19_funct_4]
26. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
27. ~! [X0,X1,X2] : ((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2)) => ! [X3] : ((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3)) => (v1_partfun1(X3,X0) => k1_funct_4(X2,X3) = X3))) [flattening 2]
28. ! [X0,X1] : ((v4_relat_1(X1,X0) & v1_relat_1(X1)) => (v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0)) [flattening 9]
29. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 26]
31. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => v1_relat_1(X2)) [flattening 3]
32. ! [X0,X1,X2] : (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) => (v5_relat_1(X2,X1) & v4_relat_1(X2,X0))) [flattening 5]
33. ? [X0,X1,X2] : (? [X3] : ((k1_funct_4(X2,X3) != X3 & v1_partfun1(X3,X0)) & (m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3))) & (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2))) [ennf transformation 27]
34. ? [X0,X1,X2] : (? [X3] : (k1_funct_4(X2,X3) != X3 & v1_partfun1(X3,X0) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2)) [flattening 33]
35. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
38. ! [X0] : (! [X1] : ((k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | (~v1_funct_1(X1) | ~v1_relat_1(X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 25]
39. ! [X0] : (! [X1] : (k1_funct_4(X0,X1) = X1 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_funct_1(X1) | ~v1_relat_1(X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 38]
40. ! [X0,X1] : ((v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 28]
41. ! [X0,X1] : ((v1_partfun1(X1,X0) <=> k1_relset_1(X0,X1) = X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 40]
48. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 23]
49. ! [X0,X1] : (k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 48]
50. ! [X0,X1] : (m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | (~v4_relat_1(X1,X0) | ~v1_relat_1(X1))) [ennf transformation 11]
51. ! [X0,X1] : (m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [flattening 50]
54. ! [X0,X1,X2] : (v1_relat_1(X2) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 31]
55. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v4_relat_1(X2,X0)) | ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1)))) [ennf transformation 32]
57. ? [X0,X1,X2] : (? [X3] : (k1_funct_4(X2,X3) != X3 & v1_partfun1(X3,X0) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X3)) & m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) & v1_funct_1(X2)) => (? [X3] : (k1_funct_4(sK2,X3) != X3 & v1_partfun1(X3,sK0) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(X3)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(sK2)) [choice axiom]
58. ? [X3] : (k1_funct_4(sK2,X3) != X3 & v1_partfun1(X3,sK0) & m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(X3)) => (sK3 != k1_funct_4(sK2,sK3) & v1_partfun1(sK3,sK0) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(sK3)) [choice axiom]
59. (sK3 != k1_funct_4(sK2,sK3) & v1_partfun1(sK3,sK0) & m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(sK3)) & m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) & v1_funct_1(sK2) [skolemisation 34,58,57]
60. ! [X0,X1] : (((v1_partfun1(X1,X0) | k1_relset_1(X0,X1) != X0) & (k1_relset_1(X0,X1) = X0 | ~v1_partfun1(X1,X0))) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1)) [nnf transformation 41]
61. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 29]
70. v1_funct_1(sK2) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
71. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 59] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
72. v1_funct_1(sK3) [cnf transformation 59] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
73. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [cnf transformation 59] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
74. v1_partfun1(sK3,sK0) [cnf transformation 59] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
75. sK3 != k1_funct_4(sK2,sK3) [cnf transformation 59] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
76. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | k1_funct_4(X0,X1) = X1 | ~v1_funct_1(X1) | ~v1_relat_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 39] {a:0,w:18,wCS:7200,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
79. k1_relset_1(X0,X1) = X0 | ~v1_partfun1(X1,X0) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 60] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. ~m1_subset_1(X0,k1_zfmisc_1(X1)) | r1_tarski(X0,X1) [cnf transformation 61] {a:0,w:7,wCS:2800,nSel:1,sine:4,thAx:0,allAx:1,thDist:-1}
85. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 17] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. k1_relset_1(X0,X1) = k9_xtuple_0(X1) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 49] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
91. m1_subset_1(k1_relset_1(X0,X1),k1_zfmisc_1(X0)) | ~v4_relat_1(X1,X0) | ~v1_relat_1(X1) [cnf transformation 51] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
93. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v1_relat_1(X2) [cnf transformation 54] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
94. ~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X0,X1))) | v4_relat_1(X2,X0) [cnf transformation 55] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
107. 1 <=> sK3 = k1_funct_4(sK2,sK3) [avatar definition]
109. sK3 != k1_funct_4(sK2,sK3) <- (~1) [avatar component clause 107] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. ~1 [avatar split clause 75,107]
112. 2 <=> v1_partfun1(sK3,sK0) [avatar definition]
114. v1_partfun1(sK3,sK0) <- (2) [avatar component clause 112] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
115. 2 [avatar split clause 74,112]
117. 3 <=> m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
119. m1_subset_1(sK3,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (3) [avatar component clause 117] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
120. 3 [avatar split clause 73,117]
122. 4 <=> v1_funct_1(sK3) [avatar definition]
124. v1_funct_1(sK3) <- (4) [avatar component clause 122] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
125. 4 [avatar split clause 72,122]
127. 5 <=> m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) [avatar definition]
129. m1_subset_1(sK2,k1_zfmisc_1(k2_zfmisc_1(sK0,sK1))) <- (5) [avatar component clause 127] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. 5 [avatar split clause 71,127]
132. 6 <=> v1_funct_1(sK2) [avatar definition]
134. v1_funct_1(sK2) <- (6) [avatar component clause 132] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. 6 [avatar split clause 70,132]
167. v1_relat_1(sK3) | ~v1_relat_1(k2_zfmisc_1(sK0,sK1)) <- (3) [resolution 76,119] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
171. v1_relat_1(sK3) <- (3) [subsumption resolution 167,85] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
173. 13 <=> v1_relat_1(sK3) [avatar definition]
175. v1_relat_1(sK3) <- (13) [avatar component clause 173] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
176. 13 | ~3 [avatar split clause 171,117,173]
179. 14 <=> v1_relat_1(sK2) [avatar definition]
181. v1_relat_1(sK2) <- (14) [avatar component clause 179] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
184. v1_relat_1(sK2) <- (5) [resolution 93,129] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
187. 14 | ~5 [avatar split clause 184,127,179]
188. v4_relat_1(sK3,sK0) <- (3) [resolution 94,119] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
189. v4_relat_1(sK2,sK0) <- (5) [resolution 94,129] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
193. 15 <=> v4_relat_1(sK3,sK0) [avatar definition]
195. v4_relat_1(sK3,sK0) <- (15) [avatar component clause 193] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
196. 15 | ~3 [avatar split clause 188,117,193]
198. 16 <=> v4_relat_1(sK2,sK0) [avatar definition]
200. v4_relat_1(sK2,sK0) <- (16) [avatar component clause 198] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
201. 16 | ~5 [avatar split clause 189,127,198]
222. r1_tarski(k1_relset_1(X5,X4),X5) | ~v1_relat_1(X4) | ~v4_relat_1(X4,X5) [resolution 91,81] {a:1,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
251. k9_xtuple_0(X3) = X2 | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) | ~v1_partfun1(X3,X2) | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) [superposition 90,79] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
252. k9_xtuple_0(X3) = X2 | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) | ~v1_partfun1(X3,X2) [duplicate literal removal 251] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
291. r1_tarski(k9_xtuple_0(X3),X2) | ~v1_relat_1(X3) | ~v4_relat_1(X3,X2) | ~v4_relat_1(X3,X2) | ~v1_relat_1(X3) [superposition 222,90] {a:2,w:14,wCS:5600,sine:2,thAx:0,allAx:3,thDist:-3}
292. r1_tarski(k9_xtuple_0(X3),X2) | ~v1_relat_1(X3) | ~v4_relat_1(X3,X2) [duplicate literal removal 291] {a:2,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:3,thDist:-3}
361. ~r1_tarski(k9_xtuple_0(X8),X7) | k1_funct_4(X8,X6) = X6 | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(X8) | ~v1_relat_1(X8) | ~v4_relat_1(X6,X7) | ~v1_relat_1(X6) | ~v1_partfun1(X6,X7) [superposition 78,252] {a:2,w:25,wCS:10000,sine:1,thAx:0,allAx:3,thDist:-3}
366. k1_funct_4(X8,X6) = X6 | ~r1_tarski(k9_xtuple_0(X8),X7) | ~v1_funct_1(X6) | ~v1_relat_1(X6) | ~v1_funct_1(X8) | ~v1_relat_1(X8) | ~v4_relat_1(X6,X7) | ~v1_partfun1(X6,X7) [duplicate literal removal 361] {a:2,w:23,wCS:9200,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
684. sK3 != sK3 | ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1) [superposition 109,366] {a:3,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
689. ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v1_funct_1(sK3) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1) [trivial inequality removal 684] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
690. ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v1_relat_1(sK3) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1, 4) [subsumption resolution 689,124] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
691. ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v1_funct_1(sK2) | ~v1_relat_1(sK2) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1, 4, 13) [subsumption resolution 690,175] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
692. ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v1_relat_1(sK2) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1, 4, 6, 13) [subsumption resolution 691,134] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
693. ~r1_tarski(k9_xtuple_0(sK2),X0) | ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) <- (~1, 4, 6, 13, 14) [subsumption resolution 692,181] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
705. ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) | ~v1_relat_1(sK2) | ~v4_relat_1(sK2,X0) <- (~1, 4, 6, 13, 14) [resolution 693,292] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
708. ~v4_relat_1(sK3,X0) | ~v1_partfun1(sK3,X0) | ~v4_relat_1(sK2,X0) <- (~1, 4, 6, 13, 14) [subsumption resolution 705,181] {a:4,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
732. ~v1_partfun1(sK3,sK0) | ~v4_relat_1(sK2,sK0) <- (~1, 4, 6, 13, 14, 15) [resolution 708,195] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
734. ~v4_relat_1(sK2,sK0) <- (~1, 2, 4, 6, 13, 14, 15) [subsumption resolution 732,114] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
735. $false <- (~1, 2, 4, 6, 13, 14, 15, 16) [subsumption resolution 734,200] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:13,thDist:-13}
736. 1 | ~2 | ~4 | ~6 | ~13 | ~14 | ~15 | ~16 [avatar contradiction clause 735]
737. $false [avatar sat refutation 110,115,120,125,130,135,176,187,196,201,736] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:198,thDist:-198}
% SZS output end Proof for funct_4__t36_funct_4
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 36
% Generated clauses: 550
% Active clauses: 127
% Passive clauses: 265
% Final active clauses: 127
% Final passive clauses: 118
% Main loop iterations started: 147
% 
% >>> Simplifying Inferences
% Duplicate literals: 94
% Trivial inequalities: 2
% Fw subsumption resolutions: 140
% 
% >>> Deletion Inferences
% Simple tautologies: 22
% Equational tautologies: 1
% Forward subsumptions: 68
% 
% >>> Generating Inferences
% Binary resolution: 227
% Unit resulting resolution: 1
% Forward superposition: 59
% Backward superposition: 14
% Self superposition: 1
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 14
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 28
% SAT solver unit clauses: 8
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 10234
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% forward subsumption resolution: 0.001 s
% resolution: 0.001 s
% unit resulting resolution: 0.002 s
% other: 0.008 s ( own 0.003 s ) 

