i: [1,47,0,2,1,1,0,0] t37_xxreal_0
a: [3,77,0,2,1,47]
s: 77
i: [1,46,0,2,1,1,0,0] t37_xxreal_0
a: [3,82,0,2,1,46]
s: 82
i: [1,45,0,2,1,1,0,0] t37_xxreal_0
a: [3,87,0,2,1,45]
s: 87
i: [1,48,0,3,1,1,0,0] t37_xxreal_0
a: [3,72,0,3,1,48]
s: 72
i: [1,63,0,2,1,0,0,1] rc1_xxreal_0
a: [3,107,0,2,1,63]
s: 107
i: [1,49,0,11,1,1,0,0] t37_xxreal_0
a: [3,67,0,11,1,49]
s: 67
i: [1,55,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,94,0,5,2,55]
s: 94
i: [1,58,0,9,3,0,0,1] idempotence_k3_xxreal_0
a: [3,98,0,7,2,58]
s: 98
i: [1,59,0,9,3,0,0,1] idempotence_k4_xxreal_0
a: [3,102,0,7,2,59]
s: 102
i: [1,56,0,8,3,0,0,1] fc5_xxreal_0
s: 56
i: [1,57,0,8,3,0,0,1] fc6_xxreal_0
s: 57
i: [1,62,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 62
i: [1,60,0,11,3,0,0,1] commutativity_k4_xxreal_0
s: 60
i: [1,61,0,11,3,0,0,1] commutativity_k3_xxreal_0
s: 61
i: [1,50,0,12,4,0,0,1] d10_xxreal_0
s: 50
i: [1,51,0,12,4,0,0,1] d10_xxreal_0
s: 51
d: [2,137,1,18,4,1,44,67,50]
d: [2,146,1,16,3,2,21,137,87]
a: [3,150,1,4,1,146]
s: 150
d: [2,181,2,0,0,3,52,77,82,150,56]
e: 181
d: [2,182,2,4,2,1,40,150,56]
d: [2,186,2,2,1,2,21,182,82]
d: [2,187,2,0,0,3,21,186,77]
e: 187
r: 150
i: [1,52,0,12,4,0,0,1] d9_xxreal_0
s: 52
d: [2,195,1,16,4,1,44,67,52]
d: [2,204,1,14,3,2,21,195,82]
d: [2,205,1,12,2,3,21,204,77]
a: [3,209,1,3,1,205]
s: 209
r: 209
i: [1,53,0,12,4,0,0,1] d9_xxreal_0
s: 53
i: [1,54,0,15,6,0,0,1] t2_xxreal_0
s: 54
a: [3,208,1,3,1,205]
s: 208
a: [3,149,1,4,1,146]
s: 149
a: [3,154,1,5,1,146]
s: 154
r: 154
a: [3,153,1,5,1,146]
s: 153
r: 153
s: 154
d: [2,261,2,10,4,1,44,154,52]
d: [2,271,2,7,3,2,21,261,208]
d: [2,272,2,5,2,3,21,271,82]
d: [2,273,2,3,1,4,21,272,77]
a: [3,277,2,3,1,273]
s: 277
r: 208
r: 277
s: 209
d: [2,215,2,7,3,1,40,209,62]
d: [2,216,2,5,2,2,21,215,77]
d: [2,217,2,3,1,3,21,216,82]
a: [3,221,2,3,1,217]
s: 221
a: [3,158,1,7,1,146]
s: 158
d: [2,245,1,12,5,1,40,54,72]
d: [2,250,1,10,4,2,21,245,82]
d: [2,251,1,8,3,3,21,250,87]
s: 251
a: [3,276,2,3,1,273]
s: 276
d: [2,402,3,0,0,6,52,77,209,276,251]
e: 402
d: [2,403,3,0,0,6,52,77,87,82,72,209,276,54]
e: 403
d: [2,404,3,5,2,4,40,276,251]
d: [2,408,3,2,1,5,21,404,209]
d: [2,409,3,0,0,6,21,408,77]
e: 409
r: 276
d: [2,163,1,20,4,1,44,67,51]
d: [2,175,1,18,3,2,21,163,87]
a: [3,179,1,9,1,175]
s: 179
r: 154
r: 179
d: [2,307,2,10,4,1,44,154,53]
d: [2,335,2,8,3,2,21,307,82]
d: [2,336,2,6,2,3,21,335,77]
a: [3,340,2,3,1,336]
s: 340
s: 153
d: [2,453,2,10,4,1,44,153,53]
d: [2,465,2,7,3,2,21,453,340]
d: [2,466,2,4,2,3,21,465,209]
d: [2,467,2,2,1,4,21,466,82]
d: [2,468,2,0,0,5,21,467,77]
e: 468
r: 209
s: 208
s: 276
d: [2,485,3,0,0,6,52,82,77,87,221,340,276,54]
e: 485
r: 340
a: [3,213,1,9,1,205]
s: 213
d: [2,324,3,7,3,1,40,277,62]
d: [2,325,3,5,2,2,21,324,77]
d: [2,326,3,3,1,3,21,325,87]
a: [3,330,3,3,1,326]
s: 330
d: [2,226,1,16,4,1,44,67,53]
d: [2,238,1,14,3,2,21,226,82]
d: [2,239,1,12,2,3,21,238,77]
a: [3,243,1,9,1,239]
s: 243
d: [2,119,1,15,3,1,44,67,61]
d: [2,125,1,13,2,2,21,119,82]
d: [2,126,1,11,1,3,21,125,77]
a: [3,130,1,11,1,126]
s: 130
d: [2,351,2,14,4,1,44,158,53]
d: [2,369,2,12,3,2,21,351,82]
a: [3,373,2,4,1,369]
s: 373
d: [2,569,3,0,0,3,52,87,77,373,57]
e: 569
d: [2,570,3,4,2,1,40,373,57]
d: [2,576,3,2,1,2,21,570,77]
d: [2,577,3,0,0,3,21,576,87]
e: 577
d: [2,572,3,9,4,1,44,373,50]
d: [2,573,3,7,3,1,18,572]
d: [2,579,3,5,2,2,21,573,87]
d: [2,580,3,2,1,3,21,579,276]
d: [2,581,3,0,0,4,21,580,77]
e: 581
r: 373
d: [2,139,1,16,6,0,44,60,50]
d: [2,142,1,12,4,0,18,139]
s: 142
a: [3,381,2,3,1,369]
s: 381
s: 179
d: [2,631,2,14,4,1,44,179,142]
d: [2,639,2,7,3,1,19,631]
d: [2,640,2,5,2,2,21,639,87]
d: [2,641,2,3,1,3,21,640,77]
d: [2,642,2,0,0,4,21,641,330]
e: 642
r: 330
d: [2,138,1,16,6,0,44,60,50]
d: [2,143,1,12,4,0,18,138]
s: 143
a: [3,329,3,3,1,326]
s: 329
d: [2,660,4,0,0,6,52,87,82,77,72,221,329,54]
e: 660
r: 221
d: [2,136,1,16,6,0,44,50,60]
d: [2,144,1,12,4,0,18,136]
s: 144
a: [3,372,2,4,1,369]
s: 372
d: [2,135,1,16,6,0,44,50,60]
d: [2,145,1,12,4,0,18,135]
s: 145
d: [2,166,1,16,6,0,44,60,51]
d: [2,169,1,12,4,0,18,166]
s: 169
d: [2,259,2,11,3,1,40,154,62]
d: [2,264,2,7,2,2,21,259,149]
d: [2,265,2,5,1,3,21,264,87]
a: [3,269,2,5,1,265]
s: 269
d: [2,709,3,10,4,1,44,269,52]
d: [2,714,3,7,3,2,21,709,329]
d: [2,715,3,4,2,3,21,714,208]
d: [2,716,3,2,1,4,21,715,82]
d: [2,717,3,0,0,5,21,716,77]
e: 717
r: 269
d: [2,165,1,16,6,0,44,60,51]
d: [2,170,1,12,4,0,18,165]
s: 170
d: [2,262,2,9,3,1,44,154,61]
d: [2,279,2,7,2,2,21,262,77]
d: [2,280,2,5,1,3,21,279,82]
a: [3,283,2,5,1,280]
s: 283
d: [2,162,1,16,6,0,44,51,60]
d: [2,172,1,12,4,0,18,162]
s: 172
d: [2,161,1,16,6,0,44,51,60]
d: [2,173,1,12,4,0,18,161]
s: 173
a: [3,220,2,3,1,217]
s: 220
d: [2,197,1,16,6,0,44,61,52]
d: [2,200,1,12,4,0,18,197]
s: 200
d: [2,557,2,20,4,1,44,130,51]
d: [2,559,2,18,3,2,21,557,87]
a: [3,563,2,4,1,559]
s: 563
d: [2,777,3,0,0,3,52,82,77,563,56]
e: 777
d: [2,778,3,4,2,1,40,563,56]
d: [2,786,3,2,1,2,21,778,77]
d: [2,787,3,0,0,3,21,786,82]
e: 787
d: [2,779,3,9,4,1,44,563,200]
d: [2,784,3,7,3,1,18,779]
d: [2,789,3,5,2,2,21,784,77]
d: [2,790,3,3,1,3,21,789,82]
d: [2,791,3,0,0,4,21,790,208]
e: 791
d: [2,780,3,9,4,1,44,563,53]
d: [2,783,3,7,3,1,18,780]
d: [2,793,3,5,2,2,21,783,77]
d: [2,794,3,2,1,3,21,793,220]
d: [2,795,3,0,0,4,21,794,82]
e: 795
r: 563
d: [2,196,1,16,6,0,44,61,52]
d: [2,201,1,12,4,0,18,196]
s: 201
d: [2,194,1,16,6,0,44,52,61]
d: [2,202,1,12,4,0,18,194]
s: 202
a: [3,562,2,4,1,559]
s: 562
d: [2,193,1,16,6,0,44,52,61]
d: [2,203,1,12,4,0,18,193]
s: 203
d: [2,349,2,12,4,1,44,158,51]
d: [2,355,2,10,3,2,21,349,77]
d: [2,356,2,8,2,3,21,355,87]
a: [3,360,2,5,1,356]
s: 360
d: [2,229,1,16,6,0,44,61,53]
d: [2,232,1,12,4,0,18,229]
s: 232
d: [2,228,1,16,6,0,44,61,53]
d: [2,233,1,12,4,0,18,228]
s: 233
d: [2,802,3,10,4,1,44,360,52]
d: [2,805,3,7,3,2,21,802,208]
d: [2,806,3,5,2,3,21,805,82]
d: [2,807,3,3,1,4,21,806,77]
a: [3,811,3,3,1,807]
s: 811
d: [2,225,1,16,6,0,44,53,61]
d: [2,235,1,12,4,0,18,225]
s: 235
d: [2,350,2,12,4,1,44,158,50]
d: [2,362,2,10,3,2,21,350,77]
d: [2,363,2,8,2,3,21,362,87]
a: [3,367,2,5,1,363]
s: 367
d: [2,869,3,0,0,4,52,87,82,72,367,52]
e: 869
d: [2,873,3,10,4,1,44,367,52]
d: [2,874,3,7,3,1,19,873]
d: [2,879,3,4,2,2,21,874,72]
d: [2,880,3,2,1,3,21,879,82]
d: [2,881,3,0,0,4,21,880,87]
e: 881
r: 179
r: 208
r: 220
r: 276
r: 329
r: 367
s: 209
a: [3,178,1,9,1,175]
s: 178
r: 130
r: 243
r: 213
r: 67
s: 221
d: [2,923,0,9,1,2,23,67,178]
a: [3,950,0,9,1,923]
s: 950
s: 277
d: [2,920,1,9,1,2,23,130,178]
a: [3,943,1,9,1,920]
s: 943
d: [2,1013,1,13,3,1,44,950,61]
d: [2,1029,1,11,2,2,21,1013,77]
d: [2,1030,1,9,1,3,21,1029,82]
a: [3,1034,1,9,1,1030]
s: 1034
a: [3,339,2,3,1,336]
s: 339
d: [2,224,1,16,6,0,44,53,61]
d: [2,236,1,12,4,0,18,224]
s: 236
s: 330
d: [2,247,1,19,8,0,40,54,62]
d: [2,248,1,15,6,0,18,247]
s: 248
a: [3,268,2,5,1,265]
s: 268
d: [2,1128,3,10,4,1,44,268,200]
d: [2,1133,3,7,3,2,21,1128,72]
d: [2,1134,3,5,2,3,21,1133,77]
d: [2,1135,3,3,1,4,21,1134,82]
d: [2,1136,3,0,0,5,21,1135,221]
e: 1136
d: [2,1129,3,10,4,1,44,268,53]
d: [2,1138,3,7,3,2,21,1129,72]
d: [2,1139,3,4,2,3,21,1138,209]
d: [2,1140,3,2,1,4,21,1139,82]
d: [2,1141,3,0,0,5,21,1140,77]
e: 1141
d: [2,306,2,11,3,1,40,154,62]
d: [2,311,2,7,2,2,21,306,149]
d: [2,312,2,5,1,3,21,311,87]
d: [2,413,2,14,4,1,44,179,51]
d: [2,419,2,7,3,1,19,413]
d: [2,420,2,5,2,2,21,419,77]
d: [2,421,2,3,1,3,21,420,87]
d: [2,455,2,9,3,1,44,153,61]
d: [2,470,2,7,2,2,21,455,77]
d: [2,471,2,5,1,3,21,470,82]
d: [2,586,2,20,4,1,44,67,142]
d: [2,602,2,18,3,2,21,586,87]
d: [2,603,2,14,2,3,21,602,149]
d: [2,752,3,10,4,1,44,283,200]
d: [2,895,3,8,3,2,21,752,82]
d: [2,896,3,6,2,3,21,895,77]
f: 1141,896,881,717,641,603,471,421,409,363,312,187,146,45,46,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t37_xxreal_0
% SZS output start Proof for xxreal_0__t37_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r1_xxreal_0(X0,X1) => k4_xxreal_0(X0,k3_xxreal_0(X2,X1)) = k3_xxreal_0(k4_xxreal_0(X0,X2),X1))))) [input t37_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (r1_xxreal_0(X0,X1) => k4_xxreal_0(X0,k3_xxreal_0(X2,X1)) = k3_xxreal_0(k4_xxreal_0(X0,X2),X1))))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0)) [input commutativity_k3_xxreal_0]
4. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0)) [input commutativity_k4_xxreal_0]
5. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
6. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X1,X0) => k4_xxreal_0(X0,X1) = X0)))) [input d10_xxreal_0]
7. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X0)))) [input d9_xxreal_0]
10. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k3_xxreal_0(X0,X1))) [input fc5_xxreal_0]
16. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
17. ? [X0] : (? [X1] : (? [X2] : ((k4_xxreal_0(X0,k3_xxreal_0(X2,X1)) != k3_xxreal_0(k4_xxreal_0(X0,X2),X1) & r1_xxreal_0(X0,X1)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
18. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_0(X0,k3_xxreal_0(X2,X1)) != k3_xxreal_0(k4_xxreal_0(X0,X2),X1) & r1_xxreal_0(X0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 17]
19. ! [X0] : (! [X1] : (((k4_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X1,X0)) & (k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 6]
20. ! [X0] : (! [X1] : (((k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1)) & (k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 7]
21. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 16]
22. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 21]
25. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 10]
26. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 25]
33. ! [X0,X1] : (k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 4]
34. ! [X0,X1] : (k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 33]
35. ! [X0,X1] : (k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
36. ! [X0,X1] : (k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 35]
37. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 5]
38. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 37]
39. ? [X0] : (? [X1] : (? [X2] : (k4_xxreal_0(X0,k3_xxreal_0(X2,X1)) != k3_xxreal_0(k4_xxreal_0(X0,X2),X1) & r1_xxreal_0(X0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (k4_xxreal_0(sK0,k3_xxreal_0(X2,X1)) != k3_xxreal_0(k4_xxreal_0(sK0,X2),X1) & r1_xxreal_0(sK0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
40. ? [X1] : (? [X2] : (k4_xxreal_0(sK0,k3_xxreal_0(X2,X1)) != k3_xxreal_0(k4_xxreal_0(sK0,X2),X1) & r1_xxreal_0(sK0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (k4_xxreal_0(sK0,k3_xxreal_0(X2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,X2),sK1) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
41. ? [X2] : (k4_xxreal_0(sK0,k3_xxreal_0(X2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,X2),sK1) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(X2)) => (k4_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(sK2)) [choice axiom]
42. ((k4_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 18,41,40,39]
45. v1_xxreal_0(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
46. v1_xxreal_0(sK1) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
47. v1_xxreal_0(sK2) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
48. r1_xxreal_0(sK0,sK1) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
49. k4_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) [cnf transformation 42] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
50. k4_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 19] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k4_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 19] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 20] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 20] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 22] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. k4_xxreal_0(X0,X1) = k4_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 34] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k3_xxreal_0(X0,X1) = k3_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 36] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 38] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. 1 <=> k4_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) = k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) [avatar definition]
67. k4_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) <- (~1) [avatar component clause 65] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
68. ~1 [avatar split clause 49,65]
70. 2 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
72. r1_xxreal_0(sK0,sK1) <- (2) [avatar component clause 70] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
73. 2 [avatar split clause 48,70]
75. 3 <=> v1_xxreal_0(sK2) [avatar definition]
77. v1_xxreal_0(sK2) <- (3) [avatar component clause 75] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
78. 3 [avatar split clause 47,75]
80. 4 <=> v1_xxreal_0(sK1) [avatar definition]
82. v1_xxreal_0(sK1) <- (4) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
83. 4 [avatar split clause 46,80]
85. 5 <=> v1_xxreal_0(sK0) [avatar definition]
87. v1_xxreal_0(sK0) <- (5) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
88. 5 [avatar split clause 45,85]
137. sK0 != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) | ~r1_xxreal_0(k3_xxreal_0(sK2,sK1),sK0) | ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) | ~v1_xxreal_0(sK0) <- (~1) [superposition 67,50] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
139. k4_xxreal_0(X4,X3) = X3 | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) | ~r1_xxreal_0(X4,X3) | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) [superposition 60,50] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
142. k4_xxreal_0(X4,X3) = X3 | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) | ~r1_xxreal_0(X4,X3) [duplicate literal removal 139] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
146. sK0 != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) | ~r1_xxreal_0(k3_xxreal_0(sK2,sK1),sK0) | ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) <- (~1, 5) [subsumption resolution 137,87] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
148. 12 <=> v1_xxreal_0(k3_xxreal_0(sK2,sK1)) [avatar definition]
149. v1_xxreal_0(k3_xxreal_0(sK2,sK1)) <- (12) [avatar component clause 148] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
150. ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) <- (~12) [avatar component clause 148] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
152. 13 <=> r1_xxreal_0(k3_xxreal_0(sK2,sK1),sK0) [avatar definition]
153. r1_xxreal_0(k3_xxreal_0(sK2,sK1),sK0) <- (13) [avatar component clause 152] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
154. ~r1_xxreal_0(k3_xxreal_0(sK2,sK1),sK0) <- (~13) [avatar component clause 152] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
156. 14 <=> sK0 = k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) [avatar definition]
158. sK0 != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) <- (~14) [avatar component clause 156] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
159. ~12 | ~13 | ~14 | 1 | ~5 [avatar split clause 146,85,65,156,152,148]
177. 15 <=> k3_xxreal_0(sK2,sK1) = k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) [avatar definition]
179. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) <- (~15) [avatar component clause 177] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
182. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (~12) [resolution 150,56] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
186. ~v1_xxreal_0(sK2) <- (4, ~12) [subsumption resolution 182,82] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
187. $false <- (3, 4, ~12) [subsumption resolution 186,77] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. ~3 | ~4 | 12 [avatar contradiction clause 187]
197. k3_xxreal_0(X4,X3) = X3 | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) | ~r1_xxreal_0(X3,X4) | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) [superposition 61,52] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
200. k3_xxreal_0(X4,X3) = X3 | ~v1_xxreal_0(X4) | ~v1_xxreal_0(X3) | ~r1_xxreal_0(X3,X4) [duplicate literal removal 197] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
207. 16 <=> r1_xxreal_0(sK2,sK1) [avatar definition]
208. r1_xxreal_0(sK2,sK1) <- (16) [avatar component clause 207] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
209. ~r1_xxreal_0(sK2,sK1) <- (~16) [avatar component clause 207] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
245. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (2) [resolution 54,72] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
250. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(X0) <- (2, 4) [subsumption resolution 245,82] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
251. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK1) | ~v1_xxreal_0(X0) <- (2, 4, 5) [subsumption resolution 250,87] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
267. 20 <=> r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) [avatar definition]
268. ~r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (~20) [avatar component clause 267] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
269. r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (20) [avatar component clause 267] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
275. 21 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
276. r1_xxreal_0(sK2,sK0) <- (21) [avatar component clause 275] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
282. 22 <=> r1_xxreal_0(k3_xxreal_0(sK1,sK2),sK0) [avatar definition]
283. r1_xxreal_0(k3_xxreal_0(sK1,sK2),sK0) <- (22) [avatar component clause 282] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
306. r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) | ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) | ~v1_xxreal_0(sK0) <- (~13) [resolution 154,62] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
311. r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) | ~v1_xxreal_0(sK0) <- (12, ~13) [subsumption resolution 306,149] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
312. r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (5, 12, ~13) [subsumption resolution 311,87] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
313. 20 | ~5 | ~12 | 13 [avatar split clause 312,152,148,85,267]
328. 23 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
329. ~r1_xxreal_0(sK0,sK2) <- (~23) [avatar component clause 328] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
350. sK0 != k3_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~14) [superposition 158,50] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
362. sK0 != k3_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (3, ~14) [subsumption resolution 350,77] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
363. sK0 != k3_xxreal_0(sK0,sK1) | ~r1_xxreal_0(sK2,sK0) <- (3, 5, ~14) [subsumption resolution 362,87] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
365. 26 <=> sK0 = k3_xxreal_0(sK0,sK1) [avatar definition]
367. sK0 != k3_xxreal_0(sK0,sK1) <- (~26) [avatar component clause 365] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
368. ~21 | ~26 | ~3 | ~5 | 14 [avatar split clause 363,156,85,75,365,275]
404. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) <- (2, 4, 5, 21) [resolution 276,251] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
408. ~v1_xxreal_0(sK2) <- (2, 4, 5, ~16, 21) [subsumption resolution 404,209] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
409. $false <- (2, 3, 4, 5, ~16, 21) [subsumption resolution 408,77] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
410. ~2 | ~3 | ~4 | ~5 | 16 | ~21 [avatar contradiction clause 409]
413. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(sK2,sK1) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~15) [superposition 179,51] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
419. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~15) [trivial inequality removal 413] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
420. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (3, ~15) [subsumption resolution 419,77] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
421. r1_xxreal_0(sK2,sK0) <- (3, 5, ~15) [subsumption resolution 420,87] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
422. 21 | ~3 | ~5 | 15 [avatar split clause 421,177,85,75,275]
455. r1_xxreal_0(k3_xxreal_0(sK1,sK2),sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (13) [superposition 153,61] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
470. r1_xxreal_0(k3_xxreal_0(sK1,sK2),sK0) | ~v1_xxreal_0(sK1) <- (3, 13) [subsumption resolution 455,77] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
471. r1_xxreal_0(k3_xxreal_0(sK1,sK2),sK0) <- (3, 4, 13) [subsumption resolution 470,82] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
481. 22 | ~3 | ~4 | ~13 [avatar split clause 471,152,80,75,282]
586. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) | ~r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (~1) [superposition 67,142] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
602. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) | ~v1_xxreal_0(k3_xxreal_0(sK2,sK1)) | ~r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (~1, 5) [subsumption resolution 586,87] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
603. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(k4_xxreal_0(sK0,sK2),sK1) | ~r1_xxreal_0(sK0,k3_xxreal_0(sK2,sK1)) <- (~1, 5, 12) [subsumption resolution 602,149] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
604. ~20 | ~15 | 1 | ~5 | ~12 [avatar split clause 603,148,85,65,177,267]
631. k3_xxreal_0(sK2,sK1) != k3_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | ~r1_xxreal_0(sK0,sK2) <- (~15) [superposition 179,142] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
639. ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) | ~r1_xxreal_0(sK0,sK2) <- (~15) [trivial inequality removal 631] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
640. ~v1_xxreal_0(sK2) | ~r1_xxreal_0(sK0,sK2) <- (5, ~15) [subsumption resolution 639,87] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
641. ~r1_xxreal_0(sK0,sK2) <- (3, 5, ~15) [subsumption resolution 640,77] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
651. ~23 | ~3 | ~5 | 15 [avatar split clause 641,177,85,75,328]
709. r1_xxreal_0(sK0,sK2) | ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (20) [superposition 269,52] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
714. ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (20, ~23) [subsumption resolution 709,329] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
715. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (16, 20, ~23) [subsumption resolution 714,208] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
716. ~v1_xxreal_0(sK2) <- (4, 16, 20, ~23) [subsumption resolution 715,82] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
717. $false <- (3, 4, 16, 20, ~23) [subsumption resolution 716,77] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
718. ~3 | ~4 | ~16 | ~20 | 23 [avatar contradiction clause 717]
752. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) | ~r1_xxreal_0(sK2,sK1) <- (22) [superposition 283,200] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
873. sK0 != sK0 | ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~26) [superposition 367,52] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
874. ~r1_xxreal_0(sK0,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~26) [trivial inequality removal 873] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
879. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (2, ~26) [subsumption resolution 874,72] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
880. ~v1_xxreal_0(sK0) <- (2, 4, ~26) [subsumption resolution 879,82] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
881. $false <- (2, 4, 5, ~26) [subsumption resolution 880,87] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
882. ~2 | ~4 | ~5 | 26 [avatar contradiction clause 881]
895. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~r1_xxreal_0(sK2,sK1) <- (4, 22) [subsumption resolution 752,82] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
896. r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK2,sK1) <- (3, 4, 22) [subsumption resolution 895,77] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
897. ~16 | 21 | ~3 | ~4 | ~22 [avatar split clause 896,282,80,75,275,207]
1129. ~r1_xxreal_0(sK0,sK1) | r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (~20) [superposition 268,53] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1138. r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (2, ~20) [subsumption resolution 1129,72] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1139. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK2) <- (2, ~16, ~20) [subsumption resolution 1138,209] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1140. ~v1_xxreal_0(sK2) <- (2, 4, ~16, ~20) [subsumption resolution 1139,82] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1141. $false <- (2, 3, 4, ~16, ~20) [subsumption resolution 1140,77] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1142. ~2 | ~3 | ~4 | 16 | 20 [avatar contradiction clause 1141]
1147. $false [avatar sat refutation 68,73,78,83,88,159,188,313,368,410,422,481,604,651,718,882,897,1142] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1741,thDist:-1741}
% SZS output end Proof for xxreal_0__t37_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 837
% Active clauses: 69
% Passive clauses: 132
% Final active clauses: 43
% Final passive clauses: 38
% Main loop iterations started: 81
% 
% >>> Simplifying Inferences
% Duplicate literals: 140
% Trivial inequalities: 11
% Fw subsumption resolutions: 324
% Fw demodulations: 3
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 31
% Forward subsumptions: 97
% 
% >>> Generating Inferences
% Binary resolution: 45
% Unit resulting resolution: 8
% Factoring: 1
% Forward superposition: 149
% Backward superposition: 82
% 
% >>> AVATAR
% Split clauses: 75
% Split components: 163
% Sat splitting refutations: 21
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 165
% SAT solver unit clauses: 6
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 5500
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% splitting component index usage: 0.001 s
% forward subsumption resolution: 0.002 s
% unit resulting resolution: 0.001 s ( own 0.0000 s ) 
% other: 0.006 s ( own 0.002 s ) 

