i: [1,33,0,17,6,1,0,0] t24_gate_1
a: [3,52,0,2,1,33]
s: 52
a: [3,47,0,7,1,33]
s: 47
i: [1,34,0,2,1,0,0,255] fc1_xboole_0
a: [3,78,0,2,1,34]
s: 78
i: [1,43,0,2,1,0,0,1] rc2_xboole_0
a: [3,83,0,2,1,43]
s: 83
i: [1,44,0,2,1,0,0,1] rc1_xboole_0
a: [3,88,0,2,1,44]
s: 88
i: [1,35,0,5,2,0,0,255] fc2_gate_1
s: 35
i: [1,36,0,5,2,0,0,255] fc1_gate_1
s: 36
i: [1,38,0,10,2,0,0,255] d18_gate_1
s: 38
d: [2,91,1,4,2,1,44,47,38]
d: [2,92,1,2,1,2,21,91,78]
d: [2,93,1,0,0,3,21,92,52]
e: 93
r: 52
a: [3,51,0,2,1,33]
s: 51
a: [3,56,0,2,1,33]
s: 56
i: [1,39,0,10,2,0,0,255] d18_gate_1
s: 39
d: [2,96,1,4,2,1,44,47,39]
d: [2,97,1,2,1,2,21,96,78]
d: [2,98,1,0,0,3,21,97,56]
e: 98
r: 56
a: [3,55,0,2,1,33]
s: 55
a: [3,60,0,2,1,33]
s: 60
i: [1,40,0,10,2,0,0,255] d18_gate_1
s: 40
d: [2,101,1,4,2,1,44,47,40]
d: [2,102,1,2,1,2,21,101,78]
d: [2,103,1,0,0,3,21,102,60]
e: 103
r: 60
a: [3,59,0,2,1,33]
s: 59
a: [3,64,0,2,1,33]
s: 64
i: [1,41,0,10,2,0,0,255] d18_gate_1
s: 41
d: [2,106,1,4,2,1,44,47,41]
d: [2,107,1,2,1,2,21,106,78]
d: [2,108,1,0,0,3,21,107,64]
e: 108
r: 64
a: [3,63,0,2,1,33]
s: 63
a: [3,68,0,2,1,33]
s: 68
i: [1,42,0,10,2,0,0,255] d18_gate_1
s: 42
d: [2,111,1,4,2,1,44,47,42]
d: [2,112,1,2,1,2,21,111,78]
d: [2,113,1,0,0,3,21,112,68]
e: 113
r: 47
r: 68
a: [3,67,0,2,1,33]
s: 67
a: [3,48,0,7,1,33]
s: 48
i: [1,37,0,19,6,0,0,255] d18_gate_1
s: 37
d: [2,130,1,13,6,1,44,48,37]
d: [2,131,1,11,5,2,21,130,67]
d: [2,132,1,9,4,3,21,131,63]
d: [2,133,1,7,3,4,21,132,59]
d: [2,134,1,5,2,5,21,133,55]
d: [2,135,1,3,1,6,21,134,51]
a: [3,139,1,3,1,135]
s: 139
d: [2,141,2,0,0,2,52,78,139,36]
e: 141
d: [2,142,2,2,1,1,40,139,36]
d: [2,144,2,0,0,2,21,142,78]
e: 144
i: [1,32,0,9,2,1,0,0] t24_gate_1
i: [1,31,0,9,2,1,0,0] t24_gate_1
i: [1,30,0,9,2,1,0,0] t24_gate_1
i: [1,29,0,9,2,1,0,0] t24_gate_1
i: [1,28,0,9,2,1,0,0] t24_gate_1
f: 144,135,113,107,102,97,92,34,28,29,30,31,32,33
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t24_gate_1
% SZS output start Proof for gate_1__t24_gate_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)) <=> (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input t24_gate_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)) <=> (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
9. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
13. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0) & ~(k18_gate_1(X0,X1,X2,X3,X4) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [input d18_gate_1]
14. ~! [X0,X1,X2,X3,X4] : (~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)) <=> (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 2]
15. ! [X0,X1,X2,X3,X4] : ((~(~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0) & ~(k18_gate_1(X0,X1,X2,X3,X4) != k1_gate_1(k1_xboole_0) & ~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [flattening 13]
16. ? [X0,X1,X2,X3,X4] : (~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)) <~> (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) [ennf transformation 14]
18. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 8]
19. ! [X0,X1,X2,X3,X4] : ((k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | (~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & (k18_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 15]
20. ? [X0,X1,X2,X3,X4] : (((v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) | v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4))) & ((~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)))) [nnf transformation 16]
21. ? [X0,X1,X2,X3,X4] : ((v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4))) & ((~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)))) [flattening 20]
22. ? [X0,X1,X2,X3,X4] : ((v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) | v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4))) & ((~v1_xboole_0(X4) & ~v1_xboole_0(X3) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) | ~v1_xboole_0(k18_gate_1(X0,X1,X2,X3,X4)))) => ((v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4))) & ((~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)))) [choice axiom]
23. (v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4))) & ((~v1_xboole_0(sK4) & ~v1_xboole_0(sK3) & ~v1_xboole_0(sK2) & ~v1_xboole_0(sK1) & ~v1_xboole_0(sK0)) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4))) [skolemisation 21,22]
28. ~v1_xboole_0(sK0) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
29. ~v1_xboole_0(sK1) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
30. ~v1_xboole_0(sK2) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. ~v1_xboole_0(sK3) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. ~v1_xboole_0(sK4) | ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) | v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [cnf transformation 23] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
34. v1_xboole_0(k1_xboole_0) [cnf transformation 9] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
36. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 18] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
37. k18_gate_1(X0,X1,X2,X3,X4) = k1_gate_1(k1_xboole_0) | v1_xboole_0(X4) | v1_xboole_0(X3) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 19] {a:0,w:19,wCS:7600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
38. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X0) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
39. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X1) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
40. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X2) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
41. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X3) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
42. k18_gate_1(X0,X1,X2,X3,X4) = k1_xboole_0 | ~v1_xboole_0(X4) [cnf transformation 19] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
46. 1 <=> v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) [avatar definition]
47. ~v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) <- (~1) [avatar component clause 46] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
48. v1_xboole_0(k18_gate_1(sK0,sK1,sK2,sK3,sK4)) <- (1) [avatar component clause 46] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
50. 2 <=> v1_xboole_0(sK0) [avatar definition]
51. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
54. 3 <=> v1_xboole_0(sK1) [avatar definition]
55. ~v1_xboole_0(sK1) <- (~3) [avatar component clause 54] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
58. 4 <=> v1_xboole_0(sK2) [avatar definition]
59. ~v1_xboole_0(sK2) <- (~4) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
62. 5 <=> v1_xboole_0(sK3) [avatar definition]
63. ~v1_xboole_0(sK3) <- (~5) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. 6 <=> v1_xboole_0(sK4) [avatar definition]
67. ~v1_xboole_0(sK4) <- (~6) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
68. v1_xboole_0(sK4) <- (6) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
69. 1 | 2 | 3 | 4 | 5 | 6 [avatar split clause 33,66,62,58,54,50,46]
70. ~1 | ~6 [avatar split clause 32,66,46]
71. ~1 | ~5 [avatar split clause 31,62,46]
72. ~1 | ~4 [avatar split clause 30,58,46]
73. ~1 | ~3 [avatar split clause 29,54,46]
74. ~1 | ~2 [avatar split clause 28,50,46]
76. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
78. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 76] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
79. 7 [avatar split clause 34,76]
91. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK0) <- (~1) [superposition 47,38] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
92. ~v1_xboole_0(sK0) <- (~1, 7) [subsumption resolution 91,78] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
95. ~2 | 1 | ~7 [avatar split clause 92,76,46,50]
96. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK1) <- (~1) [superposition 47,39] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
97. ~v1_xboole_0(sK1) <- (~1, 7) [subsumption resolution 96,78] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
100. ~3 | 1 | ~7 [avatar split clause 97,76,46,54]
101. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK2) <- (~1) [superposition 47,40] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. ~v1_xboole_0(sK2) <- (~1, 7) [subsumption resolution 101,78] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. ~4 | 1 | ~7 [avatar split clause 102,76,46,58]
106. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK3) <- (~1) [superposition 47,41] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
107. ~v1_xboole_0(sK3) <- (~1, 7) [subsumption resolution 106,78] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
110. ~5 | 1 | ~7 [avatar split clause 107,76,46,62]
111. ~v1_xboole_0(k1_xboole_0) | ~v1_xboole_0(sK4) <- (~1) [superposition 47,42] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
112. ~v1_xboole_0(sK4) <- (~1, 7) [subsumption resolution 111,78] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
113. $false <- (~1, 6, 7) [subsumption resolution 112,68] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
114. 1 | ~6 | ~7 [avatar contradiction clause 113]
130. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK4) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1) [superposition 48,37] {a:1,w:13,wCS:1300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
131. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~6) [subsumption resolution 130,67] {a:1,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK2) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~5, ~6) [subsumption resolution 131,63] {a:1,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
133. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (1, ~4, ~5, ~6) [subsumption resolution 132,59] {a:1,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. v1_xboole_0(k1_gate_1(k1_xboole_0)) | v1_xboole_0(sK0) <- (1, ~3, ~4, ~5, ~6) [subsumption resolution 133,55] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (1, ~2, ~3, ~4, ~5, ~6) [subsumption resolution 134,51] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
137. 10 <=> v1_xboole_0(k1_gate_1(k1_xboole_0)) [avatar definition]
139. v1_xboole_0(k1_gate_1(k1_xboole_0)) <- (10) [avatar component clause 137] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
140. 10 | ~1 | 2 | 3 | 4 | 5 | 6 [avatar split clause 135,66,62,58,54,50,46,137]
142. ~v1_xboole_0(k1_xboole_0) <- (10) [resolution 139,36] {a:2,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
144. $false <- (7, 10) [subsumption resolution 142,78] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
145. ~7 | ~10 [avatar contradiction clause 144]
146. $false [avatar sat refutation 69,70,71,72,73,74,79,95,100,105,110,114,140,145] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:94,thDist:-94}
% SZS output end Proof for gate_1__t24_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 77
% Active clauses: 24
% Passive clauses: 24
% Final active clauses: 18
% Main loop iterations started: 24
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 16
% 
% >>> Deletion Inferences
% Simple tautologies: 11
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 2
% Unit resulting resolution: 1
% Forward superposition: 10
% Backward superposition: 11
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 16
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 3
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

