i: [1,236,0,2,1,1,0,0] t49_bciideal
a: [3,2844,0,2,1,236]
s: 2844
i: [1,235,0,2,1,1,0,0] t49_bciideal
a: [3,2849,0,2,1,235]
s: 2849
i: [1,234,0,2,1,1,0,0] t49_bciideal
a: [3,2854,0,2,1,234]
s: 2854
i: [1,233,0,2,1,1,0,0] t49_bciideal
a: [3,2859,0,2,1,233]
s: 2859
i: [1,232,0,2,1,1,0,0] t49_bciideal
a: [3,2864,0,2,1,232]
s: 2864
i: [1,231,0,2,1,1,0,0] t49_bciideal
a: [3,2869,0,2,1,231]
s: 2869
i: [1,230,0,2,1,1,0,0] t49_bciideal
a: [3,2874,0,2,1,230]
s: 2874
i: [1,1436,0,85,6,1,0,0] t49_bciideal
a: [3,1586,0,4,1,1436]
s: 1586
i: [1,1430,0,95,7,1,0,0] t49_bciideal
a: [3,1604,0,4,1,1430]
s: 1604
i: [1,1426,0,95,7,1,0,0] t49_bciideal
a: [3,1612,0,4,1,1426]
s: 1612
i: [1,237,0,6,2,1,0,0] t49_bciideal
s: 237
i: [1,1438,0,97,7,1,0,0] t49_bciideal
a: [3,1564,0,6,1,1438]
s: 1564
i: [1,1434,0,97,7,1,0,0] t49_bciideal
a: [3,1596,0,6,1,1434]
s: 1596
i: [1,1439,0,2,1,0,0,255] spc1_boole
a: [3,2879,0,2,1,1439]
s: 2879
i: [1,1440,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,2884,0,2,1,1440]
s: 2884
i: [1,1529,0,2,1,0,0,255] existence_l1_bcialg_1
a: [3,2894,0,2,1,1529]
s: 2894
i: [1,1530,0,2,1,0,0,2] existence_l2_struct_0
a: [3,2899,0,2,1,1530]
s: 2899
i: [1,1531,0,2,1,0,0,1] existence_l2_bcialg_1
a: [3,2904,0,2,1,1531]
s: 2904
i: [1,1532,0,2,1,0,0,3] existence_l1_struct_0
a: [3,2909,0,2,1,1532]
s: 2909
i: [1,1535,0,2,1,0,0,2] rc20_struct_0
a: [3,2914,0,2,1,1535]
s: 2914
i: [1,1534,0,2,1,0,0,2] rc20_struct_0
a: [3,2919,0,2,1,1534]
s: 2919
i: [1,1533,0,2,1,0,0,2] rc20_struct_0
a: [3,2924,0,2,1,1533]
s: 2924
i: [1,1543,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2929,0,2,1,1543]
s: 2929
i: [1,1542,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2934,0,2,1,1542]
s: 2934
i: [1,1541,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2939,0,2,1,1541]
s: 2939
i: [1,1540,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2944,0,2,1,1540]
s: 2944
i: [1,1539,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2949,0,2,1,1539]
s: 2949
i: [1,1538,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2954,0,2,1,1538]
s: 2954
i: [1,1537,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2959,0,2,1,1537]
s: 2959
i: [1,1536,0,2,1,0,0,2] rc1_bcialg_3
a: [3,2964,0,2,1,1536]
s: 2964
a: [3,1559,0,10,1,1438]
s: 1559
i: [1,1524,0,3,1,0,0,255] reflexivity_r1_tarski
s: 1524
i: [1,1441,0,3,1,0,0,255] d2_xboole_0
a: [3,2889,0,3,1,1441]
s: 2889
r: 2884
d: [2,2966,0,2,1,2,23,2884,2889]
a: [3,2970,0,2,1,2966]
s: 2970
a: [3,1572,0,15,1,1438]
s: 1572
i: [1,1442,0,4,2,0,0,255] dt_l1_bcialg_1
s: 1442
d: [2,2972,1,2,1,1,40,1442,2894]
a: [3,2976,1,2,1,2972]
s: 2976
i: [1,1447,0,4,2,0,0,2] dt_l2_struct_0
s: 1447
d: [2,2978,1,2,1,1,40,1447,2899]
a: [3,2983,1,2,1,2978]
s: 2983
i: [1,1453,0,4,2,0,0,255] dt_l2_bcialg_1
s: 1453
i: [1,1454,0,4,2,0,0,255] dt_l2_bcialg_1
s: 1454
d: [2,2990,1,2,1,1,40,1453,2844]
a: [3,2996,1,2,1,2990]
s: 2996
i: [1,1523,0,4,1,0,0,1] existence_m1_subset_1
s: 1523
d: [2,3008,1,2,1,1,40,1454,2844]
a: [3,3014,1,2,1,3008]
s: 3014
i: [1,1443,0,5,2,0,0,255] dt_u1_bcialg_1
s: 1443
i: [1,1452,0,5,2,0,0,255] t6_boole
s: 1452
d: [2,3026,2,2,1,1,40,2996,1442]
a: [3,3030,2,2,1,3026]
s: 3030
i: [1,1520,0,5,2,0,0,255] t7_boole
s: 1520
d: [2,2979,1,2,1,1,40,1447,2924]
a: [3,2988,1,2,1,2979]
s: 2988
i: [1,1448,0,6,2,0,0,255] fc10_struct_0
s: 1448
i: [1,1451,0,6,2,0,0,255] rc17_struct_0
s: 1451
d: [2,2991,1,2,1,1,40,1453,2904]
a: [3,3001,1,2,1,2991]
s: 3001
i: [1,1455,0,6,3,0,0,255] cc7_struct_0
s: 1455
d: [2,2992,1,2,1,1,40,1453,2964]
a: [3,3006,1,2,1,2992]
s: 3006
i: [1,1456,0,6,3,0,0,255] cc5_struct_0
s: 1456
i: [1,1458,0,6,3,0,0,3] cc17_struct_0
s: 1458
d: [2,3009,1,2,1,1,40,1454,2904]
a: [3,3019,1,2,1,3009]
s: 3019
i: [1,1459,0,6,3,0,0,3] cc2_struct_0
s: 1459
d: [2,3010,1,2,1,1,40,1454,2964]
a: [3,3024,1,2,1,3010]
s: 3024
i: [1,1460,0,6,3,0,0,255] cc6_struct_0
s: 1460
i: [1,1461,0,6,3,0,0,3] cc1_struct_0
s: 1461
d: [2,3042,1,4,2,1,40,1455,2914]
d: [2,3043,1,2,1,2,21,3042,2988]
a: [3,3047,1,2,1,3043]
s: 3047
i: [1,1464,0,6,3,0,0,255] cc4_struct_0
s: 1464
d: [2,3058,1,4,2,1,40,1456,2909]
a: [3,3081,1,2,1,3058]
s: 3081
i: [1,1516,0,6,2,0,0,255] t1_subset
s: 1516
i: [1,1525,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 1525
d: [2,3036,2,2,1,1,40,3001,1442]
a: [3,3040,2,2,1,3036]
s: 3040
i: [1,1449,0,7,2,0,0,1] dt_k4_struct_0
s: 1449
d: [2,3049,2,2,1,1,40,3006,1442]
a: [3,3053,2,2,1,3049]
s: 3053
i: [1,1450,0,7,2,0,0,255] rc17_struct_0
s: 1450
i: [1,1462,0,7,3,0,0,255] cc8_struct_0
s: 1462
d: [2,3056,2,4,2,1,40,1456,2976]
a: [3,3063,2,2,1,3056]
s: 3063
i: [1,1465,0,7,3,0,0,255] cc11_struct_0
s: 1465
d: [2,3057,2,4,2,1,40,1456,2983]
a: [3,3072,2,2,1,3057]
s: 3072
i: [1,1466,0,7,3,0,0,255] cc11_struct_0
s: 1466
i: [1,1467,0,7,3,0,0,255] cc9_struct_0
s: 1467
d: [2,3111,3,4,2,1,40,3040,1458]
a: [3,3116,3,2,1,3111]
s: 3116
i: [1,1469,0,7,3,0,0,255] fc9_struct_0
s: 1469
i: [1,1472,0,7,3,0,0,255] fc6_struct_0
s: 1472
i: [1,1474,0,7,3,0,0,255] rc22_struct_0
s: 1474
i: [1,1507,0,7,3,0,0,2] fc2_struct_0
s: 1507
i: [1,1509,0,7,3,0,0,2] rc4_struct_0
s: 1509
i: [1,1511,0,7,3,0,0,255] rc21_struct_0
s: 1511
i: [1,1512,0,7,3,0,0,255] rc21_struct_0
s: 1512
i: [1,1513,0,7,3,0,0,255] fc8_struct_0
s: 1513
i: [1,1514,0,7,3,0,0,255] fc7_struct_0
s: 1514
i: [1,1515,0,7,3,0,0,2] fc1_struct_0
s: 1515
i: [1,1518,0,7,2,0,0,255] t3_subset
s: 1518
i: [1,1519,0,7,3,0,0,1] t8_boole
s: 1519
i: [1,1471,0,8,3,0,0,255] rc19_struct_0
s: 1471
i: [1,1517,0,8,3,0,0,255] t2_subset
s: 1517
r: 1586
r: 1604
r: 1612
a: [3,1585,0,4,1,1436]
s: 1585
a: [3,1603,0,4,1,1430]
s: 1603
a: [3,1611,0,4,1,1426]
s: 1611
i: [1,1422,0,86,7,1,0,0] t49_bciideal
a: [3,1620,0,4,1,1422]
s: 1620
i: [1,1406,0,86,7,1,0,0] t49_bciideal
a: [3,1640,0,4,1,1406]
s: 1640
r: 1640
a: [3,1639,0,4,1,1406]
s: 1639
i: [1,1390,0,86,7,1,0,0] t49_bciideal
a: [3,1660,0,4,1,1390]
s: 1660
r: 1660
a: [3,1659,0,4,1,1390]
s: 1659
i: [1,1374,0,86,7,1,0,0] t49_bciideal
a: [3,1680,0,4,1,1374]
s: 1680
r: 1680
a: [3,1679,0,4,1,1374]
s: 1679
i: [1,1358,0,68,6,1,0,0] t49_bciideal
a: [3,1700,0,4,1,1358]
s: 1700
r: 1559
r: 1585
r: 1700
s: 1586
i: [1,1435,0,85,6,1,0,0] t49_bciideal
a: [3,1591,0,4,1,1435]
s: 1591
r: 1591
r: 1639
r: 1659
r: 1679
s: 1640
s: 1660
d: [2,3139,1,7,2,1,40,1517,1586]
a: [3,3153,1,3,1,3139]
s: 3153
d: [2,3201,3,0,0,3,52,2874,3030,3153,1507]
e: 3201
d: [2,3202,2,4,2,1,40,3153,1507]
d: [2,3204,2,2,1,2,21,3202,3030]
d: [2,3205,2,0,0,3,21,3204,2874]
e: 3205
r: 3153
s: 1591
a: [3,3152,1,3,1,3139]
s: 3152
s: 1679
a: [3,1699,0,4,1,1358]
s: 1699
i: [1,1278,0,68,6,1,0,0] t49_bciideal
a: [3,1784,0,4,1,1278]
s: 1784
r: 1784
a: [3,1783,0,4,1,1278]
s: 1783
a: [3,3149,1,4,1,3139]
s: 3149
a: [3,1555,0,6,1,1438]
s: 1555
a: [3,1560,0,10,1,1438]
s: 1560
d: [2,3165,1,7,2,1,40,1640,1517]
a: [3,3169,1,4,1,3165]
s: 3169
a: [3,1568,0,15,1,1438]
s: 1568
d: [2,3171,1,7,2,1,40,1660,1517]
a: [3,3175,1,4,1,3171]
s: 3175
i: [1,1468,0,9,4,0,0,255] cc10_struct_0
s: 1468
i: [1,1470,0,9,3,0,0,255] rc19_struct_0
s: 1470
d: [2,3191,1,7,2,1,40,1591,1517]
a: [3,3195,1,4,1,3191]
s: 3195
i: [1,1522,0,9,3,0,0,255] t5_subset
s: 1522
d: [2,3217,2,4,1,1,40,3149,1525]
a: [3,3223,2,4,1,3217]
s: 3223
i: [1,1473,0,10,3,0,0,255] rc22_struct_0
s: 1473
i: [1,1508,0,10,3,0,0,2] rc4_struct_0
s: 1508
d: [2,3225,2,4,1,1,40,3169,1525]
a: [3,3231,2,4,1,3225]
s: 3231
i: [1,1510,0,10,3,0,0,255] rc21_struct_0
s: 1510
d: [2,3233,2,4,1,1,40,3175,1525]
a: [3,3239,2,4,1,3233]
s: 3239
i: [1,1521,0,10,3,0,0,255] t4_subset
s: 1521
i: [1,1444,0,12,2,0,0,255] dt_u1_bcialg_1
s: 1444
d: [2,3246,2,4,1,1,40,3195,1525]
a: [3,3252,2,4,1,3246]
s: 3252
i: [1,1445,0,14,2,0,0,255] dt_u1_bcialg_1
s: 1445
d: [2,3144,1,6,2,0,40,1517,1523]
s: 3144
i: [1,1505,0,16,7,0,0,1] existence_m2_bcialg_1
s: 1505
i: [1,1501,0,17,8,0,0,1] rc2_bciideal
s: 1501
d: [2,3273,1,16,7,0,40,1505,237]
d: [2,3274,1,14,6,1,21,3273,2844]
d: [2,3275,1,12,5,2,21,3274,2854]
d: [2,3276,1,10,4,3,21,3275,2859]
d: [2,3277,1,8,3,4,21,3276,2864]
d: [2,3278,1,6,2,5,21,3277,2869]
d: [2,3279,1,4,1,6,21,3278,2874]
a: [3,3283,1,4,1,3279]
s: 3283
i: [1,1503,0,17,8,0,0,1] dt_m2_bcialg_1
s: 1503
d: [2,3270,2,6,2,0,40,3144,1525]
s: 3270
i: [1,1506,0,17,7,0,0,255] rc1_bcialg_2
s: 1506
i: [1,1526,0,17,4,0,0,255] dt_k1_bcialg_1
s: 1526
d: [2,3034,1,7,3,0,44,1452,1452]
s: 3034
i: [1,1500,0,18,8,0,0,1] rc2_bciideal
s: 1500
d: [2,3291,1,18,8,0,40,1500,237]
d: [2,3294,1,16,7,1,21,3291,2844]
d: [2,3295,1,14,6,2,21,3294,2849]
d: [2,3296,1,12,5,3,21,3295,2854]
d: [2,3297,1,10,4,4,21,3296,2859]
d: [2,3298,1,8,3,5,21,3297,2864]
d: [2,3299,1,6,2,6,21,3298,2869]
d: [2,3300,1,4,1,7,21,3299,2874]
a: [3,3304,1,4,1,3300]
s: 3304
i: [1,1502,0,18,8,0,0,1] rc2_bciideal
s: 1502
i: [1,1496,0,20,9,0,0,1] t37_bciideal
s: 1496
d: [2,3306,1,20,9,0,40,1496,237]
d: [2,3309,1,18,8,1,21,3306,2844]
d: [2,3310,1,16,7,2,21,3309,2849]
d: [2,3311,1,14,6,3,21,3310,2854]
d: [2,3312,1,12,5,4,21,3311,2859]
d: [2,3313,1,10,4,5,21,3312,2864]
d: [2,3314,1,8,3,6,21,3313,2869]
d: [2,3315,1,6,2,7,21,3314,2874]
a: [3,3319,1,4,1,3315]
s: 3319
i: [1,1504,0,20,8,0,0,1] dt_m2_bcialg_1
s: 1504
d: [2,3255,1,7,2,0,40,1522,1523]
s: 3255
i: [1,1446,0,22,4,0,0,255] d1_bcialg_1
s: 1446
i: [1,1482,0,36,17,0,0,1] t37_bciideal
d: [2,1544,0,22,10,0,18,1482]
s: 1544
d: [2,3328,2,6,2,0,40,3255,3144]
s: 3328
i: [1,1497,0,24,10,0,0,1] t37_bciideal
s: 1497
d: [2,3254,1,8,3,0,40,1522,1518]
s: 3254
d: [2,3335,0,20,9,0,21,1497,1496]
s: 3335
i: [1,1527,0,29,6,0,0,255] dt_k5_binop_1
s: 1527
a: [3,3323,1,2,1,3315]
s: 3323
i: [1,1528,0,32,6,0,0,255] redefinition_k5_binop_1
s: 1528
d: [2,3266,1,8,2,0,40,1521,1523]
s: 3266
i: [1,1477,0,50,19,0,0,2] t7_bcialg_3
d: [2,1550,0,36,12,0,18,1477]
s: 1550
i: [1,1475,0,53,18,0,0,2] t10_bcialg_3
d: [2,1551,0,39,11,0,18,1475]
s: 1551
r: 1555
r: 1560
a: [3,1556,0,6,1,1438]
s: 1556
s: 1559
d: [2,3409,1,46,11,1,44,1551,1559]
d: [2,3427,1,42,10,2,21,3409,1586]
d: [2,3428,1,40,9,3,21,3427,2844]
d: [2,3429,1,38,8,4,21,3428,3323]
d: [2,3430,1,36,7,5,21,3429,2849]
d: [2,3431,1,34,6,6,21,3430,2854]
d: [2,3432,1,32,5,7,21,3431,2859]
d: [2,3433,1,30,4,8,21,3432,2864]
d: [2,3434,1,28,3,9,21,3433,2869]
d: [2,3435,1,26,2,10,21,3434,2874]
a: [3,3439,1,7,1,3435]
s: 3439
d: [2,3471,2,0,0,4,52,2996,1586,1591,3439,1526]
e: 3471
d: [2,3472,2,10,3,1,40,3439,1526]
d: [2,3474,2,6,2,2,21,3472,1591]
d: [2,3475,2,2,1,3,21,3474,1586]
d: [2,3476,2,0,0,4,21,3475,2996]
e: 3476
r: 3439
i: [1,1478,0,54,19,0,0,1] t9_bcialg_3
d: [2,1548,0,40,12,0,18,1478]
s: 1548
i: [1,1479,0,54,19,0,0,1] t9_bcialg_3
d: [2,1549,0,40,12,0,18,1479]
s: 1549
d: [2,3497,2,0,0,12,52,2874,2869,2864,2859,2854,2849,3323,2844,1591,1586,1556,1559,1549]
e: 3497
d: [2,3501,1,35,12,1,44,1549,1559]
d: [2,3504,1,30,11,1,19,3501]
d: [2,3516,1,24,10,2,21,3504,1556]
d: [2,3517,1,20,9,3,21,3516,1586]
d: [2,3518,1,16,8,4,21,3517,1591]
d: [2,3519,1,14,7,5,21,3518,2844]
d: [2,3520,1,12,6,6,21,3519,3323]
d: [2,3521,1,10,5,7,21,3520,2849]
d: [2,3522,1,8,4,8,21,3521,2854]
d: [2,3523,1,6,3,9,21,3522,2859]
d: [2,3524,1,4,2,10,21,3523,2864]
d: [2,3525,1,2,1,11,21,3524,2869]
d: [2,3526,1,0,0,12,21,3525,2874]
e: 3526
r: 1556
r: 1559
d: [2,3410,1,43,12,1,44,1550,1559]
d: [2,3445,1,39,11,2,21,3410,1586]
d: [2,3446,1,37,10,3,21,3445,2844]
d: [2,3447,1,35,9,4,21,3446,3323]
d: [2,3448,1,33,8,5,21,3447,2849]
d: [2,3449,1,31,7,6,21,3448,2854]
d: [2,3450,1,29,6,7,21,3449,2859]
d: [2,3451,1,27,5,8,21,3450,2864]
d: [2,3452,1,25,4,9,21,3451,2869]
d: [2,3453,1,23,3,10,21,3452,2874]
a: [3,3461,1,4,1,3453]
s: 3461
s: 1555
r: 1555
r: 1699
r: 1783
r: 1572
s: 1700
s: 1784
s: 1556
d: [2,3183,1,7,2,1,40,1700,1517]
a: [3,3187,1,4,1,3183]
s: 3187
s: 1560
d: [2,3210,1,7,2,1,40,1784,1517]
d: [2,3211,1,4,1,2,21,3210,3152]
a: [3,3215,1,4,1,3211]
s: 3215
a: [3,1571,0,15,1,1438]
s: 1571
a: [3,1576,0,18,1,1438]
s: 1576
d: [2,3598,2,4,1,1,40,3187,1525]
a: [3,3604,2,4,1,3598]
s: 3604
i: [1,1476,0,56,18,0,0,2] t10_bcialg_3
d: [2,1552,0,42,11,0,18,1476]
s: 1552
d: [2,3690,2,0,0,11,52,2874,2869,2864,2859,2854,2849,3323,2844,1784,1700,1576,1552]
e: 3690
d: [2,3695,1,33,11,1,44,1576,1552]
d: [2,3709,1,24,10,1,19,3695]
d: [2,3723,1,20,9,2,21,3709,1700]
d: [2,3724,1,16,8,3,21,3723,1784]
d: [2,3725,1,14,7,4,21,3724,2844]
d: [2,3726,1,12,6,5,21,3725,3323]
d: [2,3727,1,10,5,6,21,3726,2849]
d: [2,3728,1,8,4,7,21,3727,2854]
d: [2,3729,1,6,3,8,21,3728,2859]
d: [2,3730,1,4,2,9,21,3729,2864]
d: [2,3731,1,2,1,10,21,3730,2869]
d: [2,3732,1,0,0,11,21,3731,2874]
e: 3732
r: 1556
r: 1560
r: 1576
r: 1679
s: 1680
s: 1559
s: 1555
d: [2,3177,1,7,2,1,40,1680,1517]
a: [3,3181,1,4,1,3177]
s: 3181
i: [1,1480,0,61,20,0,0,1] t8_bcialg_3
d: [2,1547,0,47,13,0,18,1480]
s: 1547
d: [2,3883,2,0,0,13,52,2874,2869,2864,2859,2854,2849,3323,2844,1680,1660,1640,1620,1568,1547]
e: 3883
d: [2,3899,1,41,13,1,44,1568,1547]
d: [2,3913,1,32,12,1,19,3899]
d: [2,3978,1,28,11,2,21,3913,1620]
d: [2,3979,1,24,10,3,21,3978,1640]
d: [2,3980,1,20,9,4,21,3979,1660]
d: [2,3981,1,16,8,5,21,3980,1680]
d: [2,3982,1,14,7,6,21,3981,2844]
d: [2,3983,1,12,6,7,21,3982,3323]
d: [2,3984,1,10,5,8,21,3983,2849]
d: [2,3985,1,8,4,9,21,3984,2854]
d: [2,3986,1,6,3,10,21,3985,2859]
d: [2,3987,1,4,2,11,21,3986,2864]
d: [2,3988,1,2,1,12,21,3987,2869]
d: [2,3989,1,0,0,13,21,3988,2874]
e: 3989
r: 1555
r: 1559
r: 1568
r: 1784
r: 3215
s: 1783
i: [1,1198,0,74,7,1,0,0] t49_bciideal
a: [3,1868,0,4,1,1198]
s: 1868
i: [1,958,0,74,7,1,0,0] t49_bciideal
a: [3,2112,0,4,1,958]
s: 2112
i: [1,718,0,74,7,1,0,0] t49_bciideal
a: [3,2356,0,4,1,718]
s: 2356
i: [1,478,0,74,7,1,0,0] t49_bciideal
a: [3,2600,0,4,1,478]
s: 2600
a: [3,3214,1,4,1,3211]
s: 3214
s: 1556
s: 1560
d: [2,4067,1,7,2,1,40,2112,1517]
d: [2,4068,1,4,1,2,21,4067,3152]
a: [3,4072,1,4,1,4068]
s: 4072
a: [3,1580,0,27,1,1438]
s: 1580
d: [2,4096,1,53,13,1,44,1580,1547]
d: [2,4097,1,32,12,1,19,4096]
d: [2,4098,1,28,11,2,21,4097,1868]
d: [2,4099,1,24,10,3,21,4098,2112]
d: [2,4100,1,20,9,4,21,4099,2600]
d: [2,4101,1,16,8,5,21,4100,2356]
d: [2,4102,1,14,7,6,21,4101,2844]
d: [2,4103,1,12,6,7,21,4102,3323]
d: [2,4104,1,10,5,8,21,4103,2849]
d: [2,4105,1,8,4,9,21,4104,2854]
d: [2,4106,1,6,3,10,21,4105,2859]
d: [2,4107,1,4,2,11,21,4106,2864]
d: [2,4108,1,2,1,12,21,4107,2869]
d: [2,4109,1,0,0,13,21,4108,2874]
e: 4109
r: 1556
r: 1560
r: 1591
r: 1603
r: 1611
r: 1660
r: 1700
r: 1783
r: 2600
r: 3214
i: [1,1437,0,97,7,1,0,0] t49_bciideal
i: [1,1433,0,97,7,1,0,0] t49_bciideal
i: [1,1432,0,85,6,1,0,0] t49_bciideal
i: [1,1431,0,85,6,1,0,0] t49_bciideal
i: [1,1429,0,95,7,1,0,0] t49_bciideal
i: [1,1428,0,83,6,1,0,0] t49_bciideal
i: [1,1427,0,83,6,1,0,0] t49_bciideal
i: [1,1425,0,95,7,1,0,0] t49_bciideal
i: [1,1424,0,83,6,1,0,0] t49_bciideal
i: [1,1423,0,83,6,1,0,0] t49_bciideal
i: [1,1421,0,86,7,1,0,0] t49_bciideal
i: [1,1420,0,74,6,1,0,0] t49_bciideal
i: [1,1419,0,74,6,1,0,0] t49_bciideal
i: [1,1418,0,86,7,1,0,0] t49_bciideal
i: [1,1417,0,86,7,1,0,0] t49_bciideal
i: [1,1416,0,74,6,1,0,0] t49_bciideal
i: [1,1415,0,74,6,1,0,0] t49_bciideal
i: [1,1414,0,84,7,1,0,0] t49_bciideal
i: [1,1413,0,84,7,1,0,0] t49_bciideal
i: [1,1412,0,72,6,1,0,0] t49_bciideal
i: [1,1411,0,72,6,1,0,0] t49_bciideal
i: [1,1410,0,84,7,1,0,0] t49_bciideal
i: [1,1409,0,84,7,1,0,0] t49_bciideal
i: [1,1408,0,72,6,1,0,0] t49_bciideal
i: [1,1407,0,72,6,1,0,0] t49_bciideal
i: [1,1405,0,86,7,1,0,0] t49_bciideal
i: [1,1404,0,74,6,1,0,0] t49_bciideal
i: [1,1403,0,74,6,1,0,0] t49_bciideal
i: [1,1402,0,86,7,1,0,0] t49_bciideal
i: [1,1401,0,86,7,1,0,0] t49_bciideal
i: [1,1400,0,74,6,1,0,0] t49_bciideal
i: [1,1399,0,74,6,1,0,0] t49_bciideal
i: [1,1398,0,84,7,1,0,0] t49_bciideal
i: [1,1397,0,84,7,1,0,0] t49_bciideal
i: [1,1396,0,72,6,1,0,0] t49_bciideal
i: [1,1395,0,72,6,1,0,0] t49_bciideal
i: [1,1394,0,84,7,1,0,0] t49_bciideal
i: [1,1393,0,84,7,1,0,0] t49_bciideal
i: [1,1392,0,72,6,1,0,0] t49_bciideal
i: [1,1391,0,72,6,1,0,0] t49_bciideal
i: [1,1389,0,86,7,1,0,0] t49_bciideal
i: [1,1388,0,74,6,1,0,0] t49_bciideal
i: [1,1387,0,74,6,1,0,0] t49_bciideal
i: [1,1386,0,86,7,1,0,0] t49_bciideal
i: [1,1385,0,86,7,1,0,0] t49_bciideal
i: [1,1384,0,74,6,1,0,0] t49_bciideal
i: [1,1383,0,74,6,1,0,0] t49_bciideal
i: [1,1382,0,84,7,1,0,0] t49_bciideal
i: [1,1381,0,84,7,1,0,0] t49_bciideal
i: [1,1380,0,72,6,1,0,0] t49_bciideal
i: [1,1379,0,72,6,1,0,0] t49_bciideal
i: [1,1378,0,84,7,1,0,0] t49_bciideal
i: [1,1377,0,84,7,1,0,0] t49_bciideal
i: [1,1376,0,72,6,1,0,0] t49_bciideal
i: [1,1375,0,72,6,1,0,0] t49_bciideal
i: [1,1373,0,86,7,1,0,0] t49_bciideal
i: [1,1372,0,74,6,1,0,0] t49_bciideal
i: [1,1371,0,74,6,1,0,0] t49_bciideal
i: [1,1370,0,86,7,1,0,0] t49_bciideal
i: [1,1369,0,86,7,1,0,0] t49_bciideal
i: [1,1368,0,74,6,1,0,0] t49_bciideal
i: [1,1367,0,74,6,1,0,0] t49_bciideal
i: [1,1366,0,84,7,1,0,0] t49_bciideal
i: [1,1365,0,84,7,1,0,0] t49_bciideal
i: [1,1364,0,72,6,1,0,0] t49_bciideal
i: [1,1363,0,72,6,1,0,0] t49_bciideal
i: [1,1362,0,84,7,1,0,0] t49_bciideal
i: [1,1361,0,84,7,1,0,0] t49_bciideal
i: [1,1360,0,72,6,1,0,0] t49_bciideal
i: [1,1359,0,72,6,1,0,0] t49_bciideal
i: [1,1357,0,68,6,1,0,0] t49_bciideal
i: [1,1356,0,56,5,1,0,0] t49_bciideal
i: [1,1355,0,56,5,1,0,0] t49_bciideal
i: [1,1354,0,68,6,1,0,0] t49_bciideal
i: [1,1353,0,68,6,1,0,0] t49_bciideal
i: [1,1352,0,56,5,1,0,0] t49_bciideal
i: [1,1351,0,56,5,1,0,0] t49_bciideal
i: [1,1350,0,66,6,1,0,0] t49_bciideal
i: [1,1349,0,66,6,1,0,0] t49_bciideal
i: [1,1348,0,54,5,1,0,0] t49_bciideal
i: [1,1347,0,54,5,1,0,0] t49_bciideal
i: [1,1346,0,66,6,1,0,0] t49_bciideal
i: [1,1345,0,66,6,1,0,0] t49_bciideal
i: [1,1344,0,54,5,1,0,0] t49_bciideal
i: [1,1343,0,54,5,1,0,0] t49_bciideal
i: [1,1342,0,57,6,1,0,0] t49_bciideal
i: [1,1341,0,57,6,1,0,0] t49_bciideal
i: [1,1340,0,45,5,1,0,0] t49_bciideal
i: [1,1339,0,45,5,1,0,0] t49_bciideal
i: [1,1338,0,57,6,1,0,0] t49_bciideal
i: [1,1337,0,57,6,1,0,0] t49_bciideal
i: [1,1336,0,45,5,1,0,0] t49_bciideal
i: [1,1335,0,45,5,1,0,0] t49_bciideal
i: [1,1334,0,55,6,1,0,0] t49_bciideal
i: [1,1333,0,55,6,1,0,0] t49_bciideal
i: [1,1332,0,43,5,1,0,0] t49_bciideal
i: [1,1331,0,43,5,1,0,0] t49_bciideal
i: [1,1330,0,55,6,1,0,0] t49_bciideal
i: [1,1329,0,55,6,1,0,0] t49_bciideal
i: [1,1328,0,43,5,1,0,0] t49_bciideal
i: [1,1327,0,43,5,1,0,0] t49_bciideal
i: [1,1326,0,57,6,1,0,0] t49_bciideal
i: [1,1325,0,57,6,1,0,0] t49_bciideal
i: [1,1324,0,45,5,1,0,0] t49_bciideal
i: [1,1323,0,45,5,1,0,0] t49_bciideal
i: [1,1322,0,57,6,1,0,0] t49_bciideal
i: [1,1321,0,57,6,1,0,0] t49_bciideal
i: [1,1320,0,45,5,1,0,0] t49_bciideal
i: [1,1319,0,45,5,1,0,0] t49_bciideal
i: [1,1318,0,55,6,1,0,0] t49_bciideal
i: [1,1317,0,55,6,1,0,0] t49_bciideal
i: [1,1316,0,43,5,1,0,0] t49_bciideal
i: [1,1315,0,43,5,1,0,0] t49_bciideal
i: [1,1314,0,55,6,1,0,0] t49_bciideal
i: [1,1313,0,55,6,1,0,0] t49_bciideal
i: [1,1312,0,43,5,1,0,0] t49_bciideal
i: [1,1311,0,43,5,1,0,0] t49_bciideal
i: [1,1310,0,57,6,1,0,0] t49_bciideal
i: [1,1309,0,57,6,1,0,0] t49_bciideal
i: [1,1308,0,45,5,1,0,0] t49_bciideal
i: [1,1307,0,45,5,1,0,0] t49_bciideal
i: [1,1306,0,57,6,1,0,0] t49_bciideal
i: [1,1305,0,57,6,1,0,0] t49_bciideal
i: [1,1304,0,45,5,1,0,0] t49_bciideal
i: [1,1303,0,45,5,1,0,0] t49_bciideal
i: [1,1302,0,55,6,1,0,0] t49_bciideal
i: [1,1301,0,55,6,1,0,0] t49_bciideal
i: [1,1300,0,43,5,1,0,0] t49_bciideal
i: [1,1299,0,43,5,1,0,0] t49_bciideal
i: [1,1298,0,55,6,1,0,0] t49_bciideal
i: [1,1297,0,55,6,1,0,0] t49_bciideal
i: [1,1296,0,43,5,1,0,0] t49_bciideal
i: [1,1295,0,43,5,1,0,0] t49_bciideal
i: [1,1294,0,57,6,1,0,0] t49_bciideal
i: [1,1293,0,57,6,1,0,0] t49_bciideal
i: [1,1292,0,45,5,1,0,0] t49_bciideal
i: [1,1291,0,45,5,1,0,0] t49_bciideal
i: [1,1290,0,57,6,1,0,0] t49_bciideal
i: [1,1289,0,57,6,1,0,0] t49_bciideal
i: [1,1288,0,45,5,1,0,0] t49_bciideal
i: [1,1287,0,45,5,1,0,0] t49_bciideal
i: [1,1286,0,55,6,1,0,0] t49_bciideal
i: [1,1285,0,55,6,1,0,0] t49_bciideal
i: [1,1284,0,43,5,1,0,0] t49_bciideal
i: [1,1283,0,43,5,1,0,0] t49_bciideal
i: [1,1282,0,55,6,1,0,0] t49_bciideal
i: [1,1281,0,55,6,1,0,0] t49_bciideal
i: [1,1280,0,43,5,1,0,0] t49_bciideal
i: [1,1279,0,43,5,1,0,0] t49_bciideal
i: [1,1277,0,68,6,1,0,0] t49_bciideal
i: [1,1276,0,56,5,1,0,0] t49_bciideal
i: [1,1275,0,56,5,1,0,0] t49_bciideal
i: [1,1274,0,68,6,1,0,0] t49_bciideal
i: [1,1273,0,68,6,1,0,0] t49_bciideal
i: [1,1272,0,56,5,1,0,0] t49_bciideal
i: [1,1271,0,56,5,1,0,0] t49_bciideal
i: [1,1270,0,66,6,1,0,0] t49_bciideal
i: [1,1269,0,66,6,1,0,0] t49_bciideal
i: [1,1268,0,54,5,1,0,0] t49_bciideal
i: [1,1267,0,54,5,1,0,0] t49_bciideal
i: [1,1266,0,66,6,1,0,0] t49_bciideal
i: [1,1265,0,66,6,1,0,0] t49_bciideal
i: [1,1264,0,54,5,1,0,0] t49_bciideal
i: [1,1263,0,54,5,1,0,0] t49_bciideal
i: [1,1262,0,57,6,1,0,0] t49_bciideal
i: [1,1261,0,57,6,1,0,0] t49_bciideal
i: [1,1260,0,45,5,1,0,0] t49_bciideal
i: [1,1259,0,45,5,1,0,0] t49_bciideal
i: [1,1258,0,57,6,1,0,0] t49_bciideal
i: [1,1257,0,57,6,1,0,0] t49_bciideal
i: [1,1256,0,45,5,1,0,0] t49_bciideal
i: [1,1255,0,45,5,1,0,0] t49_bciideal
i: [1,1254,0,55,6,1,0,0] t49_bciideal
i: [1,1253,0,55,6,1,0,0] t49_bciideal
i: [1,1252,0,43,5,1,0,0] t49_bciideal
i: [1,1251,0,43,5,1,0,0] t49_bciideal
i: [1,1250,0,55,6,1,0,0] t49_bciideal
i: [1,1249,0,55,6,1,0,0] t49_bciideal
i: [1,1248,0,43,5,1,0,0] t49_bciideal
i: [1,1247,0,43,5,1,0,0] t49_bciideal
i: [1,1246,0,57,6,1,0,0] t49_bciideal
i: [1,1245,0,57,6,1,0,0] t49_bciideal
i: [1,1244,0,45,5,1,0,0] t49_bciideal
i: [1,1243,0,45,5,1,0,0] t49_bciideal
i: [1,1242,0,57,6,1,0,0] t49_bciideal
i: [1,1241,0,57,6,1,0,0] t49_bciideal
i: [1,1240,0,45,5,1,0,0] t49_bciideal
i: [1,1239,0,45,5,1,0,0] t49_bciideal
i: [1,1238,0,55,6,1,0,0] t49_bciideal
i: [1,1237,0,55,6,1,0,0] t49_bciideal
i: [1,1236,0,43,5,1,0,0] t49_bciideal
i: [1,1235,0,43,5,1,0,0] t49_bciideal
i: [1,1234,0,55,6,1,0,0] t49_bciideal
i: [1,1233,0,55,6,1,0,0] t49_bciideal
i: [1,1232,0,43,5,1,0,0] t49_bciideal
i: [1,1231,0,43,5,1,0,0] t49_bciideal
i: [1,1230,0,57,6,1,0,0] t49_bciideal
i: [1,1229,0,57,6,1,0,0] t49_bciideal
i: [1,1228,0,45,5,1,0,0] t49_bciideal
i: [1,1227,0,45,5,1,0,0] t49_bciideal
i: [1,1226,0,57,6,1,0,0] t49_bciideal
i: [1,1225,0,57,6,1,0,0] t49_bciideal
i: [1,1224,0,45,5,1,0,0] t49_bciideal
i: [1,1223,0,45,5,1,0,0] t49_bciideal
i: [1,1222,0,55,6,1,0,0] t49_bciideal
i: [1,1221,0,55,6,1,0,0] t49_bciideal
i: [1,1220,0,43,5,1,0,0] t49_bciideal
i: [1,1219,0,43,5,1,0,0] t49_bciideal
i: [1,1218,0,55,6,1,0,0] t49_bciideal
i: [1,1217,0,55,6,1,0,0] t49_bciideal
i: [1,1216,0,43,5,1,0,0] t49_bciideal
i: [1,1215,0,43,5,1,0,0] t49_bciideal
i: [1,1214,0,57,6,1,0,0] t49_bciideal
i: [1,1213,0,57,6,1,0,0] t49_bciideal
i: [1,1212,0,45,5,1,0,0] t49_bciideal
i: [1,1211,0,45,5,1,0,0] t49_bciideal
i: [1,1210,0,57,6,1,0,0] t49_bciideal
i: [1,1209,0,57,6,1,0,0] t49_bciideal
i: [1,1208,0,45,5,1,0,0] t49_bciideal
i: [1,1207,0,45,5,1,0,0] t49_bciideal
i: [1,1206,0,55,6,1,0,0] t49_bciideal
i: [1,1205,0,55,6,1,0,0] t49_bciideal
i: [1,1204,0,43,5,1,0,0] t49_bciideal
i: [1,1203,0,43,5,1,0,0] t49_bciideal
i: [1,1202,0,55,6,1,0,0] t49_bciideal
i: [1,1201,0,55,6,1,0,0] t49_bciideal
i: [1,1200,0,43,5,1,0,0] t49_bciideal
i: [1,1199,0,43,5,1,0,0] t49_bciideal
i: [1,1197,0,74,7,1,0,0] t49_bciideal
i: [1,1196,0,62,6,1,0,0] t49_bciideal
i: [1,1195,0,62,6,1,0,0] t49_bciideal
i: [1,1194,0,74,7,1,0,0] t49_bciideal
i: [1,1193,0,74,7,1,0,0] t49_bciideal
i: [1,1192,0,62,6,1,0,0] t49_bciideal
i: [1,1191,0,62,6,1,0,0] t49_bciideal
i: [1,1190,0,72,7,1,0,0] t49_bciideal
i: [1,1189,0,72,7,1,0,0] t49_bciideal
i: [1,1188,0,60,6,1,0,0] t49_bciideal
i: [1,1187,0,60,6,1,0,0] t49_bciideal
i: [1,1186,0,72,7,1,0,0] t49_bciideal
i: [1,1185,0,72,7,1,0,0] t49_bciideal
i: [1,1184,0,60,6,1,0,0] t49_bciideal
i: [1,1183,0,60,6,1,0,0] t49_bciideal
i: [1,1182,0,63,7,1,0,0] t49_bciideal
i: [1,1181,0,63,7,1,0,0] t49_bciideal
i: [1,1180,0,51,6,1,0,0] t49_bciideal
i: [1,1179,0,51,6,1,0,0] t49_bciideal
i: [1,1178,0,63,7,1,0,0] t49_bciideal
i: [1,1177,0,63,7,1,0,0] t49_bciideal
i: [1,1176,0,51,6,1,0,0] t49_bciideal
i: [1,1175,0,51,6,1,0,0] t49_bciideal
i: [1,1174,0,61,7,1,0,0] t49_bciideal
i: [1,1173,0,61,7,1,0,0] t49_bciideal
i: [1,1172,0,49,6,1,0,0] t49_bciideal
i: [1,1171,0,49,6,1,0,0] t49_bciideal
i: [1,1170,0,61,7,1,0,0] t49_bciideal
i: [1,1169,0,61,7,1,0,0] t49_bciideal
i: [1,1168,0,49,6,1,0,0] t49_bciideal
i: [1,1167,0,49,6,1,0,0] t49_bciideal
i: [1,1166,0,63,7,1,0,0] t49_bciideal
i: [1,1165,0,63,7,1,0,0] t49_bciideal
i: [1,1164,0,51,6,1,0,0] t49_bciideal
i: [1,1163,0,51,6,1,0,0] t49_bciideal
i: [1,1162,0,63,7,1,0,0] t49_bciideal
i: [1,1161,0,63,7,1,0,0] t49_bciideal
i: [1,1160,0,51,6,1,0,0] t49_bciideal
i: [1,1159,0,51,6,1,0,0] t49_bciideal
i: [1,1158,0,61,7,1,0,0] t49_bciideal
i: [1,1157,0,61,7,1,0,0] t49_bciideal
i: [1,1156,0,49,6,1,0,0] t49_bciideal
i: [1,1155,0,49,6,1,0,0] t49_bciideal
i: [1,1154,0,61,7,1,0,0] t49_bciideal
i: [1,1153,0,61,7,1,0,0] t49_bciideal
i: [1,1152,0,49,6,1,0,0] t49_bciideal
i: [1,1151,0,49,6,1,0,0] t49_bciideal
i: [1,1150,0,63,7,1,0,0] t49_bciideal
i: [1,1149,0,63,7,1,0,0] t49_bciideal
i: [1,1148,0,51,6,1,0,0] t49_bciideal
i: [1,1147,0,51,6,1,0,0] t49_bciideal
i: [1,1146,0,63,7,1,0,0] t49_bciideal
i: [1,1145,0,63,7,1,0,0] t49_bciideal
i: [1,1144,0,51,6,1,0,0] t49_bciideal
i: [1,1143,0,51,6,1,0,0] t49_bciideal
i: [1,1142,0,61,7,1,0,0] t49_bciideal
i: [1,1141,0,61,7,1,0,0] t49_bciideal
i: [1,1140,0,49,6,1,0,0] t49_bciideal
i: [1,1139,0,49,6,1,0,0] t49_bciideal
i: [1,1138,0,61,7,1,0,0] t49_bciideal
i: [1,1137,0,61,7,1,0,0] t49_bciideal
i: [1,1136,0,49,6,1,0,0] t49_bciideal
i: [1,1135,0,49,6,1,0,0] t49_bciideal
i: [1,1134,0,63,7,1,0,0] t49_bciideal
i: [1,1133,0,63,7,1,0,0] t49_bciideal
i: [1,1132,0,51,6,1,0,0] t49_bciideal
i: [1,1131,0,51,6,1,0,0] t49_bciideal
i: [1,1130,0,63,7,1,0,0] t49_bciideal
i: [1,1129,0,63,7,1,0,0] t49_bciideal
i: [1,1128,0,51,6,1,0,0] t49_bciideal
i: [1,1127,0,51,6,1,0,0] t49_bciideal
i: [1,1126,0,61,7,1,0,0] t49_bciideal
i: [1,1125,0,61,7,1,0,0] t49_bciideal
i: [1,1124,0,49,6,1,0,0] t49_bciideal
i: [1,1123,0,49,6,1,0,0] t49_bciideal
i: [1,1122,0,61,7,1,0,0] t49_bciideal
i: [1,1121,0,61,7,1,0,0] t49_bciideal
i: [1,1120,0,49,6,1,0,0] t49_bciideal
i: [1,1119,0,49,6,1,0,0] t49_bciideal
i: [1,1118,0,45,6,1,0,0] t49_bciideal
i: [1,1117,0,45,6,1,0,0] t49_bciideal
i: [1,1116,0,33,5,1,0,0] t49_bciideal
i: [1,1115,0,33,5,1,0,0] t49_bciideal
i: [1,1114,0,45,6,1,0,0] t49_bciideal
i: [1,1113,0,45,6,1,0,0] t49_bciideal
i: [1,1112,0,33,5,1,0,0] t49_bciideal
i: [1,1111,0,33,5,1,0,0] t49_bciideal
i: [1,1110,0,43,6,1,0,0] t49_bciideal
i: [1,1109,0,43,6,1,0,0] t49_bciideal
i: [1,1108,0,31,5,1,0,0] t49_bciideal
i: [1,1107,0,31,5,1,0,0] t49_bciideal
i: [1,1106,0,43,6,1,0,0] t49_bciideal
i: [1,1105,0,43,6,1,0,0] t49_bciideal
i: [1,1104,0,31,5,1,0,0] t49_bciideal
i: [1,1103,0,31,5,1,0,0] t49_bciideal
i: [1,1102,0,34,6,1,0,0] t49_bciideal
i: [1,1101,0,34,6,1,0,0] t49_bciideal
i: [1,1100,0,22,5,1,0,0] t49_bciideal
i: [1,1099,0,22,5,1,0,0] t49_bciideal
i: [1,1098,0,34,6,1,0,0] t49_bciideal
i: [1,1097,0,34,6,1,0,0] t49_bciideal
i: [1,1096,0,22,5,1,0,0] t49_bciideal
i: [1,1095,0,22,5,1,0,0] t49_bciideal
i: [1,1094,0,32,6,1,0,0] t49_bciideal
i: [1,1093,0,32,6,1,0,0] t49_bciideal
i: [1,1092,0,20,5,1,0,0] t49_bciideal
i: [1,1091,0,20,5,1,0,0] t49_bciideal
i: [1,1090,0,32,6,1,0,0] t49_bciideal
i: [1,1089,0,32,6,1,0,0] t49_bciideal
i: [1,1088,0,20,5,1,0,0] t49_bciideal
i: [1,1087,0,20,5,1,0,0] t49_bciideal
i: [1,1086,0,34,6,1,0,0] t49_bciideal
i: [1,1085,0,34,6,1,0,0] t49_bciideal
i: [1,1084,0,22,5,1,0,0] t49_bciideal
i: [1,1083,0,22,5,1,0,0] t49_bciideal
i: [1,1082,0,34,6,1,0,0] t49_bciideal
i: [1,1081,0,34,6,1,0,0] t49_bciideal
i: [1,1080,0,22,5,1,0,0] t49_bciideal
i: [1,1079,0,22,5,1,0,0] t49_bciideal
i: [1,1078,0,32,6,1,0,0] t49_bciideal
i: [1,1077,0,32,6,1,0,0] t49_bciideal
i: [1,1076,0,20,5,1,0,0] t49_bciideal
i: [1,1075,0,20,5,1,0,0] t49_bciideal
i: [1,1074,0,32,6,1,0,0] t49_bciideal
i: [1,1073,0,32,6,1,0,0] t49_bciideal
i: [1,1072,0,20,5,1,0,0] t49_bciideal
i: [1,1071,0,20,5,1,0,0] t49_bciideal
i: [1,1070,0,34,6,1,0,0] t49_bciideal
i: [1,1069,0,34,6,1,0,0] t49_bciideal
i: [1,1068,0,22,5,1,0,0] t49_bciideal
i: [1,1067,0,22,5,1,0,0] t49_bciideal
i: [1,1066,0,34,6,1,0,0] t49_bciideal
i: [1,1065,0,34,6,1,0,0] t49_bciideal
i: [1,1064,0,22,5,1,0,0] t49_bciideal
i: [1,1063,0,22,5,1,0,0] t49_bciideal
i: [1,1062,0,32,6,1,0,0] t49_bciideal
i: [1,1061,0,32,6,1,0,0] t49_bciideal
i: [1,1060,0,20,5,1,0,0] t49_bciideal
i: [1,1059,0,20,5,1,0,0] t49_bciideal
i: [1,1058,0,32,6,1,0,0] t49_bciideal
i: [1,1057,0,32,6,1,0,0] t49_bciideal
i: [1,1056,0,20,5,1,0,0] t49_bciideal
i: [1,1055,0,20,5,1,0,0] t49_bciideal
i: [1,1054,0,34,6,1,0,0] t49_bciideal
i: [1,1053,0,34,6,1,0,0] t49_bciideal
i: [1,1052,0,22,5,1,0,0] t49_bciideal
i: [1,1051,0,22,5,1,0,0] t49_bciideal
i: [1,1050,0,34,6,1,0,0] t49_bciideal
i: [1,1049,0,34,6,1,0,0] t49_bciideal
i: [1,1048,0,22,5,1,0,0] t49_bciideal
i: [1,1047,0,22,5,1,0,0] t49_bciideal
i: [1,1046,0,32,6,1,0,0] t49_bciideal
i: [1,1045,0,32,6,1,0,0] t49_bciideal
i: [1,1044,0,20,5,1,0,0] t49_bciideal
i: [1,1043,0,20,5,1,0,0] t49_bciideal
i: [1,1042,0,32,6,1,0,0] t49_bciideal
i: [1,1041,0,32,6,1,0,0] t49_bciideal
i: [1,1040,0,20,5,1,0,0] t49_bciideal
i: [1,1039,0,20,5,1,0,0] t49_bciideal
i: [1,1038,0,45,6,1,0,0] t49_bciideal
i: [1,1037,0,45,6,1,0,0] t49_bciideal
i: [1,1036,0,33,5,1,0,0] t49_bciideal
i: [1,1035,0,33,5,1,0,0] t49_bciideal
i: [1,1034,0,45,6,1,0,0] t49_bciideal
i: [1,1033,0,45,6,1,0,0] t49_bciideal
i: [1,1032,0,33,5,1,0,0] t49_bciideal
i: [1,1031,0,33,5,1,0,0] t49_bciideal
i: [1,1030,0,43,6,1,0,0] t49_bciideal
i: [1,1029,0,43,6,1,0,0] t49_bciideal
i: [1,1028,0,31,5,1,0,0] t49_bciideal
i: [1,1027,0,31,5,1,0,0] t49_bciideal
i: [1,1026,0,43,6,1,0,0] t49_bciideal
i: [1,1025,0,43,6,1,0,0] t49_bciideal
i: [1,1024,0,31,5,1,0,0] t49_bciideal
i: [1,1023,0,31,5,1,0,0] t49_bciideal
i: [1,1022,0,34,6,1,0,0] t49_bciideal
i: [1,1021,0,34,6,1,0,0] t49_bciideal
i: [1,1020,0,22,5,1,0,0] t49_bciideal
i: [1,1019,0,22,5,1,0,0] t49_bciideal
i: [1,1018,0,34,6,1,0,0] t49_bciideal
i: [1,1017,0,34,6,1,0,0] t49_bciideal
i: [1,1016,0,22,5,1,0,0] t49_bciideal
i: [1,1015,0,22,5,1,0,0] t49_bciideal
i: [1,1014,0,32,6,1,0,0] t49_bciideal
i: [1,1013,0,32,6,1,0,0] t49_bciideal
i: [1,1012,0,20,5,1,0,0] t49_bciideal
i: [1,1011,0,20,5,1,0,0] t49_bciideal
i: [1,1010,0,32,6,1,0,0] t49_bciideal
i: [1,1009,0,32,6,1,0,0] t49_bciideal
i: [1,1008,0,20,5,1,0,0] t49_bciideal
i: [1,1007,0,20,5,1,0,0] t49_bciideal
i: [1,1006,0,34,6,1,0,0] t49_bciideal
i: [1,1005,0,34,6,1,0,0] t49_bciideal
i: [1,1004,0,22,5,1,0,0] t49_bciideal
i: [1,1003,0,22,5,1,0,0] t49_bciideal
i: [1,1002,0,34,6,1,0,0] t49_bciideal
i: [1,1001,0,34,6,1,0,0] t49_bciideal
i: [1,1000,0,22,5,1,0,0] t49_bciideal
i: [1,999,0,22,5,1,0,0] t49_bciideal
i: [1,998,0,32,6,1,0,0] t49_bciideal
i: [1,997,0,32,6,1,0,0] t49_bciideal
i: [1,996,0,20,5,1,0,0] t49_bciideal
i: [1,995,0,20,5,1,0,0] t49_bciideal
i: [1,994,0,32,6,1,0,0] t49_bciideal
i: [1,993,0,32,6,1,0,0] t49_bciideal
i: [1,992,0,20,5,1,0,0] t49_bciideal
i: [1,991,0,20,5,1,0,0] t49_bciideal
i: [1,990,0,34,6,1,0,0] t49_bciideal
i: [1,989,0,34,6,1,0,0] t49_bciideal
i: [1,988,0,22,5,1,0,0] t49_bciideal
i: [1,987,0,22,5,1,0,0] t49_bciideal
i: [1,986,0,34,6,1,0,0] t49_bciideal
i: [1,985,0,34,6,1,0,0] t49_bciideal
i: [1,984,0,22,5,1,0,0] t49_bciideal
i: [1,983,0,22,5,1,0,0] t49_bciideal
i: [1,982,0,32,6,1,0,0] t49_bciideal
i: [1,981,0,32,6,1,0,0] t49_bciideal
i: [1,980,0,20,5,1,0,0] t49_bciideal
i: [1,979,0,20,5,1,0,0] t49_bciideal
i: [1,978,0,32,6,1,0,0] t49_bciideal
i: [1,977,0,32,6,1,0,0] t49_bciideal
i: [1,976,0,20,5,1,0,0] t49_bciideal
i: [1,975,0,20,5,1,0,0] t49_bciideal
i: [1,974,0,34,6,1,0,0] t49_bciideal
i: [1,973,0,34,6,1,0,0] t49_bciideal
i: [1,972,0,22,5,1,0,0] t49_bciideal
i: [1,971,0,22,5,1,0,0] t49_bciideal
i: [1,970,0,34,6,1,0,0] t49_bciideal
i: [1,969,0,34,6,1,0,0] t49_bciideal
i: [1,968,0,22,5,1,0,0] t49_bciideal
i: [1,967,0,22,5,1,0,0] t49_bciideal
i: [1,966,0,32,6,1,0,0] t49_bciideal
i: [1,965,0,32,6,1,0,0] t49_bciideal
i: [1,964,0,20,5,1,0,0] t49_bciideal
i: [1,963,0,20,5,1,0,0] t49_bciideal
i: [1,962,0,32,6,1,0,0] t49_bciideal
i: [1,961,0,32,6,1,0,0] t49_bciideal
i: [1,960,0,20,5,1,0,0] t49_bciideal
i: [1,959,0,20,5,1,0,0] t49_bciideal
i: [1,957,0,74,7,1,0,0] t49_bciideal
i: [1,956,0,62,6,1,0,0] t49_bciideal
i: [1,955,0,62,6,1,0,0] t49_bciideal
i: [1,954,0,74,7,1,0,0] t49_bciideal
i: [1,953,0,74,7,1,0,0] t49_bciideal
i: [1,952,0,62,6,1,0,0] t49_bciideal
i: [1,951,0,62,6,1,0,0] t49_bciideal
i: [1,950,0,72,7,1,0,0] t49_bciideal
i: [1,949,0,72,7,1,0,0] t49_bciideal
i: [1,948,0,60,6,1,0,0] t49_bciideal
i: [1,947,0,60,6,1,0,0] t49_bciideal
i: [1,946,0,72,7,1,0,0] t49_bciideal
i: [1,945,0,72,7,1,0,0] t49_bciideal
i: [1,944,0,60,6,1,0,0] t49_bciideal
i: [1,943,0,60,6,1,0,0] t49_bciideal
i: [1,942,0,63,7,1,0,0] t49_bciideal
i: [1,941,0,63,7,1,0,0] t49_bciideal
i: [1,940,0,51,6,1,0,0] t49_bciideal
i: [1,939,0,51,6,1,0,0] t49_bciideal
i: [1,938,0,63,7,1,0,0] t49_bciideal
i: [1,937,0,63,7,1,0,0] t49_bciideal
i: [1,936,0,51,6,1,0,0] t49_bciideal
i: [1,935,0,51,6,1,0,0] t49_bciideal
i: [1,934,0,61,7,1,0,0] t49_bciideal
i: [1,933,0,61,7,1,0,0] t49_bciideal
i: [1,932,0,49,6,1,0,0] t49_bciideal
i: [1,931,0,49,6,1,0,0] t49_bciideal
i: [1,930,0,61,7,1,0,0] t49_bciideal
i: [1,929,0,61,7,1,0,0] t49_bciideal
i: [1,928,0,49,6,1,0,0] t49_bciideal
i: [1,927,0,49,6,1,0,0] t49_bciideal
i: [1,926,0,63,7,1,0,0] t49_bciideal
i: [1,925,0,63,7,1,0,0] t49_bciideal
i: [1,924,0,51,6,1,0,0] t49_bciideal
i: [1,923,0,51,6,1,0,0] t49_bciideal
i: [1,922,0,63,7,1,0,0] t49_bciideal
i: [1,921,0,63,7,1,0,0] t49_bciideal
i: [1,920,0,51,6,1,0,0] t49_bciideal
i: [1,919,0,51,6,1,0,0] t49_bciideal
i: [1,918,0,61,7,1,0,0] t49_bciideal
i: [1,917,0,61,7,1,0,0] t49_bciideal
i: [1,916,0,49,6,1,0,0] t49_bciideal
i: [1,915,0,49,6,1,0,0] t49_bciideal
i: [1,914,0,61,7,1,0,0] t49_bciideal
i: [1,913,0,61,7,1,0,0] t49_bciideal
i: [1,912,0,49,6,1,0,0] t49_bciideal
i: [1,911,0,49,6,1,0,0] t49_bciideal
i: [1,910,0,63,7,1,0,0] t49_bciideal
i: [1,909,0,63,7,1,0,0] t49_bciideal
i: [1,908,0,51,6,1,0,0] t49_bciideal
i: [1,907,0,51,6,1,0,0] t49_bciideal
i: [1,906,0,63,7,1,0,0] t49_bciideal
i: [1,905,0,63,7,1,0,0] t49_bciideal
i: [1,904,0,51,6,1,0,0] t49_bciideal
i: [1,903,0,51,6,1,0,0] t49_bciideal
i: [1,902,0,61,7,1,0,0] t49_bciideal
i: [1,901,0,61,7,1,0,0] t49_bciideal
i: [1,900,0,49,6,1,0,0] t49_bciideal
i: [1,899,0,49,6,1,0,0] t49_bciideal
i: [1,898,0,61,7,1,0,0] t49_bciideal
i: [1,897,0,61,7,1,0,0] t49_bciideal
i: [1,896,0,49,6,1,0,0] t49_bciideal
i: [1,895,0,49,6,1,0,0] t49_bciideal
i: [1,894,0,63,7,1,0,0] t49_bciideal
i: [1,893,0,63,7,1,0,0] t49_bciideal
i: [1,892,0,51,6,1,0,0] t49_bciideal
i: [1,891,0,51,6,1,0,0] t49_bciideal
i: [1,890,0,63,7,1,0,0] t49_bciideal
i: [1,889,0,63,7,1,0,0] t49_bciideal
i: [1,888,0,51,6,1,0,0] t49_bciideal
i: [1,887,0,51,6,1,0,0] t49_bciideal
i: [1,886,0,61,7,1,0,0] t49_bciideal
i: [1,885,0,61,7,1,0,0] t49_bciideal
i: [1,884,0,49,6,1,0,0] t49_bciideal
i: [1,883,0,49,6,1,0,0] t49_bciideal
i: [1,882,0,61,7,1,0,0] t49_bciideal
i: [1,881,0,61,7,1,0,0] t49_bciideal
i: [1,880,0,49,6,1,0,0] t49_bciideal
i: [1,879,0,49,6,1,0,0] t49_bciideal
i: [1,878,0,45,6,1,0,0] t49_bciideal
i: [1,877,0,45,6,1,0,0] t49_bciideal
i: [1,876,0,33,5,1,0,0] t49_bciideal
i: [1,875,0,33,5,1,0,0] t49_bciideal
i: [1,874,0,45,6,1,0,0] t49_bciideal
i: [1,873,0,45,6,1,0,0] t49_bciideal
i: [1,872,0,33,5,1,0,0] t49_bciideal
i: [1,871,0,33,5,1,0,0] t49_bciideal
i: [1,870,0,43,6,1,0,0] t49_bciideal
i: [1,869,0,43,6,1,0,0] t49_bciideal
i: [1,868,0,31,5,1,0,0] t49_bciideal
i: [1,867,0,31,5,1,0,0] t49_bciideal
i: [1,866,0,43,6,1,0,0] t49_bciideal
i: [1,865,0,43,6,1,0,0] t49_bciideal
i: [1,864,0,31,5,1,0,0] t49_bciideal
i: [1,863,0,31,5,1,0,0] t49_bciideal
i: [1,862,0,34,6,1,0,0] t49_bciideal
i: [1,861,0,34,6,1,0,0] t49_bciideal
i: [1,860,0,22,5,1,0,0] t49_bciideal
i: [1,859,0,22,5,1,0,0] t49_bciideal
i: [1,858,0,34,6,1,0,0] t49_bciideal
i: [1,857,0,34,6,1,0,0] t49_bciideal
i: [1,856,0,22,5,1,0,0] t49_bciideal
i: [1,855,0,22,5,1,0,0] t49_bciideal
i: [1,854,0,32,6,1,0,0] t49_bciideal
i: [1,853,0,32,6,1,0,0] t49_bciideal
i: [1,852,0,20,5,1,0,0] t49_bciideal
i: [1,851,0,20,5,1,0,0] t49_bciideal
i: [1,850,0,32,6,1,0,0] t49_bciideal
i: [1,849,0,32,6,1,0,0] t49_bciideal
i: [1,848,0,20,5,1,0,0] t49_bciideal
i: [1,847,0,20,5,1,0,0] t49_bciideal
i: [1,846,0,34,6,1,0,0] t49_bciideal
i: [1,845,0,34,6,1,0,0] t49_bciideal
i: [1,844,0,22,5,1,0,0] t49_bciideal
i: [1,843,0,22,5,1,0,0] t49_bciideal
i: [1,842,0,34,6,1,0,0] t49_bciideal
i: [1,841,0,34,6,1,0,0] t49_bciideal
i: [1,840,0,22,5,1,0,0] t49_bciideal
i: [1,839,0,22,5,1,0,0] t49_bciideal
i: [1,838,0,32,6,1,0,0] t49_bciideal
i: [1,837,0,32,6,1,0,0] t49_bciideal
i: [1,836,0,20,5,1,0,0] t49_bciideal
i: [1,835,0,20,5,1,0,0] t49_bciideal
i: [1,834,0,32,6,1,0,0] t49_bciideal
i: [1,833,0,32,6,1,0,0] t49_bciideal
i: [1,832,0,20,5,1,0,0] t49_bciideal
i: [1,831,0,20,5,1,0,0] t49_bciideal
i: [1,830,0,34,6,1,0,0] t49_bciideal
i: [1,829,0,34,6,1,0,0] t49_bciideal
i: [1,828,0,22,5,1,0,0] t49_bciideal
i: [1,827,0,22,5,1,0,0] t49_bciideal
i: [1,826,0,34,6,1,0,0] t49_bciideal
i: [1,825,0,34,6,1,0,0] t49_bciideal
i: [1,824,0,22,5,1,0,0] t49_bciideal
i: [1,823,0,22,5,1,0,0] t49_bciideal
i: [1,822,0,32,6,1,0,0] t49_bciideal
i: [1,821,0,32,6,1,0,0] t49_bciideal
i: [1,820,0,20,5,1,0,0] t49_bciideal
i: [1,819,0,20,5,1,0,0] t49_bciideal
i: [1,818,0,32,6,1,0,0] t49_bciideal
i: [1,817,0,32,6,1,0,0] t49_bciideal
i: [1,816,0,20,5,1,0,0] t49_bciideal
i: [1,815,0,20,5,1,0,0] t49_bciideal
i: [1,814,0,34,6,1,0,0] t49_bciideal
i: [1,813,0,34,6,1,0,0] t49_bciideal
i: [1,812,0,22,5,1,0,0] t49_bciideal
i: [1,811,0,22,5,1,0,0] t49_bciideal
i: [1,810,0,34,6,1,0,0] t49_bciideal
i: [1,809,0,34,6,1,0,0] t49_bciideal
i: [1,808,0,22,5,1,0,0] t49_bciideal
i: [1,807,0,22,5,1,0,0] t49_bciideal
i: [1,806,0,32,6,1,0,0] t49_bciideal
i: [1,805,0,32,6,1,0,0] t49_bciideal
i: [1,804,0,20,5,1,0,0] t49_bciideal
i: [1,803,0,20,5,1,0,0] t49_bciideal
i: [1,802,0,32,6,1,0,0] t49_bciideal
i: [1,801,0,32,6,1,0,0] t49_bciideal
i: [1,800,0,20,5,1,0,0] t49_bciideal
i: [1,799,0,20,5,1,0,0] t49_bciideal
i: [1,798,0,45,6,1,0,0] t49_bciideal
i: [1,797,0,45,6,1,0,0] t49_bciideal
i: [1,796,0,33,5,1,0,0] t49_bciideal
i: [1,795,0,33,5,1,0,0] t49_bciideal
i: [1,794,0,45,6,1,0,0] t49_bciideal
i: [1,793,0,45,6,1,0,0] t49_bciideal
i: [1,792,0,33,5,1,0,0] t49_bciideal
i: [1,791,0,33,5,1,0,0] t49_bciideal
i: [1,790,0,43,6,1,0,0] t49_bciideal
i: [1,789,0,43,6,1,0,0] t49_bciideal
i: [1,788,0,31,5,1,0,0] t49_bciideal
i: [1,787,0,31,5,1,0,0] t49_bciideal
i: [1,786,0,43,6,1,0,0] t49_bciideal
i: [1,785,0,43,6,1,0,0] t49_bciideal
i: [1,784,0,31,5,1,0,0] t49_bciideal
i: [1,783,0,31,5,1,0,0] t49_bciideal
i: [1,782,0,34,6,1,0,0] t49_bciideal
i: [1,781,0,34,6,1,0,0] t49_bciideal
i: [1,780,0,22,5,1,0,0] t49_bciideal
i: [1,779,0,22,5,1,0,0] t49_bciideal
i: [1,778,0,34,6,1,0,0] t49_bciideal
i: [1,777,0,34,6,1,0,0] t49_bciideal
i: [1,776,0,22,5,1,0,0] t49_bciideal
i: [1,775,0,22,5,1,0,0] t49_bciideal
i: [1,774,0,32,6,1,0,0] t49_bciideal
i: [1,773,0,32,6,1,0,0] t49_bciideal
i: [1,772,0,20,5,1,0,0] t49_bciideal
i: [1,771,0,20,5,1,0,0] t49_bciideal
i: [1,770,0,32,6,1,0,0] t49_bciideal
i: [1,769,0,32,6,1,0,0] t49_bciideal
i: [1,768,0,20,5,1,0,0] t49_bciideal
i: [1,767,0,20,5,1,0,0] t49_bciideal
i: [1,766,0,34,6,1,0,0] t49_bciideal
i: [1,765,0,34,6,1,0,0] t49_bciideal
i: [1,764,0,22,5,1,0,0] t49_bciideal
i: [1,763,0,22,5,1,0,0] t49_bciideal
i: [1,762,0,34,6,1,0,0] t49_bciideal
i: [1,761,0,34,6,1,0,0] t49_bciideal
i: [1,760,0,22,5,1,0,0] t49_bciideal
i: [1,759,0,22,5,1,0,0] t49_bciideal
i: [1,758,0,32,6,1,0,0] t49_bciideal
i: [1,757,0,32,6,1,0,0] t49_bciideal
i: [1,756,0,20,5,1,0,0] t49_bciideal
i: [1,755,0,20,5,1,0,0] t49_bciideal
i: [1,754,0,32,6,1,0,0] t49_bciideal
i: [1,753,0,32,6,1,0,0] t49_bciideal
i: [1,752,0,20,5,1,0,0] t49_bciideal
i: [1,751,0,20,5,1,0,0] t49_bciideal
i: [1,750,0,34,6,1,0,0] t49_bciideal
i: [1,749,0,34,6,1,0,0] t49_bciideal
i: [1,748,0,22,5,1,0,0] t49_bciideal
i: [1,747,0,22,5,1,0,0] t49_bciideal
i: [1,746,0,34,6,1,0,0] t49_bciideal
i: [1,745,0,34,6,1,0,0] t49_bciideal
i: [1,744,0,22,5,1,0,0] t49_bciideal
i: [1,743,0,22,5,1,0,0] t49_bciideal
i: [1,742,0,32,6,1,0,0] t49_bciideal
i: [1,741,0,32,6,1,0,0] t49_bciideal
i: [1,740,0,20,5,1,0,0] t49_bciideal
i: [1,739,0,20,5,1,0,0] t49_bciideal
i: [1,738,0,32,6,1,0,0] t49_bciideal
i: [1,737,0,32,6,1,0,0] t49_bciideal
i: [1,736,0,20,5,1,0,0] t49_bciideal
i: [1,735,0,20,5,1,0,0] t49_bciideal
i: [1,734,0,34,6,1,0,0] t49_bciideal
i: [1,733,0,34,6,1,0,0] t49_bciideal
i: [1,732,0,22,5,1,0,0] t49_bciideal
i: [1,731,0,22,5,1,0,0] t49_bciideal
i: [1,730,0,34,6,1,0,0] t49_bciideal
i: [1,729,0,34,6,1,0,0] t49_bciideal
i: [1,728,0,22,5,1,0,0] t49_bciideal
i: [1,727,0,22,5,1,0,0] t49_bciideal
i: [1,726,0,32,6,1,0,0] t49_bciideal
i: [1,725,0,32,6,1,0,0] t49_bciideal
i: [1,724,0,20,5,1,0,0] t49_bciideal
i: [1,723,0,20,5,1,0,0] t49_bciideal
i: [1,722,0,32,6,1,0,0] t49_bciideal
i: [1,721,0,32,6,1,0,0] t49_bciideal
i: [1,720,0,20,5,1,0,0] t49_bciideal
i: [1,719,0,20,5,1,0,0] t49_bciideal
i: [1,717,0,74,7,1,0,0] t49_bciideal
i: [1,716,0,62,6,1,0,0] t49_bciideal
i: [1,715,0,62,6,1,0,0] t49_bciideal
i: [1,714,0,74,7,1,0,0] t49_bciideal
i: [1,713,0,74,7,1,0,0] t49_bciideal
i: [1,712,0,62,6,1,0,0] t49_bciideal
i: [1,711,0,62,6,1,0,0] t49_bciideal
i: [1,710,0,72,7,1,0,0] t49_bciideal
i: [1,709,0,72,7,1,0,0] t49_bciideal
i: [1,708,0,60,6,1,0,0] t49_bciideal
i: [1,707,0,60,6,1,0,0] t49_bciideal
i: [1,706,0,72,7,1,0,0] t49_bciideal
i: [1,705,0,72,7,1,0,0] t49_bciideal
i: [1,704,0,60,6,1,0,0] t49_bciideal
i: [1,703,0,60,6,1,0,0] t49_bciideal
i: [1,702,0,63,7,1,0,0] t49_bciideal
i: [1,701,0,63,7,1,0,0] t49_bciideal
i: [1,700,0,51,6,1,0,0] t49_bciideal
i: [1,699,0,51,6,1,0,0] t49_bciideal
i: [1,698,0,63,7,1,0,0] t49_bciideal
i: [1,697,0,63,7,1,0,0] t49_bciideal
i: [1,696,0,51,6,1,0,0] t49_bciideal
i: [1,695,0,51,6,1,0,0] t49_bciideal
i: [1,694,0,61,7,1,0,0] t49_bciideal
i: [1,693,0,61,7,1,0,0] t49_bciideal
i: [1,692,0,49,6,1,0,0] t49_bciideal
i: [1,691,0,49,6,1,0,0] t49_bciideal
i: [1,690,0,61,7,1,0,0] t49_bciideal
i: [1,689,0,61,7,1,0,0] t49_bciideal
i: [1,688,0,49,6,1,0,0] t49_bciideal
i: [1,687,0,49,6,1,0,0] t49_bciideal
i: [1,686,0,63,7,1,0,0] t49_bciideal
i: [1,685,0,63,7,1,0,0] t49_bciideal
i: [1,684,0,51,6,1,0,0] t49_bciideal
i: [1,683,0,51,6,1,0,0] t49_bciideal
i: [1,682,0,63,7,1,0,0] t49_bciideal
i: [1,681,0,63,7,1,0,0] t49_bciideal
i: [1,680,0,51,6,1,0,0] t49_bciideal
i: [1,679,0,51,6,1,0,0] t49_bciideal
i: [1,678,0,61,7,1,0,0] t49_bciideal
i: [1,677,0,61,7,1,0,0] t49_bciideal
i: [1,676,0,49,6,1,0,0] t49_bciideal
i: [1,675,0,49,6,1,0,0] t49_bciideal
i: [1,674,0,61,7,1,0,0] t49_bciideal
i: [1,673,0,61,7,1,0,0] t49_bciideal
i: [1,672,0,49,6,1,0,0] t49_bciideal
i: [1,671,0,49,6,1,0,0] t49_bciideal
i: [1,670,0,63,7,1,0,0] t49_bciideal
i: [1,669,0,63,7,1,0,0] t49_bciideal
i: [1,668,0,51,6,1,0,0] t49_bciideal
i: [1,667,0,51,6,1,0,0] t49_bciideal
i: [1,666,0,63,7,1,0,0] t49_bciideal
i: [1,665,0,63,7,1,0,0] t49_bciideal
i: [1,664,0,51,6,1,0,0] t49_bciideal
i: [1,663,0,51,6,1,0,0] t49_bciideal
i: [1,662,0,61,7,1,0,0] t49_bciideal
i: [1,661,0,61,7,1,0,0] t49_bciideal
i: [1,660,0,49,6,1,0,0] t49_bciideal
i: [1,659,0,49,6,1,0,0] t49_bciideal
i: [1,658,0,61,7,1,0,0] t49_bciideal
i: [1,657,0,61,7,1,0,0] t49_bciideal
i: [1,656,0,49,6,1,0,0] t49_bciideal
i: [1,655,0,49,6,1,0,0] t49_bciideal
i: [1,654,0,63,7,1,0,0] t49_bciideal
i: [1,653,0,63,7,1,0,0] t49_bciideal
i: [1,652,0,51,6,1,0,0] t49_bciideal
i: [1,651,0,51,6,1,0,0] t49_bciideal
i: [1,650,0,63,7,1,0,0] t49_bciideal
i: [1,649,0,63,7,1,0,0] t49_bciideal
i: [1,648,0,51,6,1,0,0] t49_bciideal
i: [1,647,0,51,6,1,0,0] t49_bciideal
i: [1,646,0,61,7,1,0,0] t49_bciideal
i: [1,645,0,61,7,1,0,0] t49_bciideal
i: [1,644,0,49,6,1,0,0] t49_bciideal
i: [1,643,0,49,6,1,0,0] t49_bciideal
i: [1,642,0,61,7,1,0,0] t49_bciideal
i: [1,641,0,61,7,1,0,0] t49_bciideal
i: [1,640,0,49,6,1,0,0] t49_bciideal
i: [1,639,0,49,6,1,0,0] t49_bciideal
i: [1,638,0,45,6,1,0,0] t49_bciideal
i: [1,637,0,45,6,1,0,0] t49_bciideal
i: [1,636,0,33,5,1,0,0] t49_bciideal
i: [1,635,0,33,5,1,0,0] t49_bciideal
i: [1,634,0,45,6,1,0,0] t49_bciideal
i: [1,633,0,45,6,1,0,0] t49_bciideal
i: [1,632,0,33,5,1,0,0] t49_bciideal
i: [1,631,0,33,5,1,0,0] t49_bciideal
i: [1,630,0,43,6,1,0,0] t49_bciideal
i: [1,629,0,43,6,1,0,0] t49_bciideal
i: [1,628,0,31,5,1,0,0] t49_bciideal
i: [1,627,0,31,5,1,0,0] t49_bciideal
i: [1,626,0,43,6,1,0,0] t49_bciideal
i: [1,625,0,43,6,1,0,0] t49_bciideal
i: [1,624,0,31,5,1,0,0] t49_bciideal
i: [1,623,0,31,5,1,0,0] t49_bciideal
i: [1,622,0,34,6,1,0,0] t49_bciideal
i: [1,621,0,34,6,1,0,0] t49_bciideal
i: [1,620,0,22,5,1,0,0] t49_bciideal
i: [1,619,0,22,5,1,0,0] t49_bciideal
i: [1,618,0,34,6,1,0,0] t49_bciideal
i: [1,617,0,34,6,1,0,0] t49_bciideal
i: [1,616,0,22,5,1,0,0] t49_bciideal
i: [1,615,0,22,5,1,0,0] t49_bciideal
i: [1,614,0,32,6,1,0,0] t49_bciideal
i: [1,613,0,32,6,1,0,0] t49_bciideal
i: [1,612,0,20,5,1,0,0] t49_bciideal
i: [1,611,0,20,5,1,0,0] t49_bciideal
i: [1,610,0,32,6,1,0,0] t49_bciideal
i: [1,609,0,32,6,1,0,0] t49_bciideal
i: [1,608,0,20,5,1,0,0] t49_bciideal
i: [1,607,0,20,5,1,0,0] t49_bciideal
i: [1,606,0,34,6,1,0,0] t49_bciideal
i: [1,605,0,34,6,1,0,0] t49_bciideal
i: [1,604,0,22,5,1,0,0] t49_bciideal
i: [1,603,0,22,5,1,0,0] t49_bciideal
i: [1,602,0,34,6,1,0,0] t49_bciideal
i: [1,601,0,34,6,1,0,0] t49_bciideal
i: [1,600,0,22,5,1,0,0] t49_bciideal
i: [1,599,0,22,5,1,0,0] t49_bciideal
i: [1,598,0,32,6,1,0,0] t49_bciideal
i: [1,597,0,32,6,1,0,0] t49_bciideal
i: [1,596,0,20,5,1,0,0] t49_bciideal
i: [1,595,0,20,5,1,0,0] t49_bciideal
i: [1,594,0,32,6,1,0,0] t49_bciideal
i: [1,593,0,32,6,1,0,0] t49_bciideal
i: [1,592,0,20,5,1,0,0] t49_bciideal
i: [1,591,0,20,5,1,0,0] t49_bciideal
i: [1,590,0,34,6,1,0,0] t49_bciideal
i: [1,589,0,34,6,1,0,0] t49_bciideal
i: [1,588,0,22,5,1,0,0] t49_bciideal
i: [1,587,0,22,5,1,0,0] t49_bciideal
i: [1,586,0,34,6,1,0,0] t49_bciideal
i: [1,585,0,34,6,1,0,0] t49_bciideal
i: [1,584,0,22,5,1,0,0] t49_bciideal
i: [1,583,0,22,5,1,0,0] t49_bciideal
i: [1,582,0,32,6,1,0,0] t49_bciideal
i: [1,581,0,32,6,1,0,0] t49_bciideal
i: [1,580,0,20,5,1,0,0] t49_bciideal
i: [1,579,0,20,5,1,0,0] t49_bciideal
i: [1,578,0,32,6,1,0,0] t49_bciideal
i: [1,577,0,32,6,1,0,0] t49_bciideal
i: [1,576,0,20,5,1,0,0] t49_bciideal
i: [1,575,0,20,5,1,0,0] t49_bciideal
i: [1,574,0,34,6,1,0,0] t49_bciideal
i: [1,573,0,34,6,1,0,0] t49_bciideal
i: [1,572,0,22,5,1,0,0] t49_bciideal
i: [1,571,0,22,5,1,0,0] t49_bciideal
i: [1,570,0,34,6,1,0,0] t49_bciideal
i: [1,569,0,34,6,1,0,0] t49_bciideal
i: [1,568,0,22,5,1,0,0] t49_bciideal
i: [1,567,0,22,5,1,0,0] t49_bciideal
i: [1,566,0,32,6,1,0,0] t49_bciideal
i: [1,565,0,32,6,1,0,0] t49_bciideal
i: [1,564,0,20,5,1,0,0] t49_bciideal
i: [1,563,0,20,5,1,0,0] t49_bciideal
i: [1,562,0,32,6,1,0,0] t49_bciideal
i: [1,561,0,32,6,1,0,0] t49_bciideal
i: [1,560,0,20,5,1,0,0] t49_bciideal
i: [1,559,0,20,5,1,0,0] t49_bciideal
i: [1,558,0,45,6,1,0,0] t49_bciideal
i: [1,557,0,45,6,1,0,0] t49_bciideal
i: [1,556,0,33,5,1,0,0] t49_bciideal
i: [1,555,0,33,5,1,0,0] t49_bciideal
i: [1,554,0,45,6,1,0,0] t49_bciideal
i: [1,553,0,45,6,1,0,0] t49_bciideal
i: [1,552,0,33,5,1,0,0] t49_bciideal
i: [1,551,0,33,5,1,0,0] t49_bciideal
i: [1,550,0,43,6,1,0,0] t49_bciideal
i: [1,549,0,43,6,1,0,0] t49_bciideal
i: [1,548,0,31,5,1,0,0] t49_bciideal
i: [1,547,0,31,5,1,0,0] t49_bciideal
i: [1,546,0,43,6,1,0,0] t49_bciideal
i: [1,545,0,43,6,1,0,0] t49_bciideal
i: [1,544,0,31,5,1,0,0] t49_bciideal
i: [1,543,0,31,5,1,0,0] t49_bciideal
i: [1,542,0,34,6,1,0,0] t49_bciideal
i: [1,541,0,34,6,1,0,0] t49_bciideal
i: [1,540,0,22,5,1,0,0] t49_bciideal
i: [1,539,0,22,5,1,0,0] t49_bciideal
i: [1,538,0,34,6,1,0,0] t49_bciideal
i: [1,537,0,34,6,1,0,0] t49_bciideal
i: [1,536,0,22,5,1,0,0] t49_bciideal
i: [1,535,0,22,5,1,0,0] t49_bciideal
i: [1,534,0,32,6,1,0,0] t49_bciideal
i: [1,533,0,32,6,1,0,0] t49_bciideal
i: [1,532,0,20,5,1,0,0] t49_bciideal
i: [1,531,0,20,5,1,0,0] t49_bciideal
i: [1,530,0,32,6,1,0,0] t49_bciideal
i: [1,529,0,32,6,1,0,0] t49_bciideal
i: [1,528,0,20,5,1,0,0] t49_bciideal
i: [1,527,0,20,5,1,0,0] t49_bciideal
i: [1,526,0,34,6,1,0,0] t49_bciideal
i: [1,525,0,34,6,1,0,0] t49_bciideal
i: [1,524,0,22,5,1,0,0] t49_bciideal
i: [1,523,0,22,5,1,0,0] t49_bciideal
i: [1,522,0,34,6,1,0,0] t49_bciideal
i: [1,521,0,34,6,1,0,0] t49_bciideal
i: [1,520,0,22,5,1,0,0] t49_bciideal
i: [1,519,0,22,5,1,0,0] t49_bciideal
i: [1,518,0,32,6,1,0,0] t49_bciideal
i: [1,517,0,32,6,1,0,0] t49_bciideal
i: [1,516,0,20,5,1,0,0] t49_bciideal
i: [1,515,0,20,5,1,0,0] t49_bciideal
i: [1,514,0,32,6,1,0,0] t49_bciideal
i: [1,513,0,32,6,1,0,0] t49_bciideal
i: [1,512,0,20,5,1,0,0] t49_bciideal
i: [1,511,0,20,5,1,0,0] t49_bciideal
i: [1,510,0,34,6,1,0,0] t49_bciideal
i: [1,509,0,34,6,1,0,0] t49_bciideal
i: [1,508,0,22,5,1,0,0] t49_bciideal
i: [1,507,0,22,5,1,0,0] t49_bciideal
i: [1,506,0,34,6,1,0,0] t49_bciideal
i: [1,505,0,34,6,1,0,0] t49_bciideal
i: [1,504,0,22,5,1,0,0] t49_bciideal
i: [1,503,0,22,5,1,0,0] t49_bciideal
i: [1,502,0,32,6,1,0,0] t49_bciideal
i: [1,501,0,32,6,1,0,0] t49_bciideal
i: [1,500,0,20,5,1,0,0] t49_bciideal
i: [1,499,0,20,5,1,0,0] t49_bciideal
i: [1,498,0,32,6,1,0,0] t49_bciideal
i: [1,497,0,32,6,1,0,0] t49_bciideal
i: [1,496,0,20,5,1,0,0] t49_bciideal
i: [1,495,0,20,5,1,0,0] t49_bciideal
i: [1,494,0,34,6,1,0,0] t49_bciideal
i: [1,493,0,34,6,1,0,0] t49_bciideal
i: [1,492,0,22,5,1,0,0] t49_bciideal
i: [1,491,0,22,5,1,0,0] t49_bciideal
i: [1,490,0,34,6,1,0,0] t49_bciideal
i: [1,489,0,34,6,1,0,0] t49_bciideal
i: [1,488,0,22,5,1,0,0] t49_bciideal
i: [1,487,0,22,5,1,0,0] t49_bciideal
i: [1,486,0,32,6,1,0,0] t49_bciideal
i: [1,485,0,32,6,1,0,0] t49_bciideal
i: [1,484,0,20,5,1,0,0] t49_bciideal
i: [1,483,0,20,5,1,0,0] t49_bciideal
i: [1,482,0,32,6,1,0,0] t49_bciideal
i: [1,481,0,32,6,1,0,0] t49_bciideal
i: [1,480,0,20,5,1,0,0] t49_bciideal
i: [1,479,0,20,5,1,0,0] t49_bciideal
i: [1,477,0,74,7,1,0,0] t49_bciideal
i: [1,476,0,62,6,1,0,0] t49_bciideal
i: [1,475,0,62,6,1,0,0] t49_bciideal
i: [1,474,0,74,7,1,0,0] t49_bciideal
i: [1,473,0,74,7,1,0,0] t49_bciideal
i: [1,472,0,62,6,1,0,0] t49_bciideal
i: [1,471,0,62,6,1,0,0] t49_bciideal
i: [1,470,0,72,7,1,0,0] t49_bciideal
i: [1,469,0,72,7,1,0,0] t49_bciideal
i: [1,468,0,60,6,1,0,0] t49_bciideal
i: [1,467,0,60,6,1,0,0] t49_bciideal
i: [1,466,0,72,7,1,0,0] t49_bciideal
i: [1,465,0,72,7,1,0,0] t49_bciideal
i: [1,464,0,60,6,1,0,0] t49_bciideal
i: [1,463,0,60,6,1,0,0] t49_bciideal
i: [1,462,0,63,7,1,0,0] t49_bciideal
i: [1,461,0,63,7,1,0,0] t49_bciideal
i: [1,460,0,51,6,1,0,0] t49_bciideal
i: [1,459,0,51,6,1,0,0] t49_bciideal
i: [1,458,0,63,7,1,0,0] t49_bciideal
i: [1,457,0,63,7,1,0,0] t49_bciideal
i: [1,456,0,51,6,1,0,0] t49_bciideal
i: [1,455,0,51,6,1,0,0] t49_bciideal
i: [1,454,0,61,7,1,0,0] t49_bciideal
i: [1,453,0,61,7,1,0,0] t49_bciideal
i: [1,452,0,49,6,1,0,0] t49_bciideal
i: [1,451,0,49,6,1,0,0] t49_bciideal
i: [1,450,0,61,7,1,0,0] t49_bciideal
i: [1,449,0,61,7,1,0,0] t49_bciideal
i: [1,448,0,49,6,1,0,0] t49_bciideal
i: [1,447,0,49,6,1,0,0] t49_bciideal
i: [1,446,0,63,7,1,0,0] t49_bciideal
i: [1,445,0,63,7,1,0,0] t49_bciideal
i: [1,444,0,51,6,1,0,0] t49_bciideal
i: [1,443,0,51,6,1,0,0] t49_bciideal
i: [1,442,0,63,7,1,0,0] t49_bciideal
i: [1,441,0,63,7,1,0,0] t49_bciideal
i: [1,440,0,51,6,1,0,0] t49_bciideal
i: [1,439,0,51,6,1,0,0] t49_bciideal
i: [1,438,0,61,7,1,0,0] t49_bciideal
i: [1,437,0,61,7,1,0,0] t49_bciideal
i: [1,436,0,49,6,1,0,0] t49_bciideal
i: [1,435,0,49,6,1,0,0] t49_bciideal
i: [1,434,0,61,7,1,0,0] t49_bciideal
i: [1,433,0,61,7,1,0,0] t49_bciideal
i: [1,432,0,49,6,1,0,0] t49_bciideal
i: [1,431,0,49,6,1,0,0] t49_bciideal
i: [1,430,0,63,7,1,0,0] t49_bciideal
i: [1,429,0,63,7,1,0,0] t49_bciideal
i: [1,428,0,51,6,1,0,0] t49_bciideal
i: [1,427,0,51,6,1,0,0] t49_bciideal
i: [1,426,0,63,7,1,0,0] t49_bciideal
i: [1,425,0,63,7,1,0,0] t49_bciideal
i: [1,424,0,51,6,1,0,0] t49_bciideal
i: [1,423,0,51,6,1,0,0] t49_bciideal
i: [1,422,0,61,7,1,0,0] t49_bciideal
i: [1,421,0,61,7,1,0,0] t49_bciideal
i: [1,420,0,49,6,1,0,0] t49_bciideal
i: [1,419,0,49,6,1,0,0] t49_bciideal
i: [1,418,0,61,7,1,0,0] t49_bciideal
i: [1,417,0,61,7,1,0,0] t49_bciideal
i: [1,416,0,49,6,1,0,0] t49_bciideal
i: [1,415,0,49,6,1,0,0] t49_bciideal
i: [1,414,0,63,7,1,0,0] t49_bciideal
i: [1,413,0,63,7,1,0,0] t49_bciideal
i: [1,412,0,51,6,1,0,0] t49_bciideal
i: [1,411,0,51,6,1,0,0] t49_bciideal
i: [1,410,0,63,7,1,0,0] t49_bciideal
i: [1,409,0,63,7,1,0,0] t49_bciideal
i: [1,408,0,51,6,1,0,0] t49_bciideal
i: [1,407,0,51,6,1,0,0] t49_bciideal
i: [1,406,0,61,7,1,0,0] t49_bciideal
i: [1,405,0,61,7,1,0,0] t49_bciideal
i: [1,404,0,49,6,1,0,0] t49_bciideal
i: [1,403,0,49,6,1,0,0] t49_bciideal
i: [1,402,0,61,7,1,0,0] t49_bciideal
i: [1,401,0,61,7,1,0,0] t49_bciideal
i: [1,400,0,49,6,1,0,0] t49_bciideal
i: [1,399,0,49,6,1,0,0] t49_bciideal
i: [1,398,0,45,6,1,0,0] t49_bciideal
i: [1,397,0,45,6,1,0,0] t49_bciideal
i: [1,396,0,33,5,1,0,0] t49_bciideal
i: [1,395,0,33,5,1,0,0] t49_bciideal
i: [1,394,0,45,6,1,0,0] t49_bciideal
i: [1,393,0,45,6,1,0,0] t49_bciideal
i: [1,392,0,33,5,1,0,0] t49_bciideal
i: [1,391,0,33,5,1,0,0] t49_bciideal
i: [1,390,0,43,6,1,0,0] t49_bciideal
i: [1,389,0,43,6,1,0,0] t49_bciideal
i: [1,388,0,31,5,1,0,0] t49_bciideal
i: [1,387,0,31,5,1,0,0] t49_bciideal
i: [1,386,0,43,6,1,0,0] t49_bciideal
i: [1,385,0,43,6,1,0,0] t49_bciideal
i: [1,384,0,31,5,1,0,0] t49_bciideal
i: [1,383,0,31,5,1,0,0] t49_bciideal
i: [1,382,0,34,6,1,0,0] t49_bciideal
i: [1,381,0,34,6,1,0,0] t49_bciideal
i: [1,380,0,22,5,1,0,0] t49_bciideal
i: [1,379,0,22,5,1,0,0] t49_bciideal
i: [1,378,0,34,6,1,0,0] t49_bciideal
i: [1,377,0,34,6,1,0,0] t49_bciideal
i: [1,376,0,22,5,1,0,0] t49_bciideal
i: [1,375,0,22,5,1,0,0] t49_bciideal
i: [1,374,0,32,6,1,0,0] t49_bciideal
i: [1,373,0,32,6,1,0,0] t49_bciideal
i: [1,372,0,20,5,1,0,0] t49_bciideal
i: [1,371,0,20,5,1,0,0] t49_bciideal
i: [1,370,0,32,6,1,0,0] t49_bciideal
i: [1,369,0,32,6,1,0,0] t49_bciideal
i: [1,368,0,20,5,1,0,0] t49_bciideal
i: [1,367,0,20,5,1,0,0] t49_bciideal
i: [1,366,0,34,6,1,0,0] t49_bciideal
i: [1,365,0,34,6,1,0,0] t49_bciideal
i: [1,364,0,22,5,1,0,0] t49_bciideal
i: [1,363,0,22,5,1,0,0] t49_bciideal
i: [1,362,0,34,6,1,0,0] t49_bciideal
i: [1,361,0,34,6,1,0,0] t49_bciideal
i: [1,360,0,22,5,1,0,0] t49_bciideal
i: [1,359,0,22,5,1,0,0] t49_bciideal
i: [1,358,0,32,6,1,0,0] t49_bciideal
i: [1,357,0,32,6,1,0,0] t49_bciideal
i: [1,356,0,20,5,1,0,0] t49_bciideal
i: [1,355,0,20,5,1,0,0] t49_bciideal
i: [1,354,0,32,6,1,0,0] t49_bciideal
i: [1,353,0,32,6,1,0,0] t49_bciideal
i: [1,352,0,20,5,1,0,0] t49_bciideal
i: [1,351,0,20,5,1,0,0] t49_bciideal
i: [1,350,0,34,6,1,0,0] t49_bciideal
i: [1,349,0,34,6,1,0,0] t49_bciideal
i: [1,348,0,22,5,1,0,0] t49_bciideal
i: [1,347,0,22,5,1,0,0] t49_bciideal
i: [1,346,0,34,6,1,0,0] t49_bciideal
i: [1,345,0,34,6,1,0,0] t49_bciideal
i: [1,344,0,22,5,1,0,0] t49_bciideal
i: [1,343,0,22,5,1,0,0] t49_bciideal
i: [1,342,0,32,6,1,0,0] t49_bciideal
i: [1,341,0,32,6,1,0,0] t49_bciideal
i: [1,340,0,20,5,1,0,0] t49_bciideal
i: [1,339,0,20,5,1,0,0] t49_bciideal
i: [1,338,0,32,6,1,0,0] t49_bciideal
i: [1,337,0,32,6,1,0,0] t49_bciideal
i: [1,336,0,20,5,1,0,0] t49_bciideal
i: [1,335,0,20,5,1,0,0] t49_bciideal
i: [1,334,0,34,6,1,0,0] t49_bciideal
i: [1,333,0,34,6,1,0,0] t49_bciideal
i: [1,332,0,22,5,1,0,0] t49_bciideal
i: [1,331,0,22,5,1,0,0] t49_bciideal
i: [1,330,0,34,6,1,0,0] t49_bciideal
i: [1,329,0,34,6,1,0,0] t49_bciideal
i: [1,328,0,22,5,1,0,0] t49_bciideal
i: [1,327,0,22,5,1,0,0] t49_bciideal
i: [1,326,0,32,6,1,0,0] t49_bciideal
i: [1,325,0,32,6,1,0,0] t49_bciideal
i: [1,324,0,20,5,1,0,0] t49_bciideal
i: [1,323,0,20,5,1,0,0] t49_bciideal
i: [1,322,0,32,6,1,0,0] t49_bciideal
i: [1,321,0,32,6,1,0,0] t49_bciideal
i: [1,320,0,20,5,1,0,0] t49_bciideal
i: [1,319,0,20,5,1,0,0] t49_bciideal
i: [1,318,0,45,6,1,0,0] t49_bciideal
i: [1,317,0,45,6,1,0,0] t49_bciideal
i: [1,316,0,33,5,1,0,0] t49_bciideal
i: [1,315,0,33,5,1,0,0] t49_bciideal
i: [1,314,0,45,6,1,0,0] t49_bciideal
i: [1,313,0,45,6,1,0,0] t49_bciideal
i: [1,312,0,33,5,1,0,0] t49_bciideal
i: [1,311,0,33,5,1,0,0] t49_bciideal
i: [1,310,0,43,6,1,0,0] t49_bciideal
i: [1,309,0,43,6,1,0,0] t49_bciideal
i: [1,308,0,31,5,1,0,0] t49_bciideal
i: [1,307,0,31,5,1,0,0] t49_bciideal
i: [1,306,0,43,6,1,0,0] t49_bciideal
i: [1,305,0,43,6,1,0,0] t49_bciideal
i: [1,304,0,31,5,1,0,0] t49_bciideal
i: [1,303,0,31,5,1,0,0] t49_bciideal
i: [1,302,0,34,6,1,0,0] t49_bciideal
i: [1,301,0,34,6,1,0,0] t49_bciideal
i: [1,300,0,22,5,1,0,0] t49_bciideal
i: [1,299,0,22,5,1,0,0] t49_bciideal
i: [1,298,0,34,6,1,0,0] t49_bciideal
i: [1,297,0,34,6,1,0,0] t49_bciideal
i: [1,296,0,22,5,1,0,0] t49_bciideal
i: [1,295,0,22,5,1,0,0] t49_bciideal
i: [1,294,0,32,6,1,0,0] t49_bciideal
i: [1,293,0,32,6,1,0,0] t49_bciideal
i: [1,292,0,20,5,1,0,0] t49_bciideal
i: [1,291,0,20,5,1,0,0] t49_bciideal
i: [1,290,0,32,6,1,0,0] t49_bciideal
i: [1,289,0,32,6,1,0,0] t49_bciideal
i: [1,288,0,20,5,1,0,0] t49_bciideal
i: [1,287,0,20,5,1,0,0] t49_bciideal
i: [1,286,0,34,6,1,0,0] t49_bciideal
i: [1,285,0,34,6,1,0,0] t49_bciideal
i: [1,284,0,22,5,1,0,0] t49_bciideal
i: [1,283,0,22,5,1,0,0] t49_bciideal
i: [1,282,0,34,6,1,0,0] t49_bciideal
i: [1,281,0,34,6,1,0,0] t49_bciideal
i: [1,280,0,22,5,1,0,0] t49_bciideal
i: [1,279,0,22,5,1,0,0] t49_bciideal
i: [1,278,0,32,6,1,0,0] t49_bciideal
i: [1,277,0,32,6,1,0,0] t49_bciideal
i: [1,276,0,20,5,1,0,0] t49_bciideal
i: [1,275,0,20,5,1,0,0] t49_bciideal
i: [1,274,0,32,6,1,0,0] t49_bciideal
i: [1,273,0,32,6,1,0,0] t49_bciideal
i: [1,272,0,20,5,1,0,0] t49_bciideal
i: [1,271,0,20,5,1,0,0] t49_bciideal
i: [1,270,0,34,6,1,0,0] t49_bciideal
i: [1,269,0,34,6,1,0,0] t49_bciideal
i: [1,268,0,22,5,1,0,0] t49_bciideal
i: [1,267,0,22,5,1,0,0] t49_bciideal
i: [1,266,0,34,6,1,0,0] t49_bciideal
i: [1,265,0,34,6,1,0,0] t49_bciideal
i: [1,264,0,22,5,1,0,0] t49_bciideal
i: [1,263,0,22,5,1,0,0] t49_bciideal
i: [1,262,0,32,6,1,0,0] t49_bciideal
i: [1,261,0,32,6,1,0,0] t49_bciideal
i: [1,260,0,20,5,1,0,0] t49_bciideal
i: [1,259,0,20,5,1,0,0] t49_bciideal
i: [1,258,0,32,6,1,0,0] t49_bciideal
i: [1,257,0,32,6,1,0,0] t49_bciideal
i: [1,256,0,20,5,1,0,0] t49_bciideal
i: [1,255,0,20,5,1,0,0] t49_bciideal
i: [1,254,0,34,6,1,0,0] t49_bciideal
i: [1,253,0,34,6,1,0,0] t49_bciideal
i: [1,252,0,22,5,1,0,0] t49_bciideal
i: [1,251,0,22,5,1,0,0] t49_bciideal
i: [1,250,0,34,6,1,0,0] t49_bciideal
i: [1,249,0,34,6,1,0,0] t49_bciideal
i: [1,248,0,22,5,1,0,0] t49_bciideal
i: [1,247,0,22,5,1,0,0] t49_bciideal
i: [1,246,0,32,6,1,0,0] t49_bciideal
i: [1,245,0,32,6,1,0,0] t49_bciideal
i: [1,244,0,20,5,1,0,0] t49_bciideal
i: [1,243,0,20,5,1,0,0] t49_bciideal
i: [1,242,0,32,6,1,0,0] t49_bciideal
i: [1,241,0,32,6,1,0,0] t49_bciideal
i: [1,240,0,20,5,1,0,0] t49_bciideal
i: [1,239,0,20,5,1,0,0] t49_bciideal
d: [2,3337,2,16,8,1,40,3335,3319]
d: [2,3338,2,14,7,2,21,3337,2844]
d: [2,3339,2,12,6,3,21,3338,2849]
d: [2,3340,2,10,5,4,21,3339,2854]
d: [2,3341,2,8,4,5,21,3340,2859]
d: [2,3342,2,6,3,6,21,3341,2864]
d: [2,3343,2,4,2,7,21,3342,2869]
d: [2,3344,2,2,1,8,21,3343,2874]
d: [2,3529,1,37,12,1,44,1548,1555]
d: [2,3534,1,34,11,1,19,3529]
d: [2,3550,1,30,10,2,21,3534,1586]
d: [2,3551,1,26,9,3,21,3550,1591]
d: [2,3552,1,24,8,4,21,3551,2844]
d: [2,3553,1,22,7,5,21,3552,3323]
d: [2,3554,1,20,6,6,21,3553,2849]
d: [2,3555,1,18,5,7,21,3554,2854]
d: [2,3556,1,16,4,8,21,3555,2859]
d: [2,3557,1,14,3,9,21,3556,2864]
d: [2,3558,1,12,2,10,21,3557,2869]
d: [2,3559,1,10,1,11,21,3558,2874]
d: [2,3385,1,33,11,1,44,1572,1551]
d: [2,3390,1,24,10,1,19,3385]
d: [2,3583,1,22,9,2,21,3390,2844]
d: [2,3584,1,20,8,3,21,3583,3323]
d: [2,3585,1,18,7,4,21,3584,2849]
d: [2,3586,1,16,6,5,21,3585,2854]
d: [2,3587,1,14,5,6,21,3586,2859]
d: [2,3588,1,12,4,7,21,3587,2864]
d: [2,3589,1,10,3,8,21,3588,2869]
d: [2,3590,1,8,2,9,21,3589,2874]
d: [2,3747,1,35,12,1,44,1549,1559]
d: [2,3755,1,30,11,1,19,3747]
d: [2,3756,1,26,10,2,21,3755,1586]
d: [2,3757,1,22,9,3,21,3756,1591]
d: [2,3758,1,20,8,4,21,3757,2844]
d: [2,3759,1,18,7,5,21,3758,3323]
d: [2,3760,1,16,6,6,21,3759,2849]
d: [2,3761,1,14,5,7,21,3760,2854]
d: [2,3762,1,12,4,8,21,3761,2859]
d: [2,3763,1,10,3,9,21,3762,2864]
d: [2,3764,1,8,2,10,21,3763,2869]
d: [2,3765,1,6,1,11,21,3764,2874]
d: [2,3362,1,33,12,1,44,1550,1596]
d: [2,3365,1,30,11,1,19,3362]
d: [2,4138,1,24,10,2,21,3365,1564]
d: [2,4139,1,22,9,3,21,4138,2844]
d: [2,4140,1,20,8,4,21,4139,3323]
d: [2,4141,1,18,7,5,21,4140,2849]
d: [2,4142,1,16,6,6,21,4141,2854]
d: [2,4143,1,14,5,7,21,4142,2859]
d: [2,4144,1,12,4,8,21,4143,2864]
d: [2,4145,1,10,3,9,21,4144,2869]
d: [2,4146,1,8,2,10,21,4145,2874]
d: [2,4205,1,20,9,4,21,4099,2356]
d: [2,4206,1,18,8,5,21,4205,2844]
d: [2,4207,1,16,7,6,21,4206,3323]
d: [2,4208,1,14,6,7,21,4207,2849]
d: [2,4209,1,12,5,8,21,4208,2854]
d: [2,4210,1,10,4,9,21,4209,2859]
d: [2,4211,1,8,3,10,21,4210,2864]
d: [2,4212,1,6,2,11,21,4211,2869]
d: [2,4213,1,4,1,12,21,4212,2874]
f: 4213,4146,3989,3765,3732,3590,3559,3344,3315,230,231,232,233,234,235,236,239,240,241,242,243,244,245,246,247,248,249,250,251,252,253,254,255,256,257,258,259,260,261,262,263,264,265,266,267,268,269,270,271,272,273,274,275,276,277,278,279,280,281,282,283,284,285,286,287,288,289,290,291,292,293,294,295,296,297,298,299,300,301,302,303,304,305,306,307,308,309,310,311,312,313,314,315,316,317,318,319,320,321,322,323,324,325,326,327,328,329,330,331,332,333,334,335,336,337,338,339,340,341,342,343,344,345,346,347,348,349,350,351,352,353,354,355,356,357,358,359,360,361,362,363,364,365,366,367,368,369,370,371,372,373,374,375,376,377,378,379,380,381,382,383,384,385,386,387,388,389,390,391,392,393,394,395,396,397,398,399,400,401,402,403,404,405,406,407,408,409,410,411,412,413,414,415,416,417,418,419,420,421,422,423,424,425,426,427,428,429,430,431,432,433,434,435,436,437,438,439,440,441,442,443,444,445,446,447,448,449,450,451,452,453,454,455,456,457,458,459,460,461,462,463,464,465,466,467,468,469,470,471,472,473,474,475,476,477,478,479,480,481,482,483,484,485,486,487,488,489,490,491,492,493,494,495,496,497,498,499,500,501,502,503,504,505,506,507,508,509,510,511,512,513,514,515,516,517,518,519,520,521,522,523,524,525,526,527,528,529,530,531,532,533,534,535,536,537,538,539,540,541,542,543,544,545,546,547,548,549,550,551,552,553,554,555,556,557,558,559,560,561,562,563,564,565,566,567,568,569,570,571,572,573,574,575,576,577,578,579,580,581,582,583,584,585,586,587,588,589,590,591,592,593,594,595,596,597,598,599,600,601,602,603,604,605,606,607,608,609,610,611,612,613,614,615,616,617,618,619,620,621,622,623,624,625,626,627,628,629,630,631,632,633,634,635,636,637,638,639,640,641,642,643,644,645,646,647,648,649,650,651,652,653,654,655,656,657,658,659,660,661,662,663,664,665,666,667,668,669,670,671,672,673,674,675,676,677,678,679,680,681,682,683,684,685,686,687,688,689,690,691,692,693,694,695,696,697,698,699,700,701,702,703,704,705,706,707,708,709,710,711,712,713,714,715,716,717,718,719,720,721,722,723,724,725,726,727,728,729,730,731,732,733,734,735,736,737,738,739,740,741,742,743,744,745,746,747,748,749,750,751,752,753,754,755,756,757,758,759,760,761,762,763,764,765,766,767,768,769,770,771,772,773,774,775,776,777,778,779,780,781,782,783,784,785,786,787,788,789,790,791,792,793,794,795,796,797,798,799,800,801,802,803,804,805,806,807,808,809,810,811,812,813,814,815,816,817,818,819,820,821,822,823,824,825,826,827,828,829,830,831,832,833,834,835,836,837,838,839,840,841,842,843,844,845,846,847,848,849,850,851,852,853,854,855,856,857,858,859,860,861,862,863,864,865,866,867,868,869,870,871,872,873,874,875,876,877,878,879,880,881,882,883,884,885,886,887,888,889,890,891,892,893,894,895,896,897,898,899,900,901,902,903,904,905,906,907,908,909,910,911,912,913,914,915,916,917,918,919,920,921,922,923,924,925,926,927,928,929,930,931,932,933,934,935,936,937,938,939,940,941,942,943,944,945,946,947,948,949,950,951,952,953,954,955,956,957,958,959,960,961,962,963,964,965,966,967,968,969,970,971,972,973,974,975,976,977,978,979,980,981,982,983,984,985,986,987,988,989,990,991,992,993,994,995,996,997,998,999,1000,1001,1002,1003,1004,1005,1006,1007,1008,1009,1010,1011,1012,1013,1014,1015,1016,1017,1018,1019,1020,1021,1022,1023,1024,1025,1026,1027,1028,1029,1030,1031,1032,1033,1034,1035,1036,1037,1038,1039,1040,1041,1042,1043,1044,1045,1046,1047,1048,1049,1050,1051,1052,1053,1054,1055,1056,1057,1058,1059,1060,1061,1062,1063,1064,1065,1066,1067,1068,1069,1070,1071,1072,1073,1074,1075,1076,1077,1078,1079,1080,1081,1082,1083,1084,1085,1086,1087,1088,1089,1090,1091,1092,1093,1094,1095,1096,1097,1098,1099,1100,1101,1102,1103,1104,1105,1106,1107,1108,1109,1110,1111,1112,1113,1114,1115,1116,1117,1118,1119,1120,1121,1122,1123,1124,1125,1126,1127,1128,1129,1130,1131,1132,1133,1134,1135,1136,1137,1138,1139,1140,1141,1142,1143,1144,1145,1146,1147,1148,1149,1150,1151,1152,1153,1154,1155,1156,1157,1158,1159,1160,1161,1162,1163,1164,1165,1166,1167,1168,1169,1170,1171,1172,1173,1174,1175,1176,1177,1178,1179,1180,1181,1182,1183,1184,1185,1186,1187,1188,1189,1190,1191,1192,1193,1194,1195,1196,1197,1198,1199,1200,1201,1202,1203,1204,1205,1206,1207,1208,1209,1210,1211,1212,1213,1214,1215,1216,1217,1218,1219,1220,1221,1222,1223,1224,1225,1226,1227,1228,1229,1230,1231,1232,1233,1234,1235,1236,1237,1238,1239,1240,1241,1242,1243,1244,1245,1246,1247,1248,1249,1250,1251,1252,1253,1254,1255,1256,1257,1258,1259,1260,1261,1262,1263,1264,1265,1266,1267,1268,1269,1270,1271,1272,1273,1274,1275,1276,1277,1278,1279,1280,1281,1282,1283,1284,1285,1286,1287,1288,1289,1290,1291,1292,1293,1294,1295,1296,1297,1298,1299,1300,1301,1302,1303,1304,1305,1306,1307,1308,1309,1310,1311,1312,1313,1314,1315,1316,1317,1318,1319,1320,1321,1322,1323,1324,1325,1326,1327,1328,1329,1330,1331,1332,1333,1334,1335,1336,1337,1338,1339,1340,1341,1342,1343,1344,1345,1346,1347,1348,1349,1350,1351,1352,1353,1354,1355,1356,1357,1358,1359,1360,1361,1362,1363,1364,1365,1366,1367,1368,1369,1370,1371,1372,1373,1374,1375,1376,1377,1378,1379,1380,1381,1382,1383,1384,1385,1386,1387,1388,1389,1390,1391,1392,1393,1394,1395,1396,1397,1398,1399,1400,1401,1402,1403,1404,1405,1406,1407,1408,1409,1410,1411,1412,1413,1414,1415,1416,1417,1418,1419,1420,1421,1422,1423,1424,1425,1426,1427,1428,1429,1430,1431,1432,1433,1434,1435,1436,1437,1438
% Refutation found. Thanks to Tanya!
% SZS status Theorem for bciideal__t49_bciideal
% SZS output start Proof for bciideal__t49_bciideal
1. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => (! [X1] : (m2_bcialg_1(X1,X0) => (m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0))) => (! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X1,X3) => k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0))))))) [input t49_bciideal]
2. ~! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => (! [X1] : (m2_bcialg_1(X1,X0) => (m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0))) => (! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X1,X3) => k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))) & ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0))))))) [negated conjecture 1]
60. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))))))) [input t10_bcialg_3]
64. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) <=> ! [X1] : (m2_bcialg_1(X1,X0) => (m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0))))) [input t37_bciideal]
69. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 => k1_bcialg_1(X0,X2,X1) = X2))))) [input t7_bcialg_3]
71. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X3) => k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)))))))) [input t8_bcialg_3]
73. ! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k1_bcialg_1(X0,X1,X2) = X1 <=> k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0)))))) [input t9_bcialg_3]
74. ~! [X0] : ((l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => (! [X1] : (m2_bcialg_1(X1,X0) => (m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0))) => (! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r1_bcialg_1(X0,X2,X4) => k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X4,X2))) = k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X2,X3)))))) & ! [X5] : (m1_subset_1(X5,u1_struct_0(X0)) => ! [X6] : (m1_subset_1(X6,u1_struct_0(X0)) => (k1_bcialg_1(X0,X5,X6) = k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),X5)) & k1_bcialg_1(X0,X5,X6) = k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X6,k1_bcialg_1(X0,X6,X5)))))) & ! [X7] : (m1_subset_1(X7,u1_struct_0(X0)) => ! [X8] : (m1_subset_1(X8,u1_struct_0(X0)) => ! [X9] : (m1_subset_1(X9,u1_struct_0(X0)) => (r1_bcialg_1(X0,X8,X9) => k1_bcialg_1(X0,X8,X7) = k1_bcialg_1(X0,k1_bcialg_1(X0,X9,X7),k1_bcialg_1(X0,X9,X8)))))) & ! [X10] : (m1_subset_1(X10,u1_struct_0(X0)) => ! [X11] : (m1_subset_1(X11,u1_struct_0(X0)) => (k1_bcialg_1(X0,X10,X11) = X10 => k1_bcialg_1(X0,X11,X10) = X11))) & ! [X12] : (m1_subset_1(X12,u1_struct_0(X0)) => ! [X13] : (m1_subset_1(X13,u1_struct_0(X0)) => (k1_bcialg_1(X0,X12,X13) = X12 <=> k4_struct_0(X0) = k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12)))))))) [rectify 2]
88. ? [X0] : (((? [X2] : (? [X3] : (? [X4] : ((k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X4,X2))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X2,X3)) & r1_bcialg_1(X0,X2,X4)) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) | ? [X5] : (? [X6] : ((k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),X5)) | k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X6,k1_bcialg_1(X0,X6,X5)))) & m1_subset_1(X6,u1_struct_0(X0))) & m1_subset_1(X5,u1_struct_0(X0))) | ? [X7] : (? [X8] : (? [X9] : ((k1_bcialg_1(X0,X8,X7) != k1_bcialg_1(X0,k1_bcialg_1(X0,X9,X7),k1_bcialg_1(X0,X9,X8)) & r1_bcialg_1(X0,X8,X9)) & m1_subset_1(X9,u1_struct_0(X0))) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) | ? [X10] : (? [X11] : ((k1_bcialg_1(X0,X11,X10) != X11 & k1_bcialg_1(X0,X10,X11) = X10) & m1_subset_1(X11,u1_struct_0(X0))) & m1_subset_1(X10,u1_struct_0(X0))) | ? [X12] : (? [X13] : ((k1_bcialg_1(X0,X12,X13) = X12 <~> k4_struct_0(X0) = k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12))) & m1_subset_1(X13,u1_struct_0(X0))) & m1_subset_1(X12,u1_struct_0(X0)))) & ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0))) & (l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0))) [ennf transformation 74]
89. ? [X0] : ((? [X2] : (? [X3] : (? [X4] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X4,X2))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X2,X3)) & r1_bcialg_1(X0,X2,X4) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) | ? [X5] : (? [X6] : ((k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),X5)) | k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X6,k1_bcialg_1(X0,X6,X5)))) & m1_subset_1(X6,u1_struct_0(X0))) & m1_subset_1(X5,u1_struct_0(X0))) | ? [X7] : (? [X8] : (? [X9] : (k1_bcialg_1(X0,X8,X7) != k1_bcialg_1(X0,k1_bcialg_1(X0,X9,X7),k1_bcialg_1(X0,X9,X8)) & r1_bcialg_1(X0,X8,X9) & m1_subset_1(X9,u1_struct_0(X0))) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) | ? [X10] : (? [X11] : (k1_bcialg_1(X0,X11,X10) != X11 & k1_bcialg_1(X0,X10,X11) = X10 & m1_subset_1(X11,u1_struct_0(X0))) & m1_subset_1(X10,u1_struct_0(X0))) | ? [X12] : (? [X13] : ((k1_bcialg_1(X0,X12,X13) = X12 <~> k4_struct_0(X0) = k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12))) & m1_subset_1(X13,u1_struct_0(X0))) & m1_subset_1(X12,u1_struct_0(X0)))) & ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0)) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [flattening 88]
129. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 60]
130. ! [X0] : (! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) & k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 129]
131. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 69]
132. ! [X0] : (! [X1] : (! [X2] : (k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 131]
133. ! [X0] : ((! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = X1 <=> k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 73]
134. ! [X0] : (! [X1] : (! [X2] : ((k1_bcialg_1(X0,X1,X2) = X1 <=> k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 133]
135. ! [X0] : ((! [X1] : (! [X2] : (! [X3] : ((k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 71]
136. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 135]
139. ! [X0] : (((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) <=> ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0))) | (~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) [ennf transformation 64]
140. ! [X0] : (((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) <=> ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 139]
177. ? [X0] : ((? [X2] : (? [X3] : (? [X4] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X4,X2))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X2,X3)) & r1_bcialg_1(X0,X2,X4) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) | ? [X5] : (? [X6] : ((k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),X5)) | k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X6,k1_bcialg_1(X0,X6,X5)))) & m1_subset_1(X6,u1_struct_0(X0))) & m1_subset_1(X5,u1_struct_0(X0))) | ? [X7] : (? [X8] : (? [X9] : (k1_bcialg_1(X0,X8,X7) != k1_bcialg_1(X0,k1_bcialg_1(X0,X9,X7),k1_bcialg_1(X0,X9,X8)) & r1_bcialg_1(X0,X8,X9) & m1_subset_1(X9,u1_struct_0(X0))) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) | ? [X10] : (? [X11] : (k1_bcialg_1(X0,X11,X10) != X11 & k1_bcialg_1(X0,X10,X11) = X10 & m1_subset_1(X11,u1_struct_0(X0))) & m1_subset_1(X10,u1_struct_0(X0))) | ? [X12] : (? [X13] : (((k4_struct_0(X0) != k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12)) | k1_bcialg_1(X0,X12,X13) != X12) & (k4_struct_0(X0) = k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12)) | k1_bcialg_1(X0,X12,X13) = X12)) & m1_subset_1(X13,u1_struct_0(X0))) & m1_subset_1(X12,u1_struct_0(X0)))) & ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0)) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [nnf transformation 89]
178. ? [X0] : ((? [X2] : (? [X3] : (? [X4] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X4,X2))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X3),k1_bcialg_1(X0,X2,X3)) & r1_bcialg_1(X0,X2,X4) & m1_subset_1(X4,u1_struct_0(X0))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) | ? [X5] : (? [X6] : ((k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),k1_bcialg_1(X0,k1_bcialg_1(X0,X5,X6),X5)) | k1_bcialg_1(X0,X5,X6) != k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X6,k1_bcialg_1(X0,X6,X5)))) & m1_subset_1(X6,u1_struct_0(X0))) & m1_subset_1(X5,u1_struct_0(X0))) | ? [X7] : (? [X8] : (? [X9] : (k1_bcialg_1(X0,X8,X7) != k1_bcialg_1(X0,k1_bcialg_1(X0,X9,X7),k1_bcialg_1(X0,X9,X8)) & r1_bcialg_1(X0,X8,X9) & m1_subset_1(X9,u1_struct_0(X0))) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) | ? [X10] : (? [X11] : (k1_bcialg_1(X0,X11,X10) != X11 & k1_bcialg_1(X0,X10,X11) = X10 & m1_subset_1(X11,u1_struct_0(X0))) & m1_subset_1(X10,u1_struct_0(X0))) | ? [X12] : (? [X13] : ((k4_struct_0(X0) != k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12)) | k1_bcialg_1(X0,X12,X13) != X12) & (k4_struct_0(X0) = k1_bcialg_1(X0,X13,k1_bcialg_1(X0,X13,X12)) | k1_bcialg_1(X0,X12,X13) = X12) & m1_subset_1(X13,u1_struct_0(X0))) & m1_subset_1(X12,u1_struct_0(X0)))) & ! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0)) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [flattening 177]
179. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k4_struct_0(X0) != k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) != X11) & (k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & ! [X13] : ((m2_bcialg_1(X13,X0) & v3_bciideal(X13,X0)) | ~m2_bcialg_1(X13,X0)) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) [rectify 178]
180. ? [X0] : ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X3,X1))) != k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X2),k1_bcialg_1(X0,X1,X2)) & r1_bcialg_1(X0,X1,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),k1_bcialg_1(X0,k1_bcialg_1(X0,X4,X5),X4)) | k1_bcialg_1(X0,X4,X5) != k1_bcialg_1(X0,X4,k1_bcialg_1(X0,X5,k1_bcialg_1(X0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(X0))) & m1_subset_1(X4,u1_struct_0(X0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(X0,X7,X6) != k1_bcialg_1(X0,k1_bcialg_1(X0,X8,X6),k1_bcialg_1(X0,X8,X7)) & r1_bcialg_1(X0,X7,X8) & m1_subset_1(X8,u1_struct_0(X0))) & m1_subset_1(X7,u1_struct_0(X0))) & m1_subset_1(X6,u1_struct_0(X0))) | ? [X9] : (? [X10] : (k1_bcialg_1(X0,X10,X9) != X10 & k1_bcialg_1(X0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(X0))) & m1_subset_1(X9,u1_struct_0(X0))) | ? [X11] : (? [X12] : ((k4_struct_0(X0) != k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) != X11) & (k4_struct_0(X0) = k1_bcialg_1(X0,X12,k1_bcialg_1(X0,X12,X11)) | k1_bcialg_1(X0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(X0))) & m1_subset_1(X11,u1_struct_0(X0)))) & ! [X13] : ((m2_bcialg_1(X13,X0) & v3_bciideal(X13,X0)) | ~m2_bcialg_1(X13,X0)) & l2_bcialg_1(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) => ((? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,X1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X1,X2)) & r1_bcialg_1(sK0,X1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) | ? [X4] : (? [X5] : ((k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),X4)) | k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,X4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(X4,u1_struct_0(sK0))) | ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,X6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,X6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(X6,u1_struct_0(sK0))) | ? [X9] : (? [X10] : (k1_bcialg_1(sK0,X10,X9) != X10 & k1_bcialg_1(sK0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(X9,u1_struct_0(sK0))) | ? [X11] : (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) != X11) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(X11,u1_struct_0(sK0)))) & ! [X13] : ((m2_bcialg_1(X13,sK0) & v3_bciideal(X13,sK0)) | ~m2_bcialg_1(X13,sK0)) & l2_bcialg_1(sK0) & v8_bcialg_1(sK0) & v7_bcialg_1(sK0) & v5_bcialg_1(sK0) & v4_bcialg_1(sK0) & v3_bcialg_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
181. ? [X1] : (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,X1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X1,X2)) & r1_bcialg_1(sK0,X1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,sK1,X2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
182. ? [X2] : (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,X2),k1_bcialg_1(sK0,sK1,X2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
183. ? [X3] : (k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,X3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => (k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
184. ? [X4] : (? [X5] : ((k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,X4,X5),X4)) | k1_bcialg_1(sK0,X4,X5) != k1_bcialg_1(sK0,X4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,X4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(X4,u1_struct_0(sK0))) => (? [X5] : ((k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),sK4)) | k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,sK4)))) & m1_subset_1(X5,u1_struct_0(sK0))) & m1_subset_1(sK4,u1_struct_0(sK0))) [choice axiom]
185. ? [X5] : ((k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,X5),sK4)) | k1_bcialg_1(sK0,sK4,X5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,X5,k1_bcialg_1(sK0,X5,sK4)))) & m1_subset_1(X5,u1_struct_0(sK0))) => ((k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4)))) & m1_subset_1(sK5,u1_struct_0(sK0))) [choice axiom]
186. ? [X6] : (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,X6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,X6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(X6,u1_struct_0(sK0))) => (? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) & m1_subset_1(sK6,u1_struct_0(sK0))) [choice axiom]
187. ? [X7] : (? [X8] : (k1_bcialg_1(sK0,X7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,X7)) & r1_bcialg_1(sK0,X7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(X7,u1_struct_0(sK0))) => (? [X8] : (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,sK7)) & r1_bcialg_1(sK0,sK7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) & m1_subset_1(sK7,u1_struct_0(sK0))) [choice axiom]
188. ? [X8] : (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,X8,sK6),k1_bcialg_1(sK0,X8,sK7)) & r1_bcialg_1(sK0,sK7,X8) & m1_subset_1(X8,u1_struct_0(sK0))) => (k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) & r1_bcialg_1(sK0,sK7,sK8) & m1_subset_1(sK8,u1_struct_0(sK0))) [choice axiom]
189. ? [X9] : (? [X10] : (k1_bcialg_1(sK0,X10,X9) != X10 & k1_bcialg_1(sK0,X9,X10) = X9 & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(X9,u1_struct_0(sK0))) => (? [X10] : (k1_bcialg_1(sK0,X10,sK9) != X10 & sK9 = k1_bcialg_1(sK0,sK9,X10) & m1_subset_1(X10,u1_struct_0(sK0))) & m1_subset_1(sK9,u1_struct_0(sK0))) [choice axiom]
190. ? [X10] : (k1_bcialg_1(sK0,X10,sK9) != X10 & sK9 = k1_bcialg_1(sK0,sK9,X10) & m1_subset_1(X10,u1_struct_0(sK0))) => (sK10 != k1_bcialg_1(sK0,sK10,sK9) & sK9 = k1_bcialg_1(sK0,sK9,sK10) & m1_subset_1(sK10,u1_struct_0(sK0))) [choice axiom]
191. ? [X11] : (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) != X11) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,X11)) | k1_bcialg_1(sK0,X11,X12) = X11) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(X11,u1_struct_0(sK0))) => (? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,X12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,X12)) & m1_subset_1(X12,u1_struct_0(sK0))) & m1_subset_1(sK11,u1_struct_0(sK0))) [choice axiom]
192. ? [X12] : ((k4_struct_0(sK0) != k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,X12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,X12,k1_bcialg_1(sK0,X12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,X12)) & m1_subset_1(X12,u1_struct_0(sK0))) => ((k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12)) & m1_subset_1(sK12,u1_struct_0(sK0))) [choice axiom]
193. ((((k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) & r1_bcialg_1(sK0,sK1,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) | (((k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4)))) & m1_subset_1(sK5,u1_struct_0(sK0))) & m1_subset_1(sK4,u1_struct_0(sK0))) | (((k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) & r1_bcialg_1(sK0,sK7,sK8) & m1_subset_1(sK8,u1_struct_0(sK0))) & m1_subset_1(sK7,u1_struct_0(sK0))) & m1_subset_1(sK6,u1_struct_0(sK0))) | ((sK10 != k1_bcialg_1(sK0,sK10,sK9) & sK9 = k1_bcialg_1(sK0,sK9,sK10) & m1_subset_1(sK10,u1_struct_0(sK0))) & m1_subset_1(sK9,u1_struct_0(sK0))) | (((k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12)) & (k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12)) & m1_subset_1(sK12,u1_struct_0(sK0))) & m1_subset_1(sK11,u1_struct_0(sK0)))) & ! [X13] : ((m2_bcialg_1(X13,sK0) & v3_bciideal(X13,sK0)) | ~m2_bcialg_1(X13,sK0)) & l2_bcialg_1(sK0) & v8_bcialg_1(sK0) & v7_bcialg_1(sK0) & v5_bcialg_1(sK0) & v4_bcialg_1(sK0) & v3_bcialg_1(sK0) & ~v2_struct_0(sK0) [skolemisation 179,192,191,190,189,188,187,186,185,184,183,182,181,180]
200. ! [X0] : (! [X1] : (! [X2] : (((k1_bcialg_1(X0,X1,X2) = X1 | k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) != k4_struct_0(X0)) & (k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0) | k1_bcialg_1(X0,X1,X2) != X1)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [nnf transformation 134]
201. ! [X0] : ((((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | ? [X1] : ((~m2_bcialg_1(X1,X0) | ~v3_bciideal(X1,X0)) & m2_bcialg_1(X1,X0))) & (! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0)) | (~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [nnf transformation 140]
202. ! [X0] : ((((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | ? [X1] : ((~m2_bcialg_1(X1,X0) | ~v3_bciideal(X1,X0)) & m2_bcialg_1(X1,X0))) & (! [X1] : ((m2_bcialg_1(X1,X0) & v3_bciideal(X1,X0)) | ~m2_bcialg_1(X1,X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [flattening 201]
203. ! [X0] : ((((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | ? [X1] : ((~m2_bcialg_1(X1,X0) | ~v3_bciideal(X1,X0)) & m2_bcialg_1(X1,X0))) & (! [X2] : ((m2_bcialg_1(X2,X0) & v3_bciideal(X2,X0)) | ~m2_bcialg_1(X2,X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [rectify 202]
204. ! [X0] : (? [X1] : ((~m2_bcialg_1(X1,X0) | ~v3_bciideal(X1,X0)) & m2_bcialg_1(X1,X0)) => ((~m2_bcialg_1(sK16(X0),X0) | ~v3_bciideal(sK16(X0),X0)) & m2_bcialg_1(sK16(X0),X0))) [choice axiom]
205. ! [X0] : ((((l2_bcialg_1(X0) & v1_bcialg_3(X0) & v8_bcialg_1(X0) & v7_bcialg_1(X0) & v5_bcialg_1(X0) & v4_bcialg_1(X0) & v3_bcialg_1(X0) & ~v2_struct_0(X0)) | ((~m2_bcialg_1(sK16(X0),X0) | ~v3_bciideal(sK16(X0),X0)) & m2_bcialg_1(sK16(X0),X0))) & (! [X2] : ((m2_bcialg_1(X2,X0) & v3_bciideal(X2,X0)) | ~m2_bcialg_1(X2,X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0))) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0)) [skolemisation 203,204]
230. ~v2_struct_0(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
231. v3_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
232. v4_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
233. v5_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
234. v7_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
235. v8_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
236. l2_bcialg_1(sK0) [cnf transformation 193] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
237. ~m2_bcialg_1(X13,sK0) | v3_bciideal(X13,sK0) [cnf transformation 193] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
239. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
240. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
241. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
242. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
243. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
244. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
245. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
246. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
247. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
248. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
249. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
250. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
251. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
252. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
253. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
254. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
255. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
256. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
257. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
258. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
259. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
260. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
261. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
262. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
263. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
264. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
265. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
266. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
267. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
268. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
269. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
270. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
271. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
272. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
273. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
274. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
275. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
276. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
277. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
278. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
279. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
280. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
281. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
282. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
283. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
284. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
285. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
286. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
287. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
288. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
289. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
290. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
291. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
292. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
293. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
294. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
295. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
296. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
297. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
298. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
299. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
300. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
301. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
302. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
303. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
304. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
305. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
306. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
307. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
308. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
309. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
310. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
311. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
312. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
313. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
314. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
315. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
316. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
317. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
318. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
319. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
320. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
321. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
322. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
323. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
324. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
325. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
326. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
327. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
328. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
329. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
330. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
331. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
332. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
333. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
334. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
335. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
336. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
337. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
338. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
339. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
340. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
341. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
342. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
343. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
344. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
345. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
346. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
347. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
348. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
349. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
350. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
351. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
352. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
353. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
354. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
355. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
356. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
357. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
358. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
359. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
360. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
361. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
362. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
363. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
364. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
365. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
366. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
367. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
368. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
369. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
370. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
371. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
372. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
373. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
374. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
375. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
376. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
377. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
378. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
379. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
380. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
381. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
382. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
383. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
384. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
385. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
386. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
387. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
388. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
389. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
390. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
391. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
392. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
393. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
394. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
395. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
396. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
397. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
398. m1_subset_1(sK1,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
399. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
400. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
401. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
402. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
403. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
404. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
405. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
406. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
407. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
408. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
409. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
410. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
411. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
412. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
413. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
414. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
415. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
416. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
417. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
418. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
419. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
420. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
421. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
422. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
423. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
424. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
425. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
426. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
427. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
428. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
429. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
430. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
431. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
432. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
433. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
434. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
435. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
436. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
437. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
438. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
439. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
440. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
441. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
442. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
443. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
444. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
445. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
446. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
447. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
448. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
449. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
450. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
451. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
452. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
453. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
454. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
455. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
456. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
457. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
458. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
459. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
460. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
461. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
462. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
463. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
464. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
465. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
466. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
467. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
468. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
469. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
470. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
471. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
472. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
473. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
474. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
475. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
476. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
477. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
478. m1_subset_1(sK1,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
479. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
480. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
481. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
482. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
483. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
484. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
485. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
486. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
487. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
488. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
489. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
490. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
491. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
492. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
493. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
494. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
495. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
496. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
497. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
498. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
499. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
500. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
501. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
502. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
503. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
504. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
505. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
506. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
507. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
508. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
509. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
510. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
511. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
512. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
513. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
514. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
515. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
516. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
517. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
518. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
519. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
520. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
521. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
522. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
523. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
524. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
525. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
526. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
527. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
528. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
529. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
530. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
531. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
532. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
533. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
534. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
535. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
536. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
537. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
538. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
539. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
540. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
541. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
542. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
543. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
544. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
545. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
546. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
547. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
548. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
549. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
550. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
551. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
552. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
553. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
554. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
555. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
556. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
557. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
558. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
559. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
560. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
561. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
562. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
563. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
564. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
565. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
566. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
567. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
568. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
569. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
570. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
571. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
572. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
573. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
574. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
575. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
576. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
577. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
578. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
579. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
580. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
581. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
582. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
583. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
584. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
585. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
586. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
587. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
588. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
589. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
590. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
591. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
592. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
593. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
594. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
595. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
596. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
597. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
598. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
599. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
600. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
601. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
602. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
603. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
604. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
605. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
606. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
607. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
608. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
609. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
610. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
611. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
612. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
613. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
614. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
615. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
616. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
617. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
618. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
619. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
620. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
621. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
622. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
623. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
624. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
625. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
626. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
627. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
628. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
629. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
630. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
631. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
632. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
633. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
634. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
635. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
636. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
637. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
638. m1_subset_1(sK2,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
639. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
640. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
641. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
642. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
643. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
644. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
645. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
646. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
647. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
648. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
649. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
650. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
651. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
652. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
653. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
654. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
655. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
656. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
657. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
658. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
659. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
660. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
661. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
662. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
663. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
664. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
665. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
666. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
667. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
668. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
669. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
670. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
671. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
672. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
673. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
674. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
675. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
676. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
677. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
678. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
679. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
680. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
681. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
682. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
683. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
684. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
685. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
686. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
687. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
688. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
689. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
690. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
691. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
692. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
693. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
694. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
695. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
696. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
697. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
698. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
699. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
700. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
701. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
702. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
703. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
704. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
705. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
706. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
707. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
708. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
709. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
710. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
711. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
712. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
713. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
714. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
715. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
716. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
717. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
718. m1_subset_1(sK2,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
719. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
720. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
721. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
722. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
723. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
724. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
725. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
726. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
727. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
728. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
729. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
730. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
731. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
732. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
733. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
734. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
735. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
736. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
737. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
738. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
739. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
740. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
741. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
742. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
743. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
744. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
745. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
746. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
747. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
748. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
749. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
750. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
751. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
752. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
753. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
754. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
755. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
756. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
757. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
758. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
759. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
760. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
761. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
762. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
763. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
764. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
765. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
766. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
767. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
768. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
769. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
770. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
771. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
772. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
773. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
774. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
775. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
776. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
777. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
778. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
779. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
780. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
781. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
782. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
783. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
784. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
785. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
786. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
787. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
788. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
789. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
790. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
791. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
792. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
793. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
794. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
795. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
796. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
797. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
798. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
799. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
800. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
801. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
802. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
803. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
804. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
805. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
806. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
807. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
808. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
809. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
810. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
811. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
812. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
813. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
814. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
815. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
816. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
817. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
818. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
819. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
820. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
821. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
822. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
823. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
824. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
825. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
826. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
827. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
828. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
829. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
830. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
831. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
832. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
833. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
834. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
835. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
836. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
837. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
838. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
839. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
840. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
841. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
842. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
843. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
844. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
845. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
846. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
847. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
848. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
849. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
850. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
851. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
852. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
853. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
854. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
855. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
856. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
857. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
858. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
859. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
860. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
861. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
862. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
863. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
864. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
865. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
866. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
867. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
868. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
869. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
870. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
871. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
872. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
873. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
874. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
875. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
876. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
877. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
878. m1_subset_1(sK3,u1_struct_0(sK0)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
879. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
880. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
881. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
882. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
883. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
884. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
885. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
886. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
887. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
888. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
889. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
890. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
891. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
892. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
893. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
894. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
895. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
896. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
897. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
898. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
899. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
900. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
901. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
902. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
903. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
904. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
905. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
906. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
907. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
908. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
909. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
910. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
911. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
912. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
913. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
914. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
915. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
916. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
917. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
918. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
919. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
920. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
921. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
922. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
923. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
924. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
925. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
926. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
927. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
928. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
929. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
930. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
931. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
932. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
933. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
934. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
935. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
936. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
937. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
938. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
939. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
940. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
941. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
942. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
943. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
944. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
945. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
946. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
947. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
948. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
949. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
950. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
951. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
952. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
953. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
954. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
955. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
956. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
957. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
958. m1_subset_1(sK3,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
959. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
960. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
961. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
962. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
963. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
964. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
965. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
966. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
967. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
968. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
969. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
970. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
971. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
972. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
973. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
974. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
975. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
976. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
977. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
978. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
979. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
980. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
981. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
982. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
983. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
984. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
985. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
986. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
987. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
988. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
989. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
990. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
991. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
992. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
993. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
994. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
995. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
996. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
997. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
998. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
999. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1000. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1001. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1002. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1003. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1004. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1005. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1006. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1007. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1008. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1009. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1010. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1011. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1012. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1013. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1014. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1015. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1016. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1017. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1018. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1019. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1020. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1021. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1022. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1023. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1024. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1025. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1026. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1027. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1028. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1029. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1030. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1031. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1032. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1033. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1034. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1035. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1036. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1037. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1038. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1039. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1040. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1041. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1042. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1043. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1044. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1045. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1046. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1047. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1048. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1049. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1050. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1051. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1052. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1053. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1054. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1055. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1056. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1057. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1058. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1059. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1060. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1061. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1062. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1063. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1064. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1065. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1066. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1067. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1068. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1069. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1070. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1071. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1072. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1073. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1074. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1075. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1076. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1077. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1078. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1079. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1080. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1081. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1082. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1083. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1084. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1085. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1086. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1087. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1088. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1089. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1090. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1091. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1092. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1093. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1094. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:32,wCS:3200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1095. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1096. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1097. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1098. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1099. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1100. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:22,wCS:2200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1101. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1102. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:34,wCS:3400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1103. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1104. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1105. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1106. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1107. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1108. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1109. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1110. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1111. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1112. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1113. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1114. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1115. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1116. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:33,wCS:3300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1117. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1118. r1_bcialg_1(sK0,sK1,sK3) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1119. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1120. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1121. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1122. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1123. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1124. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1125. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1126. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1127. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1128. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1129. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1130. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1131. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1132. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1133. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1134. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1135. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1136. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1137. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1138. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1139. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1140. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1141. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1142. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1143. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1144. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1145. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1146. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1147. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1148. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1149. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1150. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1151. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1152. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1153. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1154. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1155. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1156. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1157. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1158. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1159. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1160. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1161. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1162. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1163. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1164. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1165. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1166. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1167. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1168. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1169. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1170. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1171. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1172. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:49,wCS:4900,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1173. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1174. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:61,wCS:6100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1175. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1176. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1177. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1178. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1179. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1180. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:51,wCS:5100,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1181. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1182. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:63,wCS:6300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1183. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1184. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1185. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1186. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1187. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1188. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:60,wCS:6000,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1189. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1190. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1191. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1192. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1193. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1194. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1195. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1196. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:62,wCS:6200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1197. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1198. r1_bcialg_1(sK0,sK1,sK3) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1199. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1200. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1201. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1202. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1203. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1204. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1205. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1206. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1207. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1208. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1209. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1210. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1211. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1212. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1213. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1214. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1215. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1216. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1217. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1218. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1219. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1220. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1221. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1222. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1223. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1224. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1225. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1226. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1227. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1228. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1229. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1230. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1231. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1232. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1233. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1234. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1235. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1236. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1237. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1238. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1239. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1240. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1241. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1242. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1243. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1244. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1245. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1246. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1247. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1248. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1249. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1250. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1251. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1252. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1253. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1254. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1255. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1256. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1257. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1258. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1259. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1260. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1261. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1262. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1263. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1264. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1265. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1266. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1267. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1268. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1269. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1270. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1271. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1272. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1273. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1274. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1275. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1276. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1277. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1278. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK4,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1279. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1280. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1281. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1282. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1283. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1284. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1285. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1286. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1287. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1288. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1289. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1290. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1291. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1292. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1293. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1294. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1295. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1296. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1297. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1298. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1299. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1300. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1301. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1302. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1303. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1304. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1305. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1306. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1307. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1308. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1309. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1310. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1311. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1312. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1313. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1314. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1315. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1316. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1317. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1318. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1319. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1320. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1321. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1322. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1323. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1324. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1325. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1326. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1327. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1328. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1329. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1330. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1331. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1332. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:43,wCS:4300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1333. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1334. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:55,wCS:5500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1335. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1336. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1337. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1338. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1339. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1340. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:45,wCS:4500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1341. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1342. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:57,wCS:5700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1343. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1344. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1345. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1346. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1347. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1348. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:54,wCS:5400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1349. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1350. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:66,wCS:6600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1351. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1352. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1353. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1354. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1355. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1356. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:56,wCS:5600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1357. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1358. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | m1_subset_1(sK5,u1_struct_0(sK0)) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:68,wCS:6800,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1359. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1360. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1361. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1362. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1363. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1364. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1365. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1366. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1367. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1368. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1369. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1370. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1371. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1372. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1373. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1374. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK6,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1375. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1376. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1377. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1378. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1379. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1380. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1381. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1382. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1383. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1384. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1385. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1386. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1387. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1388. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1389. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1390. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK7,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1391. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1392. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1393. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1394. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1395. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1396. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1397. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1398. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1399. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1400. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1401. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1402. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1403. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1404. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1405. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1406. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | m1_subset_1(sK8,u1_struct_0(sK0)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1407. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1408. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1409. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1410. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1411. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1412. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:72,wCS:7200,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1413. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1414. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:84,wCS:8400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1415. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1416. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1417. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1418. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1419. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1420. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:74,wCS:7400,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1421. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1422. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | r1_bcialg_1(sK0,sK7,sK8) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:86,wCS:8600,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1423. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1424. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1425. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1426. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK9,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1427. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1428. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:83,wCS:8300,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1429. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1430. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | m1_subset_1(sK10,u1_struct_0(sK0)) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:95,wCS:9500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1431. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1432. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1433. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1434. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK9 = k1_bcialg_1(sK0,sK9,sK10) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1435. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK11,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1436. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | m1_subset_1(sK12,u1_struct_0(sK0)) [cnf transformation 193] {a:0,w:85,wCS:8500,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1437. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 = k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1438. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) | k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) | k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) | sK10 != k1_bcialg_1(sK0,sK10,sK9) | k4_struct_0(sK0) != k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | sK11 != k1_bcialg_1(sK0,sK11,sK12) [cnf transformation 193] {a:0,w:97,wCS:9700,goal:1,sine:0,thAx:0,allAx:14,thDist:-14}
1475. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 130] {a:0,w:53,wCS:21200,sine:2,thAx:0,allAx:1,thDist:-1}
1476. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 130] {a:0,w:56,wCS:22400,sine:2,thAx:0,allAx:1,thDist:-1}
1477. k1_bcialg_1(X0,X2,X1) = X2 | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 132] {a:0,w:50,wCS:20000,sine:2,thAx:0,allAx:1,thDist:-1}
1478. k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0) | k1_bcialg_1(X0,X1,X2) != X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 200] {a:0,w:54,wCS:21600,sine:1,thAx:0,allAx:1,thDist:-1}
1479. k1_bcialg_1(X0,X1,X2) = X1 | k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) != k4_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 200] {a:0,w:54,wCS:21600,sine:1,thAx:0,allAx:1,thDist:-1}
1480. k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 136] {a:0,w:61,wCS:24400,sine:1,thAx:0,allAx:1,thDist:-1}
1496. m2_bcialg_1(sK16(X0),X0) | v1_bcialg_3(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 205] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1497. v1_bcialg_3(X0) | ~m2_bcialg_1(sK16(X0),X0) | ~v3_bciideal(sK16(X0),X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [cnf transformation 205] {a:0,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
1547. k1_bcialg_1(X0,X2,X1) = k1_bcialg_1(X0,k1_bcialg_1(X0,X3,X1),k1_bcialg_1(X0,X3,X2)) | ~r1_bcialg_1(X0,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1480] {a:0,w:47,wCS:18800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1548. k1_bcialg_1(X0,X1,X2) != X1 | k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) = k4_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1478] {a:0,w:40,wCS:16000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1549. k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1)) != k4_struct_0(X0) | k1_bcialg_1(X0,X1,X2) = X1 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1479] {a:0,w:40,wCS:16000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
1550. k1_bcialg_1(X0,X1,X2) != X1 | k1_bcialg_1(X0,X2,X1) = X2 | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1477] {a:0,w:36,wCS:14400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1551. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,X1,k1_bcialg_1(X0,X2,k1_bcialg_1(X0,X2,X1))) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1475] {a:0,w:39,wCS:15600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1552. k1_bcialg_1(X0,X1,X2) = k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),k1_bcialg_1(X0,k1_bcialg_1(X0,X1,X2),X1)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l2_bcialg_1(X0) | ~v1_bcialg_3(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [duplicate literal removal 1476] {a:0,w:42,wCS:16800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
1554. 1 <=> sK11 = k1_bcialg_1(sK0,sK11,sK12) [avatar definition]
1555. sK11 = k1_bcialg_1(sK0,sK11,sK12) <- (1) [avatar component clause 1554] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1558. 2 <=> k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) [avatar definition]
1559. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) <- (2) [avatar component clause 1558] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1562. 3 <=> sK10 = k1_bcialg_1(sK0,sK10,sK9) [avatar definition]
1564. sK10 != k1_bcialg_1(sK0,sK10,sK9) <- (~3) [avatar component clause 1562] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1566. 4 <=> k1_bcialg_1(sK0,sK7,sK6) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) [avatar definition]
1568. k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK8,sK6),k1_bcialg_1(sK0,sK8,sK7)) <- (~4) [avatar component clause 1566] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1570. 5 <=> k1_bcialg_1(sK0,sK4,sK5) = k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) [avatar definition]
1572. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,k1_bcialg_1(sK0,sK5,k1_bcialg_1(sK0,sK5,sK4))) <- (~5) [avatar component clause 1570] {a:0,w:15,wCS:1500,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1574. 6 <=> k1_bcialg_1(sK0,sK4,sK5) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) [avatar definition]
1576. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK4,sK5),sK4)) <- (~6) [avatar component clause 1574] {a:0,w:18,wCS:1800,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1578. 7 <=> k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) = k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) [avatar definition]
1580. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK3,sK1))) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) <- (~7) [avatar component clause 1578] {a:0,w:27,wCS:2700,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1581. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1438,1578,1574,1570,1566,1562,1558,1554]
1582. 1 | 2 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1437,1578,1574,1570,1566,1562,1558,1554]
1584. 8 <=> m1_subset_1(sK12,u1_struct_0(sK0)) [avatar definition]
1586. m1_subset_1(sK12,u1_struct_0(sK0)) <- (8) [avatar component clause 1584] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1587. 8 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1436,1578,1574,1570,1566,1562,1584]
1589. 9 <=> m1_subset_1(sK11,u1_struct_0(sK0)) [avatar definition]
1591. m1_subset_1(sK11,u1_struct_0(sK0)) <- (9) [avatar component clause 1589] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1592. 9 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1435,1578,1574,1570,1566,1562,1589]
1594. 10 <=> sK9 = k1_bcialg_1(sK0,sK9,sK10) [avatar definition]
1596. sK9 = k1_bcialg_1(sK0,sK9,sK10) <- (10) [avatar component clause 1594] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1597. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1434,1578,1574,1570,1566,1594,1558,1554]
1598. 1 | 2 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1433,1578,1574,1570,1566,1594,1558,1554]
1599. 8 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1432,1578,1574,1570,1566,1594,1584]
1600. 9 | 10 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1431,1578,1574,1570,1566,1594,1589]
1602. 11 <=> m1_subset_1(sK10,u1_struct_0(sK0)) [avatar definition]
1605. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1430,1578,1574,1570,1566,1602,1558,1554]
1606. 1 | 2 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1429,1578,1574,1570,1566,1602,1558,1554]
1607. 8 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1428,1578,1574,1570,1566,1602,1584]
1608. 9 | 11 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1427,1578,1574,1570,1566,1602,1589]
1610. 12 <=> m1_subset_1(sK9,u1_struct_0(sK0)) [avatar definition]
1613. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1426,1578,1574,1570,1566,1610,1558,1554]
1614. 1 | 2 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1425,1578,1574,1570,1566,1610,1558,1554]
1615. 8 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1424,1578,1574,1570,1566,1610,1584]
1616. 9 | 12 | ~4 | ~5 | ~6 | ~7 [avatar split clause 1423,1578,1574,1570,1566,1610,1589]
1618. 13 <=> r1_bcialg_1(sK0,sK7,sK8) [avatar definition]
1620. r1_bcialg_1(sK0,sK7,sK8) <- (13) [avatar component clause 1618] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1621. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1422,1578,1574,1570,1618,1562,1558,1554]
1622. 1 | 2 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1421,1578,1574,1570,1618,1562,1558,1554]
1623. 8 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1420,1578,1574,1570,1618,1562,1584]
1624. 9 | ~3 | 13 | ~5 | ~6 | ~7 [avatar split clause 1419,1578,1574,1570,1618,1562,1589]
1625. ~1 | ~2 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1418,1578,1574,1570,1618,1594,1558,1554]
1626. 1 | 2 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1417,1578,1574,1570,1618,1594,1558,1554]
1627. 8 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1416,1578,1574,1570,1618,1594,1584]
1628. 9 | 10 | 13 | ~5 | ~6 | ~7 [avatar split clause 1415,1578,1574,1570,1618,1594,1589]
1629. ~1 | ~2 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1414,1578,1574,1570,1618,1602,1558,1554]
1630. 1 | 2 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1413,1578,1574,1570,1618,1602,1558,1554]
1631. 8 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1412,1578,1574,1570,1618,1602,1584]
1632. 9 | 11 | 13 | ~5 | ~6 | ~7 [avatar split clause 1411,1578,1574,1570,1618,1602,1589]
1633. ~1 | ~2 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1410,1578,1574,1570,1618,1610,1558,1554]
1634. 1 | 2 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1409,1578,1574,1570,1618,1610,1558,1554]
1635. 8 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1408,1578,1574,1570,1618,1610,1584]
1636. 9 | 12 | 13 | ~5 | ~6 | ~7 [avatar split clause 1407,1578,1574,1570,1618,1610,1589]
1638. 14 <=> m1_subset_1(sK8,u1_struct_0(sK0)) [avatar definition]
1640. m1_subset_1(sK8,u1_struct_0(sK0)) <- (14) [avatar component clause 1638] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1641. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1406,1578,1574,1570,1638,1562,1558,1554]
1642. 1 | 2 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1405,1578,1574,1570,1638,1562,1558,1554]
1643. 8 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1404,1578,1574,1570,1638,1562,1584]
1644. 9 | ~3 | 14 | ~5 | ~6 | ~7 [avatar split clause 1403,1578,1574,1570,1638,1562,1589]
1645. ~1 | ~2 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1402,1578,1574,1570,1638,1594,1558,1554]
1646. 1 | 2 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1401,1578,1574,1570,1638,1594,1558,1554]
1647. 8 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1400,1578,1574,1570,1638,1594,1584]
1648. 9 | 10 | 14 | ~5 | ~6 | ~7 [avatar split clause 1399,1578,1574,1570,1638,1594,1589]
1649. ~1 | ~2 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1398,1578,1574,1570,1638,1602,1558,1554]
1650. 1 | 2 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1397,1578,1574,1570,1638,1602,1558,1554]
1651. 8 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1396,1578,1574,1570,1638,1602,1584]
1652. 9 | 11 | 14 | ~5 | ~6 | ~7 [avatar split clause 1395,1578,1574,1570,1638,1602,1589]
1653. ~1 | ~2 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1394,1578,1574,1570,1638,1610,1558,1554]
1654. 1 | 2 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1393,1578,1574,1570,1638,1610,1558,1554]
1655. 8 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1392,1578,1574,1570,1638,1610,1584]
1656. 9 | 12 | 14 | ~5 | ~6 | ~7 [avatar split clause 1391,1578,1574,1570,1638,1610,1589]
1658. 15 <=> m1_subset_1(sK7,u1_struct_0(sK0)) [avatar definition]
1660. m1_subset_1(sK7,u1_struct_0(sK0)) <- (15) [avatar component clause 1658] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1661. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1390,1578,1574,1570,1658,1562,1558,1554]
1662. 1 | 2 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1389,1578,1574,1570,1658,1562,1558,1554]
1663. 8 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1388,1578,1574,1570,1658,1562,1584]
1664. 9 | ~3 | 15 | ~5 | ~6 | ~7 [avatar split clause 1387,1578,1574,1570,1658,1562,1589]
1665. ~1 | ~2 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1386,1578,1574,1570,1658,1594,1558,1554]
1666. 1 | 2 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1385,1578,1574,1570,1658,1594,1558,1554]
1667. 8 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1384,1578,1574,1570,1658,1594,1584]
1668. 9 | 10 | 15 | ~5 | ~6 | ~7 [avatar split clause 1383,1578,1574,1570,1658,1594,1589]
1669. ~1 | ~2 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1382,1578,1574,1570,1658,1602,1558,1554]
1670. 1 | 2 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1381,1578,1574,1570,1658,1602,1558,1554]
1671. 8 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1380,1578,1574,1570,1658,1602,1584]
1672. 9 | 11 | 15 | ~5 | ~6 | ~7 [avatar split clause 1379,1578,1574,1570,1658,1602,1589]
1673. ~1 | ~2 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1378,1578,1574,1570,1658,1610,1558,1554]
1674. 1 | 2 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1377,1578,1574,1570,1658,1610,1558,1554]
1675. 8 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1376,1578,1574,1570,1658,1610,1584]
1676. 9 | 12 | 15 | ~5 | ~6 | ~7 [avatar split clause 1375,1578,1574,1570,1658,1610,1589]
1678. 16 <=> m1_subset_1(sK6,u1_struct_0(sK0)) [avatar definition]
1680. m1_subset_1(sK6,u1_struct_0(sK0)) <- (16) [avatar component clause 1678] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1681. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1374,1578,1574,1570,1678,1562,1558,1554]
1682. 1 | 2 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1373,1578,1574,1570,1678,1562,1558,1554]
1683. 8 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1372,1578,1574,1570,1678,1562,1584]
1684. 9 | ~3 | 16 | ~5 | ~6 | ~7 [avatar split clause 1371,1578,1574,1570,1678,1562,1589]
1685. ~1 | ~2 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1370,1578,1574,1570,1678,1594,1558,1554]
1686. 1 | 2 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1369,1578,1574,1570,1678,1594,1558,1554]
1687. 8 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1368,1578,1574,1570,1678,1594,1584]
1688. 9 | 10 | 16 | ~5 | ~6 | ~7 [avatar split clause 1367,1578,1574,1570,1678,1594,1589]
1689. ~1 | ~2 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1366,1578,1574,1570,1678,1602,1558,1554]
1690. 1 | 2 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1365,1578,1574,1570,1678,1602,1558,1554]
1691. 8 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1364,1578,1574,1570,1678,1602,1584]
1692. 9 | 11 | 16 | ~5 | ~6 | ~7 [avatar split clause 1363,1578,1574,1570,1678,1602,1589]
1693. ~1 | ~2 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1362,1578,1574,1570,1678,1610,1558,1554]
1694. 1 | 2 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1361,1578,1574,1570,1678,1610,1558,1554]
1695. 8 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1360,1578,1574,1570,1678,1610,1584]
1696. 9 | 12 | 16 | ~5 | ~6 | ~7 [avatar split clause 1359,1578,1574,1570,1678,1610,1589]
1698. 17 <=> m1_subset_1(sK5,u1_struct_0(sK0)) [avatar definition]
1700. m1_subset_1(sK5,u1_struct_0(sK0)) <- (17) [avatar component clause 1698] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1701. ~1 | ~2 | ~3 | ~4 | 17 | ~7 [avatar split clause 1358,1578,1698,1566,1562,1558,1554]
1702. 1 | 2 | ~3 | ~4 | 17 | ~7 [avatar split clause 1357,1578,1698,1566,1562,1558,1554]
1703. 8 | ~3 | ~4 | 17 | ~7 [avatar split clause 1356,1578,1698,1566,1562,1584]
1704. 9 | ~3 | ~4 | 17 | ~7 [avatar split clause 1355,1578,1698,1566,1562,1589]
1705. ~1 | ~2 | 10 | ~4 | 17 | ~7 [avatar split clause 1354,1578,1698,1566,1594,1558,1554]
1706. 1 | 2 | 10 | ~4 | 17 | ~7 [avatar split clause 1353,1578,1698,1566,1594,1558,1554]
1707. 8 | 10 | ~4 | 17 | ~7 [avatar split clause 1352,1578,1698,1566,1594,1584]
1708. 9 | 10 | ~4 | 17 | ~7 [avatar split clause 1351,1578,1698,1566,1594,1589]
1709. ~1 | ~2 | 11 | ~4 | 17 | ~7 [avatar split clause 1350,1578,1698,1566,1602,1558,1554]
1710. 1 | 2 | 11 | ~4 | 17 | ~7 [avatar split clause 1349,1578,1698,1566,1602,1558,1554]
1711. 8 | 11 | ~4 | 17 | ~7 [avatar split clause 1348,1578,1698,1566,1602,1584]
1712. 9 | 11 | ~4 | 17 | ~7 [avatar split clause 1347,1578,1698,1566,1602,1589]
1713. ~1 | ~2 | 12 | ~4 | 17 | ~7 [avatar split clause 1346,1578,1698,1566,1610,1558,1554]
1714. 1 | 2 | 12 | ~4 | 17 | ~7 [avatar split clause 1345,1578,1698,1566,1610,1558,1554]
1715. 8 | 12 | ~4 | 17 | ~7 [avatar split clause 1344,1578,1698,1566,1610,1584]
1716. 9 | 12 | ~4 | 17 | ~7 [avatar split clause 1343,1578,1698,1566,1610,1589]
1717. ~1 | ~2 | ~3 | 13 | 17 | ~7 [avatar split clause 1342,1578,1698,1618,1562,1558,1554]
1718. 1 | 2 | ~3 | 13 | 17 | ~7 [avatar split clause 1341,1578,1698,1618,1562,1558,1554]
1719. 8 | ~3 | 13 | 17 | ~7 [avatar split clause 1340,1578,1698,1618,1562,1584]
1720. 9 | ~3 | 13 | 17 | ~7 [avatar split clause 1339,1578,1698,1618,1562,1589]
1721. ~1 | ~2 | 10 | 13 | 17 | ~7 [avatar split clause 1338,1578,1698,1618,1594,1558,1554]
1722. 1 | 2 | 10 | 13 | 17 | ~7 [avatar split clause 1337,1578,1698,1618,1594,1558,1554]
1723. 8 | 10 | 13 | 17 | ~7 [avatar split clause 1336,1578,1698,1618,1594,1584]
1724. 9 | 10 | 13 | 17 | ~7 [avatar split clause 1335,1578,1698,1618,1594,1589]
1725. ~1 | ~2 | 11 | 13 | 17 | ~7 [avatar split clause 1334,1578,1698,1618,1602,1558,1554]
1726. 1 | 2 | 11 | 13 | 17 | ~7 [avatar split clause 1333,1578,1698,1618,1602,1558,1554]
1727. 8 | 11 | 13 | 17 | ~7 [avatar split clause 1332,1578,1698,1618,1602,1584]
1728. 9 | 11 | 13 | 17 | ~7 [avatar split clause 1331,1578,1698,1618,1602,1589]
1729. ~1 | ~2 | 12 | 13 | 17 | ~7 [avatar split clause 1330,1578,1698,1618,1610,1558,1554]
1730. 1 | 2 | 12 | 13 | 17 | ~7 [avatar split clause 1329,1578,1698,1618,1610,1558,1554]
1731. 8 | 12 | 13 | 17 | ~7 [avatar split clause 1328,1578,1698,1618,1610,1584]
1732. 9 | 12 | 13 | 17 | ~7 [avatar split clause 1327,1578,1698,1618,1610,1589]
1733. ~1 | ~2 | ~3 | 14 | 17 | ~7 [avatar split clause 1326,1578,1698,1638,1562,1558,1554]
1734. 1 | 2 | ~3 | 14 | 17 | ~7 [avatar split clause 1325,1578,1698,1638,1562,1558,1554]
1735. 8 | ~3 | 14 | 17 | ~7 [avatar split clause 1324,1578,1698,1638,1562,1584]
1736. 9 | ~3 | 14 | 17 | ~7 [avatar split clause 1323,1578,1698,1638,1562,1589]
1737. ~1 | ~2 | 10 | 14 | 17 | ~7 [avatar split clause 1322,1578,1698,1638,1594,1558,1554]
1738. 1 | 2 | 10 | 14 | 17 | ~7 [avatar split clause 1321,1578,1698,1638,1594,1558,1554]
1739. 8 | 10 | 14 | 17 | ~7 [avatar split clause 1320,1578,1698,1638,1594,1584]
1740. 9 | 10 | 14 | 17 | ~7 [avatar split clause 1319,1578,1698,1638,1594,1589]
1741. ~1 | ~2 | 11 | 14 | 17 | ~7 [avatar split clause 1318,1578,1698,1638,1602,1558,1554]
1742. 1 | 2 | 11 | 14 | 17 | ~7 [avatar split clause 1317,1578,1698,1638,1602,1558,1554]
1743. 8 | 11 | 14 | 17 | ~7 [avatar split clause 1316,1578,1698,1638,1602,1584]
1744. 9 | 11 | 14 | 17 | ~7 [avatar split clause 1315,1578,1698,1638,1602,1589]
1745. ~1 | ~2 | 12 | 14 | 17 | ~7 [avatar split clause 1314,1578,1698,1638,1610,1558,1554]
1746. 1 | 2 | 12 | 14 | 17 | ~7 [avatar split clause 1313,1578,1698,1638,1610,1558,1554]
1747. 8 | 12 | 14 | 17 | ~7 [avatar split clause 1312,1578,1698,1638,1610,1584]
1748. 9 | 12 | 14 | 17 | ~7 [avatar split clause 1311,1578,1698,1638,1610,1589]
1749. ~1 | ~2 | ~3 | 15 | 17 | ~7 [avatar split clause 1310,1578,1698,1658,1562,1558,1554]
1750. 1 | 2 | ~3 | 15 | 17 | ~7 [avatar split clause 1309,1578,1698,1658,1562,1558,1554]
1751. 8 | ~3 | 15 | 17 | ~7 [avatar split clause 1308,1578,1698,1658,1562,1584]
1752. 9 | ~3 | 15 | 17 | ~7 [avatar split clause 1307,1578,1698,1658,1562,1589]
1753. ~1 | ~2 | 10 | 15 | 17 | ~7 [avatar split clause 1306,1578,1698,1658,1594,1558,1554]
1754. 1 | 2 | 10 | 15 | 17 | ~7 [avatar split clause 1305,1578,1698,1658,1594,1558,1554]
1755. 8 | 10 | 15 | 17 | ~7 [avatar split clause 1304,1578,1698,1658,1594,1584]
1756. 9 | 10 | 15 | 17 | ~7 [avatar split clause 1303,1578,1698,1658,1594,1589]
1757. ~1 | ~2 | 11 | 15 | 17 | ~7 [avatar split clause 1302,1578,1698,1658,1602,1558,1554]
1758. 1 | 2 | 11 | 15 | 17 | ~7 [avatar split clause 1301,1578,1698,1658,1602,1558,1554]
1759. 8 | 11 | 15 | 17 | ~7 [avatar split clause 1300,1578,1698,1658,1602,1584]
1760. 9 | 11 | 15 | 17 | ~7 [avatar split clause 1299,1578,1698,1658,1602,1589]
1761. ~1 | ~2 | 12 | 15 | 17 | ~7 [avatar split clause 1298,1578,1698,1658,1610,1558,1554]
1762. 1 | 2 | 12 | 15 | 17 | ~7 [avatar split clause 1297,1578,1698,1658,1610,1558,1554]
1763. 8 | 12 | 15 | 17 | ~7 [avatar split clause 1296,1578,1698,1658,1610,1584]
1764. 9 | 12 | 15 | 17 | ~7 [avatar split clause 1295,1578,1698,1658,1610,1589]
1765. ~1 | ~2 | ~3 | 16 | 17 | ~7 [avatar split clause 1294,1578,1698,1678,1562,1558,1554]
1766. 1 | 2 | ~3 | 16 | 17 | ~7 [avatar split clause 1293,1578,1698,1678,1562,1558,1554]
1767. 8 | ~3 | 16 | 17 | ~7 [avatar split clause 1292,1578,1698,1678,1562,1584]
1768. 9 | ~3 | 16 | 17 | ~7 [avatar split clause 1291,1578,1698,1678,1562,1589]
1769. ~1 | ~2 | 10 | 16 | 17 | ~7 [avatar split clause 1290,1578,1698,1678,1594,1558,1554]
1770. 1 | 2 | 10 | 16 | 17 | ~7 [avatar split clause 1289,1578,1698,1678,1594,1558,1554]
1771. 8 | 10 | 16 | 17 | ~7 [avatar split clause 1288,1578,1698,1678,1594,1584]
1772. 9 | 10 | 16 | 17 | ~7 [avatar split clause 1287,1578,1698,1678,1594,1589]
1773. ~1 | ~2 | 11 | 16 | 17 | ~7 [avatar split clause 1286,1578,1698,1678,1602,1558,1554]
1774. 1 | 2 | 11 | 16 | 17 | ~7 [avatar split clause 1285,1578,1698,1678,1602,1558,1554]
1775. 8 | 11 | 16 | 17 | ~7 [avatar split clause 1284,1578,1698,1678,1602,1584]
1776. 9 | 11 | 16 | 17 | ~7 [avatar split clause 1283,1578,1698,1678,1602,1589]
1777. ~1 | ~2 | 12 | 16 | 17 | ~7 [avatar split clause 1282,1578,1698,1678,1610,1558,1554]
1778. 1 | 2 | 12 | 16 | 17 | ~7 [avatar split clause 1281,1578,1698,1678,1610,1558,1554]
1779. 8 | 12 | 16 | 17 | ~7 [avatar split clause 1280,1578,1698,1678,1610,1584]
1780. 9 | 12 | 16 | 17 | ~7 [avatar split clause 1279,1578,1698,1678,1610,1589]
1782. 18 <=> m1_subset_1(sK4,u1_struct_0(sK0)) [avatar definition]
1784. m1_subset_1(sK4,u1_struct_0(sK0)) <- (18) [avatar component clause 1782] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1785. ~1 | ~2 | ~3 | ~4 | 18 | ~7 [avatar split clause 1278,1578,1782,1566,1562,1558,1554]
1786. 1 | 2 | ~3 | ~4 | 18 | ~7 [avatar split clause 1277,1578,1782,1566,1562,1558,1554]
1787. 8 | ~3 | ~4 | 18 | ~7 [avatar split clause 1276,1578,1782,1566,1562,1584]
1788. 9 | ~3 | ~4 | 18 | ~7 [avatar split clause 1275,1578,1782,1566,1562,1589]
1789. ~1 | ~2 | 10 | ~4 | 18 | ~7 [avatar split clause 1274,1578,1782,1566,1594,1558,1554]
1790. 1 | 2 | 10 | ~4 | 18 | ~7 [avatar split clause 1273,1578,1782,1566,1594,1558,1554]
1791. 8 | 10 | ~4 | 18 | ~7 [avatar split clause 1272,1578,1782,1566,1594,1584]
1792. 9 | 10 | ~4 | 18 | ~7 [avatar split clause 1271,1578,1782,1566,1594,1589]
1793. ~1 | ~2 | 11 | ~4 | 18 | ~7 [avatar split clause 1270,1578,1782,1566,1602,1558,1554]
1794. 1 | 2 | 11 | ~4 | 18 | ~7 [avatar split clause 1269,1578,1782,1566,1602,1558,1554]
1795. 8 | 11 | ~4 | 18 | ~7 [avatar split clause 1268,1578,1782,1566,1602,1584]
1796. 9 | 11 | ~4 | 18 | ~7 [avatar split clause 1267,1578,1782,1566,1602,1589]
1797. ~1 | ~2 | 12 | ~4 | 18 | ~7 [avatar split clause 1266,1578,1782,1566,1610,1558,1554]
1798. 1 | 2 | 12 | ~4 | 18 | ~7 [avatar split clause 1265,1578,1782,1566,1610,1558,1554]
1799. 8 | 12 | ~4 | 18 | ~7 [avatar split clause 1264,1578,1782,1566,1610,1584]
1800. 9 | 12 | ~4 | 18 | ~7 [avatar split clause 1263,1578,1782,1566,1610,1589]
1801. ~1 | ~2 | ~3 | 13 | 18 | ~7 [avatar split clause 1262,1578,1782,1618,1562,1558,1554]
1802. 1 | 2 | ~3 | 13 | 18 | ~7 [avatar split clause 1261,1578,1782,1618,1562,1558,1554]
1803. 8 | ~3 | 13 | 18 | ~7 [avatar split clause 1260,1578,1782,1618,1562,1584]
1804. 9 | ~3 | 13 | 18 | ~7 [avatar split clause 1259,1578,1782,1618,1562,1589]
1805. ~1 | ~2 | 10 | 13 | 18 | ~7 [avatar split clause 1258,1578,1782,1618,1594,1558,1554]
1806. 1 | 2 | 10 | 13 | 18 | ~7 [avatar split clause 1257,1578,1782,1618,1594,1558,1554]
1807. 8 | 10 | 13 | 18 | ~7 [avatar split clause 1256,1578,1782,1618,1594,1584]
1808. 9 | 10 | 13 | 18 | ~7 [avatar split clause 1255,1578,1782,1618,1594,1589]
1809. ~1 | ~2 | 11 | 13 | 18 | ~7 [avatar split clause 1254,1578,1782,1618,1602,1558,1554]
1810. 1 | 2 | 11 | 13 | 18 | ~7 [avatar split clause 1253,1578,1782,1618,1602,1558,1554]
1811. 8 | 11 | 13 | 18 | ~7 [avatar split clause 1252,1578,1782,1618,1602,1584]
1812. 9 | 11 | 13 | 18 | ~7 [avatar split clause 1251,1578,1782,1618,1602,1589]
1813. ~1 | ~2 | 12 | 13 | 18 | ~7 [avatar split clause 1250,1578,1782,1618,1610,1558,1554]
1814. 1 | 2 | 12 | 13 | 18 | ~7 [avatar split clause 1249,1578,1782,1618,1610,1558,1554]
1815. 8 | 12 | 13 | 18 | ~7 [avatar split clause 1248,1578,1782,1618,1610,1584]
1816. 9 | 12 | 13 | 18 | ~7 [avatar split clause 1247,1578,1782,1618,1610,1589]
1817. ~1 | ~2 | ~3 | 14 | 18 | ~7 [avatar split clause 1246,1578,1782,1638,1562,1558,1554]
1818. 1 | 2 | ~3 | 14 | 18 | ~7 [avatar split clause 1245,1578,1782,1638,1562,1558,1554]
1819. 8 | ~3 | 14 | 18 | ~7 [avatar split clause 1244,1578,1782,1638,1562,1584]
1820. 9 | ~3 | 14 | 18 | ~7 [avatar split clause 1243,1578,1782,1638,1562,1589]
1821. ~1 | ~2 | 10 | 14 | 18 | ~7 [avatar split clause 1242,1578,1782,1638,1594,1558,1554]
1822. 1 | 2 | 10 | 14 | 18 | ~7 [avatar split clause 1241,1578,1782,1638,1594,1558,1554]
1823. 8 | 10 | 14 | 18 | ~7 [avatar split clause 1240,1578,1782,1638,1594,1584]
1824. 9 | 10 | 14 | 18 | ~7 [avatar split clause 1239,1578,1782,1638,1594,1589]
1825. ~1 | ~2 | 11 | 14 | 18 | ~7 [avatar split clause 1238,1578,1782,1638,1602,1558,1554]
1826. 1 | 2 | 11 | 14 | 18 | ~7 [avatar split clause 1237,1578,1782,1638,1602,1558,1554]
1827. 8 | 11 | 14 | 18 | ~7 [avatar split clause 1236,1578,1782,1638,1602,1584]
1828. 9 | 11 | 14 | 18 | ~7 [avatar split clause 1235,1578,1782,1638,1602,1589]
1829. ~1 | ~2 | 12 | 14 | 18 | ~7 [avatar split clause 1234,1578,1782,1638,1610,1558,1554]
1830. 1 | 2 | 12 | 14 | 18 | ~7 [avatar split clause 1233,1578,1782,1638,1610,1558,1554]
1831. 8 | 12 | 14 | 18 | ~7 [avatar split clause 1232,1578,1782,1638,1610,1584]
1832. 9 | 12 | 14 | 18 | ~7 [avatar split clause 1231,1578,1782,1638,1610,1589]
1833. ~1 | ~2 | ~3 | 15 | 18 | ~7 [avatar split clause 1230,1578,1782,1658,1562,1558,1554]
1834. 1 | 2 | ~3 | 15 | 18 | ~7 [avatar split clause 1229,1578,1782,1658,1562,1558,1554]
1835. 8 | ~3 | 15 | 18 | ~7 [avatar split clause 1228,1578,1782,1658,1562,1584]
1836. 9 | ~3 | 15 | 18 | ~7 [avatar split clause 1227,1578,1782,1658,1562,1589]
1837. ~1 | ~2 | 10 | 15 | 18 | ~7 [avatar split clause 1226,1578,1782,1658,1594,1558,1554]
1838. 1 | 2 | 10 | 15 | 18 | ~7 [avatar split clause 1225,1578,1782,1658,1594,1558,1554]
1839. 8 | 10 | 15 | 18 | ~7 [avatar split clause 1224,1578,1782,1658,1594,1584]
1840. 9 | 10 | 15 | 18 | ~7 [avatar split clause 1223,1578,1782,1658,1594,1589]
1841. ~1 | ~2 | 11 | 15 | 18 | ~7 [avatar split clause 1222,1578,1782,1658,1602,1558,1554]
1842. 1 | 2 | 11 | 15 | 18 | ~7 [avatar split clause 1221,1578,1782,1658,1602,1558,1554]
1843. 8 | 11 | 15 | 18 | ~7 [avatar split clause 1220,1578,1782,1658,1602,1584]
1844. 9 | 11 | 15 | 18 | ~7 [avatar split clause 1219,1578,1782,1658,1602,1589]
1845. ~1 | ~2 | 12 | 15 | 18 | ~7 [avatar split clause 1218,1578,1782,1658,1610,1558,1554]
1846. 1 | 2 | 12 | 15 | 18 | ~7 [avatar split clause 1217,1578,1782,1658,1610,1558,1554]
1847. 8 | 12 | 15 | 18 | ~7 [avatar split clause 1216,1578,1782,1658,1610,1584]
1848. 9 | 12 | 15 | 18 | ~7 [avatar split clause 1215,1578,1782,1658,1610,1589]
1849. ~1 | ~2 | ~3 | 16 | 18 | ~7 [avatar split clause 1214,1578,1782,1678,1562,1558,1554]
1850. 1 | 2 | ~3 | 16 | 18 | ~7 [avatar split clause 1213,1578,1782,1678,1562,1558,1554]
1851. 8 | ~3 | 16 | 18 | ~7 [avatar split clause 1212,1578,1782,1678,1562,1584]
1852. 9 | ~3 | 16 | 18 | ~7 [avatar split clause 1211,1578,1782,1678,1562,1589]
1853. ~1 | ~2 | 10 | 16 | 18 | ~7 [avatar split clause 1210,1578,1782,1678,1594,1558,1554]
1854. 1 | 2 | 10 | 16 | 18 | ~7 [avatar split clause 1209,1578,1782,1678,1594,1558,1554]
1855. 8 | 10 | 16 | 18 | ~7 [avatar split clause 1208,1578,1782,1678,1594,1584]
1856. 9 | 10 | 16 | 18 | ~7 [avatar split clause 1207,1578,1782,1678,1594,1589]
1857. ~1 | ~2 | 11 | 16 | 18 | ~7 [avatar split clause 1206,1578,1782,1678,1602,1558,1554]
1858. 1 | 2 | 11 | 16 | 18 | ~7 [avatar split clause 1205,1578,1782,1678,1602,1558,1554]
1859. 8 | 11 | 16 | 18 | ~7 [avatar split clause 1204,1578,1782,1678,1602,1584]
1860. 9 | 11 | 16 | 18 | ~7 [avatar split clause 1203,1578,1782,1678,1602,1589]
1861. ~1 | ~2 | 12 | 16 | 18 | ~7 [avatar split clause 1202,1578,1782,1678,1610,1558,1554]
1862. 1 | 2 | 12 | 16 | 18 | ~7 [avatar split clause 1201,1578,1782,1678,1610,1558,1554]
1863. 8 | 12 | 16 | 18 | ~7 [avatar split clause 1200,1578,1782,1678,1610,1584]
1864. 9 | 12 | 16 | 18 | ~7 [avatar split clause 1199,1578,1782,1678,1610,1589]
1866. 19 <=> r1_bcialg_1(sK0,sK1,sK3) [avatar definition]
1868. r1_bcialg_1(sK0,sK1,sK3) <- (19) [avatar component clause 1866] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
1869. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1198,1866,1574,1570,1566,1562,1558,1554]
1870. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1197,1866,1574,1570,1566,1562,1558,1554]
1871. 8 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1196,1866,1574,1570,1566,1562,1584]
1872. 9 | ~3 | ~4 | ~5 | ~6 | 19 [avatar split clause 1195,1866,1574,1570,1566,1562,1589]
1873. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1194,1866,1574,1570,1566,1594,1558,1554]
1874. 1 | 2 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1193,1866,1574,1570,1566,1594,1558,1554]
1875. 8 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1192,1866,1574,1570,1566,1594,1584]
1876. 9 | 10 | ~4 | ~5 | ~6 | 19 [avatar split clause 1191,1866,1574,1570,1566,1594,1589]
1877. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1190,1866,1574,1570,1566,1602,1558,1554]
1878. 1 | 2 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1189,1866,1574,1570,1566,1602,1558,1554]
1879. 8 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1188,1866,1574,1570,1566,1602,1584]
1880. 9 | 11 | ~4 | ~5 | ~6 | 19 [avatar split clause 1187,1866,1574,1570,1566,1602,1589]
1881. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1186,1866,1574,1570,1566,1610,1558,1554]
1882. 1 | 2 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1185,1866,1574,1570,1566,1610,1558,1554]
1883. 8 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1184,1866,1574,1570,1566,1610,1584]
1884. 9 | 12 | ~4 | ~5 | ~6 | 19 [avatar split clause 1183,1866,1574,1570,1566,1610,1589]
1885. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1182,1866,1574,1570,1618,1562,1558,1554]
1886. 1 | 2 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1181,1866,1574,1570,1618,1562,1558,1554]
1887. 8 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1180,1866,1574,1570,1618,1562,1584]
1888. 9 | ~3 | 13 | ~5 | ~6 | 19 [avatar split clause 1179,1866,1574,1570,1618,1562,1589]
1889. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1178,1866,1574,1570,1618,1594,1558,1554]
1890. 1 | 2 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1177,1866,1574,1570,1618,1594,1558,1554]
1891. 8 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1176,1866,1574,1570,1618,1594,1584]
1892. 9 | 10 | 13 | ~5 | ~6 | 19 [avatar split clause 1175,1866,1574,1570,1618,1594,1589]
1893. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1174,1866,1574,1570,1618,1602,1558,1554]
1894. 1 | 2 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1173,1866,1574,1570,1618,1602,1558,1554]
1895. 8 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1172,1866,1574,1570,1618,1602,1584]
1896. 9 | 11 | 13 | ~5 | ~6 | 19 [avatar split clause 1171,1866,1574,1570,1618,1602,1589]
1897. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1170,1866,1574,1570,1618,1610,1558,1554]
1898. 1 | 2 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1169,1866,1574,1570,1618,1610,1558,1554]
1899. 8 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1168,1866,1574,1570,1618,1610,1584]
1900. 9 | 12 | 13 | ~5 | ~6 | 19 [avatar split clause 1167,1866,1574,1570,1618,1610,1589]
1901. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1166,1866,1574,1570,1638,1562,1558,1554]
1902. 1 | 2 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1165,1866,1574,1570,1638,1562,1558,1554]
1903. 8 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1164,1866,1574,1570,1638,1562,1584]
1904. 9 | ~3 | 14 | ~5 | ~6 | 19 [avatar split clause 1163,1866,1574,1570,1638,1562,1589]
1905. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1162,1866,1574,1570,1638,1594,1558,1554]
1906. 1 | 2 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1161,1866,1574,1570,1638,1594,1558,1554]
1907. 8 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1160,1866,1574,1570,1638,1594,1584]
1908. 9 | 10 | 14 | ~5 | ~6 | 19 [avatar split clause 1159,1866,1574,1570,1638,1594,1589]
1909. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1158,1866,1574,1570,1638,1602,1558,1554]
1910. 1 | 2 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1157,1866,1574,1570,1638,1602,1558,1554]
1911. 8 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1156,1866,1574,1570,1638,1602,1584]
1912. 9 | 11 | 14 | ~5 | ~6 | 19 [avatar split clause 1155,1866,1574,1570,1638,1602,1589]
1913. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1154,1866,1574,1570,1638,1610,1558,1554]
1914. 1 | 2 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1153,1866,1574,1570,1638,1610,1558,1554]
1915. 8 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1152,1866,1574,1570,1638,1610,1584]
1916. 9 | 12 | 14 | ~5 | ~6 | 19 [avatar split clause 1151,1866,1574,1570,1638,1610,1589]
1917. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1150,1866,1574,1570,1658,1562,1558,1554]
1918. 1 | 2 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1149,1866,1574,1570,1658,1562,1558,1554]
1919. 8 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1148,1866,1574,1570,1658,1562,1584]
1920. 9 | ~3 | 15 | ~5 | ~6 | 19 [avatar split clause 1147,1866,1574,1570,1658,1562,1589]
1921. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1146,1866,1574,1570,1658,1594,1558,1554]
1922. 1 | 2 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1145,1866,1574,1570,1658,1594,1558,1554]
1923. 8 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1144,1866,1574,1570,1658,1594,1584]
1924. 9 | 10 | 15 | ~5 | ~6 | 19 [avatar split clause 1143,1866,1574,1570,1658,1594,1589]
1925. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1142,1866,1574,1570,1658,1602,1558,1554]
1926. 1 | 2 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1141,1866,1574,1570,1658,1602,1558,1554]
1927. 8 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1140,1866,1574,1570,1658,1602,1584]
1928. 9 | 11 | 15 | ~5 | ~6 | 19 [avatar split clause 1139,1866,1574,1570,1658,1602,1589]
1929. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1138,1866,1574,1570,1658,1610,1558,1554]
1930. 1 | 2 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1137,1866,1574,1570,1658,1610,1558,1554]
1931. 8 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1136,1866,1574,1570,1658,1610,1584]
1932. 9 | 12 | 15 | ~5 | ~6 | 19 [avatar split clause 1135,1866,1574,1570,1658,1610,1589]
1933. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1134,1866,1574,1570,1678,1562,1558,1554]
1934. 1 | 2 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1133,1866,1574,1570,1678,1562,1558,1554]
1935. 8 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1132,1866,1574,1570,1678,1562,1584]
1936. 9 | ~3 | 16 | ~5 | ~6 | 19 [avatar split clause 1131,1866,1574,1570,1678,1562,1589]
1937. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1130,1866,1574,1570,1678,1594,1558,1554]
1938. 1 | 2 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1129,1866,1574,1570,1678,1594,1558,1554]
1939. 8 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1128,1866,1574,1570,1678,1594,1584]
1940. 9 | 10 | 16 | ~5 | ~6 | 19 [avatar split clause 1127,1866,1574,1570,1678,1594,1589]
1941. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1126,1866,1574,1570,1678,1602,1558,1554]
1942. 1 | 2 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1125,1866,1574,1570,1678,1602,1558,1554]
1943. 8 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1124,1866,1574,1570,1678,1602,1584]
1944. 9 | 11 | 16 | ~5 | ~6 | 19 [avatar split clause 1123,1866,1574,1570,1678,1602,1589]
1945. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1122,1866,1574,1570,1678,1610,1558,1554]
1946. 1 | 2 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1121,1866,1574,1570,1678,1610,1558,1554]
1947. 8 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1120,1866,1574,1570,1678,1610,1584]
1948. 9 | 12 | 16 | ~5 | ~6 | 19 [avatar split clause 1119,1866,1574,1570,1678,1610,1589]
1949. ~1 | ~2 | ~3 | ~4 | 17 | 19 [avatar split clause 1118,1866,1698,1566,1562,1558,1554]
1950. 1 | 2 | ~3 | ~4 | 17 | 19 [avatar split clause 1117,1866,1698,1566,1562,1558,1554]
1951. 8 | ~3 | ~4 | 17 | 19 [avatar split clause 1116,1866,1698,1566,1562,1584]
1952. 9 | ~3 | ~4 | 17 | 19 [avatar split clause 1115,1866,1698,1566,1562,1589]
1953. ~1 | ~2 | 10 | ~4 | 17 | 19 [avatar split clause 1114,1866,1698,1566,1594,1558,1554]
1954. 1 | 2 | 10 | ~4 | 17 | 19 [avatar split clause 1113,1866,1698,1566,1594,1558,1554]
1955. 8 | 10 | ~4 | 17 | 19 [avatar split clause 1112,1866,1698,1566,1594,1584]
1956. 9 | 10 | ~4 | 17 | 19 [avatar split clause 1111,1866,1698,1566,1594,1589]
1957. ~1 | ~2 | 11 | ~4 | 17 | 19 [avatar split clause 1110,1866,1698,1566,1602,1558,1554]
1958. 1 | 2 | 11 | ~4 | 17 | 19 [avatar split clause 1109,1866,1698,1566,1602,1558,1554]
1959. 8 | 11 | ~4 | 17 | 19 [avatar split clause 1108,1866,1698,1566,1602,1584]
1960. 9 | 11 | ~4 | 17 | 19 [avatar split clause 1107,1866,1698,1566,1602,1589]
1961. ~1 | ~2 | 12 | ~4 | 17 | 19 [avatar split clause 1106,1866,1698,1566,1610,1558,1554]
1962. 1 | 2 | 12 | ~4 | 17 | 19 [avatar split clause 1105,1866,1698,1566,1610,1558,1554]
1963. 8 | 12 | ~4 | 17 | 19 [avatar split clause 1104,1866,1698,1566,1610,1584]
1964. 9 | 12 | ~4 | 17 | 19 [avatar split clause 1103,1866,1698,1566,1610,1589]
1965. ~1 | ~2 | ~3 | 13 | 17 | 19 [avatar split clause 1102,1866,1698,1618,1562,1558,1554]
1966. 1 | 2 | ~3 | 13 | 17 | 19 [avatar split clause 1101,1866,1698,1618,1562,1558,1554]
1967. 8 | ~3 | 13 | 17 | 19 [avatar split clause 1100,1866,1698,1618,1562,1584]
1968. 9 | ~3 | 13 | 17 | 19 [avatar split clause 1099,1866,1698,1618,1562,1589]
1969. ~1 | ~2 | 10 | 13 | 17 | 19 [avatar split clause 1098,1866,1698,1618,1594,1558,1554]
1970. 1 | 2 | 10 | 13 | 17 | 19 [avatar split clause 1097,1866,1698,1618,1594,1558,1554]
1971. 8 | 10 | 13 | 17 | 19 [avatar split clause 1096,1866,1698,1618,1594,1584]
1972. 9 | 10 | 13 | 17 | 19 [avatar split clause 1095,1866,1698,1618,1594,1589]
1973. ~1 | ~2 | 11 | 13 | 17 | 19 [avatar split clause 1094,1866,1698,1618,1602,1558,1554]
1974. 1 | 2 | 11 | 13 | 17 | 19 [avatar split clause 1093,1866,1698,1618,1602,1558,1554]
1975. 8 | 11 | 13 | 17 | 19 [avatar split clause 1092,1866,1698,1618,1602,1584]
1976. 9 | 11 | 13 | 17 | 19 [avatar split clause 1091,1866,1698,1618,1602,1589]
1977. ~1 | ~2 | 12 | 13 | 17 | 19 [avatar split clause 1090,1866,1698,1618,1610,1558,1554]
1978. 1 | 2 | 12 | 13 | 17 | 19 [avatar split clause 1089,1866,1698,1618,1610,1558,1554]
1979. 8 | 12 | 13 | 17 | 19 [avatar split clause 1088,1866,1698,1618,1610,1584]
1980. 9 | 12 | 13 | 17 | 19 [avatar split clause 1087,1866,1698,1618,1610,1589]
1981. ~1 | ~2 | ~3 | 14 | 17 | 19 [avatar split clause 1086,1866,1698,1638,1562,1558,1554]
1982. 1 | 2 | ~3 | 14 | 17 | 19 [avatar split clause 1085,1866,1698,1638,1562,1558,1554]
1983. 8 | ~3 | 14 | 17 | 19 [avatar split clause 1084,1866,1698,1638,1562,1584]
1984. 9 | ~3 | 14 | 17 | 19 [avatar split clause 1083,1866,1698,1638,1562,1589]
1985. ~1 | ~2 | 10 | 14 | 17 | 19 [avatar split clause 1082,1866,1698,1638,1594,1558,1554]
1986. 1 | 2 | 10 | 14 | 17 | 19 [avatar split clause 1081,1866,1698,1638,1594,1558,1554]
1987. 8 | 10 | 14 | 17 | 19 [avatar split clause 1080,1866,1698,1638,1594,1584]
1988. 9 | 10 | 14 | 17 | 19 [avatar split clause 1079,1866,1698,1638,1594,1589]
1989. ~1 | ~2 | 11 | 14 | 17 | 19 [avatar split clause 1078,1866,1698,1638,1602,1558,1554]
1990. 1 | 2 | 11 | 14 | 17 | 19 [avatar split clause 1077,1866,1698,1638,1602,1558,1554]
1991. 8 | 11 | 14 | 17 | 19 [avatar split clause 1076,1866,1698,1638,1602,1584]
1992. 9 | 11 | 14 | 17 | 19 [avatar split clause 1075,1866,1698,1638,1602,1589]
1993. ~1 | ~2 | 12 | 14 | 17 | 19 [avatar split clause 1074,1866,1698,1638,1610,1558,1554]
1994. 1 | 2 | 12 | 14 | 17 | 19 [avatar split clause 1073,1866,1698,1638,1610,1558,1554]
1995. 8 | 12 | 14 | 17 | 19 [avatar split clause 1072,1866,1698,1638,1610,1584]
1996. 9 | 12 | 14 | 17 | 19 [avatar split clause 1071,1866,1698,1638,1610,1589]
1997. ~1 | ~2 | ~3 | 15 | 17 | 19 [avatar split clause 1070,1866,1698,1658,1562,1558,1554]
1998. 1 | 2 | ~3 | 15 | 17 | 19 [avatar split clause 1069,1866,1698,1658,1562,1558,1554]
1999. 8 | ~3 | 15 | 17 | 19 [avatar split clause 1068,1866,1698,1658,1562,1584]
2000. 9 | ~3 | 15 | 17 | 19 [avatar split clause 1067,1866,1698,1658,1562,1589]
2001. ~1 | ~2 | 10 | 15 | 17 | 19 [avatar split clause 1066,1866,1698,1658,1594,1558,1554]
2002. 1 | 2 | 10 | 15 | 17 | 19 [avatar split clause 1065,1866,1698,1658,1594,1558,1554]
2003. 8 | 10 | 15 | 17 | 19 [avatar split clause 1064,1866,1698,1658,1594,1584]
2004. 9 | 10 | 15 | 17 | 19 [avatar split clause 1063,1866,1698,1658,1594,1589]
2005. ~1 | ~2 | 11 | 15 | 17 | 19 [avatar split clause 1062,1866,1698,1658,1602,1558,1554]
2006. 1 | 2 | 11 | 15 | 17 | 19 [avatar split clause 1061,1866,1698,1658,1602,1558,1554]
2007. 8 | 11 | 15 | 17 | 19 [avatar split clause 1060,1866,1698,1658,1602,1584]
2008. 9 | 11 | 15 | 17 | 19 [avatar split clause 1059,1866,1698,1658,1602,1589]
2009. ~1 | ~2 | 12 | 15 | 17 | 19 [avatar split clause 1058,1866,1698,1658,1610,1558,1554]
2010. 1 | 2 | 12 | 15 | 17 | 19 [avatar split clause 1057,1866,1698,1658,1610,1558,1554]
2011. 8 | 12 | 15 | 17 | 19 [avatar split clause 1056,1866,1698,1658,1610,1584]
2012. 9 | 12 | 15 | 17 | 19 [avatar split clause 1055,1866,1698,1658,1610,1589]
2013. ~1 | ~2 | ~3 | 16 | 17 | 19 [avatar split clause 1054,1866,1698,1678,1562,1558,1554]
2014. 1 | 2 | ~3 | 16 | 17 | 19 [avatar split clause 1053,1866,1698,1678,1562,1558,1554]
2015. 8 | ~3 | 16 | 17 | 19 [avatar split clause 1052,1866,1698,1678,1562,1584]
2016. 9 | ~3 | 16 | 17 | 19 [avatar split clause 1051,1866,1698,1678,1562,1589]
2017. ~1 | ~2 | 10 | 16 | 17 | 19 [avatar split clause 1050,1866,1698,1678,1594,1558,1554]
2018. 1 | 2 | 10 | 16 | 17 | 19 [avatar split clause 1049,1866,1698,1678,1594,1558,1554]
2019. 8 | 10 | 16 | 17 | 19 [avatar split clause 1048,1866,1698,1678,1594,1584]
2020. 9 | 10 | 16 | 17 | 19 [avatar split clause 1047,1866,1698,1678,1594,1589]
2021. ~1 | ~2 | 11 | 16 | 17 | 19 [avatar split clause 1046,1866,1698,1678,1602,1558,1554]
2022. 1 | 2 | 11 | 16 | 17 | 19 [avatar split clause 1045,1866,1698,1678,1602,1558,1554]
2023. 8 | 11 | 16 | 17 | 19 [avatar split clause 1044,1866,1698,1678,1602,1584]
2024. 9 | 11 | 16 | 17 | 19 [avatar split clause 1043,1866,1698,1678,1602,1589]
2025. ~1 | ~2 | 12 | 16 | 17 | 19 [avatar split clause 1042,1866,1698,1678,1610,1558,1554]
2026. 1 | 2 | 12 | 16 | 17 | 19 [avatar split clause 1041,1866,1698,1678,1610,1558,1554]
2027. 8 | 12 | 16 | 17 | 19 [avatar split clause 1040,1866,1698,1678,1610,1584]
2028. 9 | 12 | 16 | 17 | 19 [avatar split clause 1039,1866,1698,1678,1610,1589]
2029. ~1 | ~2 | ~3 | ~4 | 18 | 19 [avatar split clause 1038,1866,1782,1566,1562,1558,1554]
2030. 1 | 2 | ~3 | ~4 | 18 | 19 [avatar split clause 1037,1866,1782,1566,1562,1558,1554]
2031. 8 | ~3 | ~4 | 18 | 19 [avatar split clause 1036,1866,1782,1566,1562,1584]
2032. 9 | ~3 | ~4 | 18 | 19 [avatar split clause 1035,1866,1782,1566,1562,1589]
2033. ~1 | ~2 | 10 | ~4 | 18 | 19 [avatar split clause 1034,1866,1782,1566,1594,1558,1554]
2034. 1 | 2 | 10 | ~4 | 18 | 19 [avatar split clause 1033,1866,1782,1566,1594,1558,1554]
2035. 8 | 10 | ~4 | 18 | 19 [avatar split clause 1032,1866,1782,1566,1594,1584]
2036. 9 | 10 | ~4 | 18 | 19 [avatar split clause 1031,1866,1782,1566,1594,1589]
2037. ~1 | ~2 | 11 | ~4 | 18 | 19 [avatar split clause 1030,1866,1782,1566,1602,1558,1554]
2038. 1 | 2 | 11 | ~4 | 18 | 19 [avatar split clause 1029,1866,1782,1566,1602,1558,1554]
2039. 8 | 11 | ~4 | 18 | 19 [avatar split clause 1028,1866,1782,1566,1602,1584]
2040. 9 | 11 | ~4 | 18 | 19 [avatar split clause 1027,1866,1782,1566,1602,1589]
2041. ~1 | ~2 | 12 | ~4 | 18 | 19 [avatar split clause 1026,1866,1782,1566,1610,1558,1554]
2042. 1 | 2 | 12 | ~4 | 18 | 19 [avatar split clause 1025,1866,1782,1566,1610,1558,1554]
2043. 8 | 12 | ~4 | 18 | 19 [avatar split clause 1024,1866,1782,1566,1610,1584]
2044. 9 | 12 | ~4 | 18 | 19 [avatar split clause 1023,1866,1782,1566,1610,1589]
2045. ~1 | ~2 | ~3 | 13 | 18 | 19 [avatar split clause 1022,1866,1782,1618,1562,1558,1554]
2046. 1 | 2 | ~3 | 13 | 18 | 19 [avatar split clause 1021,1866,1782,1618,1562,1558,1554]
2047. 8 | ~3 | 13 | 18 | 19 [avatar split clause 1020,1866,1782,1618,1562,1584]
2048. 9 | ~3 | 13 | 18 | 19 [avatar split clause 1019,1866,1782,1618,1562,1589]
2049. ~1 | ~2 | 10 | 13 | 18 | 19 [avatar split clause 1018,1866,1782,1618,1594,1558,1554]
2050. 1 | 2 | 10 | 13 | 18 | 19 [avatar split clause 1017,1866,1782,1618,1594,1558,1554]
2051. 8 | 10 | 13 | 18 | 19 [avatar split clause 1016,1866,1782,1618,1594,1584]
2052. 9 | 10 | 13 | 18 | 19 [avatar split clause 1015,1866,1782,1618,1594,1589]
2053. ~1 | ~2 | 11 | 13 | 18 | 19 [avatar split clause 1014,1866,1782,1618,1602,1558,1554]
2054. 1 | 2 | 11 | 13 | 18 | 19 [avatar split clause 1013,1866,1782,1618,1602,1558,1554]
2055. 8 | 11 | 13 | 18 | 19 [avatar split clause 1012,1866,1782,1618,1602,1584]
2056. 9 | 11 | 13 | 18 | 19 [avatar split clause 1011,1866,1782,1618,1602,1589]
2057. ~1 | ~2 | 12 | 13 | 18 | 19 [avatar split clause 1010,1866,1782,1618,1610,1558,1554]
2058. 1 | 2 | 12 | 13 | 18 | 19 [avatar split clause 1009,1866,1782,1618,1610,1558,1554]
2059. 8 | 12 | 13 | 18 | 19 [avatar split clause 1008,1866,1782,1618,1610,1584]
2060. 9 | 12 | 13 | 18 | 19 [avatar split clause 1007,1866,1782,1618,1610,1589]
2061. ~1 | ~2 | ~3 | 14 | 18 | 19 [avatar split clause 1006,1866,1782,1638,1562,1558,1554]
2062. 1 | 2 | ~3 | 14 | 18 | 19 [avatar split clause 1005,1866,1782,1638,1562,1558,1554]
2063. 8 | ~3 | 14 | 18 | 19 [avatar split clause 1004,1866,1782,1638,1562,1584]
2064. 9 | ~3 | 14 | 18 | 19 [avatar split clause 1003,1866,1782,1638,1562,1589]
2065. ~1 | ~2 | 10 | 14 | 18 | 19 [avatar split clause 1002,1866,1782,1638,1594,1558,1554]
2066. 1 | 2 | 10 | 14 | 18 | 19 [avatar split clause 1001,1866,1782,1638,1594,1558,1554]
2067. 8 | 10 | 14 | 18 | 19 [avatar split clause 1000,1866,1782,1638,1594,1584]
2068. 9 | 10 | 14 | 18 | 19 [avatar split clause 999,1866,1782,1638,1594,1589]
2069. ~1 | ~2 | 11 | 14 | 18 | 19 [avatar split clause 998,1866,1782,1638,1602,1558,1554]
2070. 1 | 2 | 11 | 14 | 18 | 19 [avatar split clause 997,1866,1782,1638,1602,1558,1554]
2071. 8 | 11 | 14 | 18 | 19 [avatar split clause 996,1866,1782,1638,1602,1584]
2072. 9 | 11 | 14 | 18 | 19 [avatar split clause 995,1866,1782,1638,1602,1589]
2073. ~1 | ~2 | 12 | 14 | 18 | 19 [avatar split clause 994,1866,1782,1638,1610,1558,1554]
2074. 1 | 2 | 12 | 14 | 18 | 19 [avatar split clause 993,1866,1782,1638,1610,1558,1554]
2075. 8 | 12 | 14 | 18 | 19 [avatar split clause 992,1866,1782,1638,1610,1584]
2076. 9 | 12 | 14 | 18 | 19 [avatar split clause 991,1866,1782,1638,1610,1589]
2077. ~1 | ~2 | ~3 | 15 | 18 | 19 [avatar split clause 990,1866,1782,1658,1562,1558,1554]
2078. 1 | 2 | ~3 | 15 | 18 | 19 [avatar split clause 989,1866,1782,1658,1562,1558,1554]
2079. 8 | ~3 | 15 | 18 | 19 [avatar split clause 988,1866,1782,1658,1562,1584]
2080. 9 | ~3 | 15 | 18 | 19 [avatar split clause 987,1866,1782,1658,1562,1589]
2081. ~1 | ~2 | 10 | 15 | 18 | 19 [avatar split clause 986,1866,1782,1658,1594,1558,1554]
2082. 1 | 2 | 10 | 15 | 18 | 19 [avatar split clause 985,1866,1782,1658,1594,1558,1554]
2083. 8 | 10 | 15 | 18 | 19 [avatar split clause 984,1866,1782,1658,1594,1584]
2084. 9 | 10 | 15 | 18 | 19 [avatar split clause 983,1866,1782,1658,1594,1589]
2085. ~1 | ~2 | 11 | 15 | 18 | 19 [avatar split clause 982,1866,1782,1658,1602,1558,1554]
2086. 1 | 2 | 11 | 15 | 18 | 19 [avatar split clause 981,1866,1782,1658,1602,1558,1554]
2087. 8 | 11 | 15 | 18 | 19 [avatar split clause 980,1866,1782,1658,1602,1584]
2088. 9 | 11 | 15 | 18 | 19 [avatar split clause 979,1866,1782,1658,1602,1589]
2089. ~1 | ~2 | 12 | 15 | 18 | 19 [avatar split clause 978,1866,1782,1658,1610,1558,1554]
2090. 1 | 2 | 12 | 15 | 18 | 19 [avatar split clause 977,1866,1782,1658,1610,1558,1554]
2091. 8 | 12 | 15 | 18 | 19 [avatar split clause 976,1866,1782,1658,1610,1584]
2092. 9 | 12 | 15 | 18 | 19 [avatar split clause 975,1866,1782,1658,1610,1589]
2093. ~1 | ~2 | ~3 | 16 | 18 | 19 [avatar split clause 974,1866,1782,1678,1562,1558,1554]
2094. 1 | 2 | ~3 | 16 | 18 | 19 [avatar split clause 973,1866,1782,1678,1562,1558,1554]
2095. 8 | ~3 | 16 | 18 | 19 [avatar split clause 972,1866,1782,1678,1562,1584]
2096. 9 | ~3 | 16 | 18 | 19 [avatar split clause 971,1866,1782,1678,1562,1589]
2097. ~1 | ~2 | 10 | 16 | 18 | 19 [avatar split clause 970,1866,1782,1678,1594,1558,1554]
2098. 1 | 2 | 10 | 16 | 18 | 19 [avatar split clause 969,1866,1782,1678,1594,1558,1554]
2099. 8 | 10 | 16 | 18 | 19 [avatar split clause 968,1866,1782,1678,1594,1584]
2100. 9 | 10 | 16 | 18 | 19 [avatar split clause 967,1866,1782,1678,1594,1589]
2101. ~1 | ~2 | 11 | 16 | 18 | 19 [avatar split clause 966,1866,1782,1678,1602,1558,1554]
2102. 1 | 2 | 11 | 16 | 18 | 19 [avatar split clause 965,1866,1782,1678,1602,1558,1554]
2103. 8 | 11 | 16 | 18 | 19 [avatar split clause 964,1866,1782,1678,1602,1584]
2104. 9 | 11 | 16 | 18 | 19 [avatar split clause 963,1866,1782,1678,1602,1589]
2105. ~1 | ~2 | 12 | 16 | 18 | 19 [avatar split clause 962,1866,1782,1678,1610,1558,1554]
2106. 1 | 2 | 12 | 16 | 18 | 19 [avatar split clause 961,1866,1782,1678,1610,1558,1554]
2107. 8 | 12 | 16 | 18 | 19 [avatar split clause 960,1866,1782,1678,1610,1584]
2108. 9 | 12 | 16 | 18 | 19 [avatar split clause 959,1866,1782,1678,1610,1589]
2110. 20 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
2112. m1_subset_1(sK3,u1_struct_0(sK0)) <- (20) [avatar component clause 2110] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2113. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 958,2110,1574,1570,1566,1562,1558,1554]
2114. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 957,2110,1574,1570,1566,1562,1558,1554]
2115. 8 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 956,2110,1574,1570,1566,1562,1584]
2116. 9 | ~3 | ~4 | ~5 | ~6 | 20 [avatar split clause 955,2110,1574,1570,1566,1562,1589]
2117. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 954,2110,1574,1570,1566,1594,1558,1554]
2118. 1 | 2 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 953,2110,1574,1570,1566,1594,1558,1554]
2119. 8 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 952,2110,1574,1570,1566,1594,1584]
2120. 9 | 10 | ~4 | ~5 | ~6 | 20 [avatar split clause 951,2110,1574,1570,1566,1594,1589]
2121. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 950,2110,1574,1570,1566,1602,1558,1554]
2122. 1 | 2 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 949,2110,1574,1570,1566,1602,1558,1554]
2123. 8 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 948,2110,1574,1570,1566,1602,1584]
2124. 9 | 11 | ~4 | ~5 | ~6 | 20 [avatar split clause 947,2110,1574,1570,1566,1602,1589]
2125. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 946,2110,1574,1570,1566,1610,1558,1554]
2126. 1 | 2 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 945,2110,1574,1570,1566,1610,1558,1554]
2127. 8 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 944,2110,1574,1570,1566,1610,1584]
2128. 9 | 12 | ~4 | ~5 | ~6 | 20 [avatar split clause 943,2110,1574,1570,1566,1610,1589]
2129. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 942,2110,1574,1570,1618,1562,1558,1554]
2130. 1 | 2 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 941,2110,1574,1570,1618,1562,1558,1554]
2131. 8 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 940,2110,1574,1570,1618,1562,1584]
2132. 9 | ~3 | 13 | ~5 | ~6 | 20 [avatar split clause 939,2110,1574,1570,1618,1562,1589]
2133. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 938,2110,1574,1570,1618,1594,1558,1554]
2134. 1 | 2 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 937,2110,1574,1570,1618,1594,1558,1554]
2135. 8 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 936,2110,1574,1570,1618,1594,1584]
2136. 9 | 10 | 13 | ~5 | ~6 | 20 [avatar split clause 935,2110,1574,1570,1618,1594,1589]
2137. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 934,2110,1574,1570,1618,1602,1558,1554]
2138. 1 | 2 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 933,2110,1574,1570,1618,1602,1558,1554]
2139. 8 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 932,2110,1574,1570,1618,1602,1584]
2140. 9 | 11 | 13 | ~5 | ~6 | 20 [avatar split clause 931,2110,1574,1570,1618,1602,1589]
2141. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 930,2110,1574,1570,1618,1610,1558,1554]
2142. 1 | 2 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 929,2110,1574,1570,1618,1610,1558,1554]
2143. 8 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 928,2110,1574,1570,1618,1610,1584]
2144. 9 | 12 | 13 | ~5 | ~6 | 20 [avatar split clause 927,2110,1574,1570,1618,1610,1589]
2145. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 926,2110,1574,1570,1638,1562,1558,1554]
2146. 1 | 2 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 925,2110,1574,1570,1638,1562,1558,1554]
2147. 8 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 924,2110,1574,1570,1638,1562,1584]
2148. 9 | ~3 | 14 | ~5 | ~6 | 20 [avatar split clause 923,2110,1574,1570,1638,1562,1589]
2149. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 922,2110,1574,1570,1638,1594,1558,1554]
2150. 1 | 2 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 921,2110,1574,1570,1638,1594,1558,1554]
2151. 8 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 920,2110,1574,1570,1638,1594,1584]
2152. 9 | 10 | 14 | ~5 | ~6 | 20 [avatar split clause 919,2110,1574,1570,1638,1594,1589]
2153. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 918,2110,1574,1570,1638,1602,1558,1554]
2154. 1 | 2 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 917,2110,1574,1570,1638,1602,1558,1554]
2155. 8 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 916,2110,1574,1570,1638,1602,1584]
2156. 9 | 11 | 14 | ~5 | ~6 | 20 [avatar split clause 915,2110,1574,1570,1638,1602,1589]
2157. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 914,2110,1574,1570,1638,1610,1558,1554]
2158. 1 | 2 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 913,2110,1574,1570,1638,1610,1558,1554]
2159. 8 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 912,2110,1574,1570,1638,1610,1584]
2160. 9 | 12 | 14 | ~5 | ~6 | 20 [avatar split clause 911,2110,1574,1570,1638,1610,1589]
2161. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 910,2110,1574,1570,1658,1562,1558,1554]
2162. 1 | 2 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 909,2110,1574,1570,1658,1562,1558,1554]
2163. 8 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 908,2110,1574,1570,1658,1562,1584]
2164. 9 | ~3 | 15 | ~5 | ~6 | 20 [avatar split clause 907,2110,1574,1570,1658,1562,1589]
2165. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 906,2110,1574,1570,1658,1594,1558,1554]
2166. 1 | 2 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 905,2110,1574,1570,1658,1594,1558,1554]
2167. 8 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 904,2110,1574,1570,1658,1594,1584]
2168. 9 | 10 | 15 | ~5 | ~6 | 20 [avatar split clause 903,2110,1574,1570,1658,1594,1589]
2169. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 902,2110,1574,1570,1658,1602,1558,1554]
2170. 1 | 2 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 901,2110,1574,1570,1658,1602,1558,1554]
2171. 8 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 900,2110,1574,1570,1658,1602,1584]
2172. 9 | 11 | 15 | ~5 | ~6 | 20 [avatar split clause 899,2110,1574,1570,1658,1602,1589]
2173. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 898,2110,1574,1570,1658,1610,1558,1554]
2174. 1 | 2 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 897,2110,1574,1570,1658,1610,1558,1554]
2175. 8 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 896,2110,1574,1570,1658,1610,1584]
2176. 9 | 12 | 15 | ~5 | ~6 | 20 [avatar split clause 895,2110,1574,1570,1658,1610,1589]
2177. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 894,2110,1574,1570,1678,1562,1558,1554]
2178. 1 | 2 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 893,2110,1574,1570,1678,1562,1558,1554]
2179. 8 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 892,2110,1574,1570,1678,1562,1584]
2180. 9 | ~3 | 16 | ~5 | ~6 | 20 [avatar split clause 891,2110,1574,1570,1678,1562,1589]
2181. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 890,2110,1574,1570,1678,1594,1558,1554]
2182. 1 | 2 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 889,2110,1574,1570,1678,1594,1558,1554]
2183. 8 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 888,2110,1574,1570,1678,1594,1584]
2184. 9 | 10 | 16 | ~5 | ~6 | 20 [avatar split clause 887,2110,1574,1570,1678,1594,1589]
2185. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 886,2110,1574,1570,1678,1602,1558,1554]
2186. 1 | 2 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 885,2110,1574,1570,1678,1602,1558,1554]
2187. 8 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 884,2110,1574,1570,1678,1602,1584]
2188. 9 | 11 | 16 | ~5 | ~6 | 20 [avatar split clause 883,2110,1574,1570,1678,1602,1589]
2189. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 882,2110,1574,1570,1678,1610,1558,1554]
2190. 1 | 2 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 881,2110,1574,1570,1678,1610,1558,1554]
2191. 8 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 880,2110,1574,1570,1678,1610,1584]
2192. 9 | 12 | 16 | ~5 | ~6 | 20 [avatar split clause 879,2110,1574,1570,1678,1610,1589]
2193. ~1 | ~2 | ~3 | ~4 | 17 | 20 [avatar split clause 878,2110,1698,1566,1562,1558,1554]
2194. 1 | 2 | ~3 | ~4 | 17 | 20 [avatar split clause 877,2110,1698,1566,1562,1558,1554]
2195. 8 | ~3 | ~4 | 17 | 20 [avatar split clause 876,2110,1698,1566,1562,1584]
2196. 9 | ~3 | ~4 | 17 | 20 [avatar split clause 875,2110,1698,1566,1562,1589]
2197. ~1 | ~2 | 10 | ~4 | 17 | 20 [avatar split clause 874,2110,1698,1566,1594,1558,1554]
2198. 1 | 2 | 10 | ~4 | 17 | 20 [avatar split clause 873,2110,1698,1566,1594,1558,1554]
2199. 8 | 10 | ~4 | 17 | 20 [avatar split clause 872,2110,1698,1566,1594,1584]
2200. 9 | 10 | ~4 | 17 | 20 [avatar split clause 871,2110,1698,1566,1594,1589]
2201. ~1 | ~2 | 11 | ~4 | 17 | 20 [avatar split clause 870,2110,1698,1566,1602,1558,1554]
2202. 1 | 2 | 11 | ~4 | 17 | 20 [avatar split clause 869,2110,1698,1566,1602,1558,1554]
2203. 8 | 11 | ~4 | 17 | 20 [avatar split clause 868,2110,1698,1566,1602,1584]
2204. 9 | 11 | ~4 | 17 | 20 [avatar split clause 867,2110,1698,1566,1602,1589]
2205. ~1 | ~2 | 12 | ~4 | 17 | 20 [avatar split clause 866,2110,1698,1566,1610,1558,1554]
2206. 1 | 2 | 12 | ~4 | 17 | 20 [avatar split clause 865,2110,1698,1566,1610,1558,1554]
2207. 8 | 12 | ~4 | 17 | 20 [avatar split clause 864,2110,1698,1566,1610,1584]
2208. 9 | 12 | ~4 | 17 | 20 [avatar split clause 863,2110,1698,1566,1610,1589]
2209. ~1 | ~2 | ~3 | 13 | 17 | 20 [avatar split clause 862,2110,1698,1618,1562,1558,1554]
2210. 1 | 2 | ~3 | 13 | 17 | 20 [avatar split clause 861,2110,1698,1618,1562,1558,1554]
2211. 8 | ~3 | 13 | 17 | 20 [avatar split clause 860,2110,1698,1618,1562,1584]
2212. 9 | ~3 | 13 | 17 | 20 [avatar split clause 859,2110,1698,1618,1562,1589]
2213. ~1 | ~2 | 10 | 13 | 17 | 20 [avatar split clause 858,2110,1698,1618,1594,1558,1554]
2214. 1 | 2 | 10 | 13 | 17 | 20 [avatar split clause 857,2110,1698,1618,1594,1558,1554]
2215. 8 | 10 | 13 | 17 | 20 [avatar split clause 856,2110,1698,1618,1594,1584]
2216. 9 | 10 | 13 | 17 | 20 [avatar split clause 855,2110,1698,1618,1594,1589]
2217. ~1 | ~2 | 11 | 13 | 17 | 20 [avatar split clause 854,2110,1698,1618,1602,1558,1554]
2218. 1 | 2 | 11 | 13 | 17 | 20 [avatar split clause 853,2110,1698,1618,1602,1558,1554]
2219. 8 | 11 | 13 | 17 | 20 [avatar split clause 852,2110,1698,1618,1602,1584]
2220. 9 | 11 | 13 | 17 | 20 [avatar split clause 851,2110,1698,1618,1602,1589]
2221. ~1 | ~2 | 12 | 13 | 17 | 20 [avatar split clause 850,2110,1698,1618,1610,1558,1554]
2222. 1 | 2 | 12 | 13 | 17 | 20 [avatar split clause 849,2110,1698,1618,1610,1558,1554]
2223. 8 | 12 | 13 | 17 | 20 [avatar split clause 848,2110,1698,1618,1610,1584]
2224. 9 | 12 | 13 | 17 | 20 [avatar split clause 847,2110,1698,1618,1610,1589]
2225. ~1 | ~2 | ~3 | 14 | 17 | 20 [avatar split clause 846,2110,1698,1638,1562,1558,1554]
2226. 1 | 2 | ~3 | 14 | 17 | 20 [avatar split clause 845,2110,1698,1638,1562,1558,1554]
2227. 8 | ~3 | 14 | 17 | 20 [avatar split clause 844,2110,1698,1638,1562,1584]
2228. 9 | ~3 | 14 | 17 | 20 [avatar split clause 843,2110,1698,1638,1562,1589]
2229. ~1 | ~2 | 10 | 14 | 17 | 20 [avatar split clause 842,2110,1698,1638,1594,1558,1554]
2230. 1 | 2 | 10 | 14 | 17 | 20 [avatar split clause 841,2110,1698,1638,1594,1558,1554]
2231. 8 | 10 | 14 | 17 | 20 [avatar split clause 840,2110,1698,1638,1594,1584]
2232. 9 | 10 | 14 | 17 | 20 [avatar split clause 839,2110,1698,1638,1594,1589]
2233. ~1 | ~2 | 11 | 14 | 17 | 20 [avatar split clause 838,2110,1698,1638,1602,1558,1554]
2234. 1 | 2 | 11 | 14 | 17 | 20 [avatar split clause 837,2110,1698,1638,1602,1558,1554]
2235. 8 | 11 | 14 | 17 | 20 [avatar split clause 836,2110,1698,1638,1602,1584]
2236. 9 | 11 | 14 | 17 | 20 [avatar split clause 835,2110,1698,1638,1602,1589]
2237. ~1 | ~2 | 12 | 14 | 17 | 20 [avatar split clause 834,2110,1698,1638,1610,1558,1554]
2238. 1 | 2 | 12 | 14 | 17 | 20 [avatar split clause 833,2110,1698,1638,1610,1558,1554]
2239. 8 | 12 | 14 | 17 | 20 [avatar split clause 832,2110,1698,1638,1610,1584]
2240. 9 | 12 | 14 | 17 | 20 [avatar split clause 831,2110,1698,1638,1610,1589]
2241. ~1 | ~2 | ~3 | 15 | 17 | 20 [avatar split clause 830,2110,1698,1658,1562,1558,1554]
2242. 1 | 2 | ~3 | 15 | 17 | 20 [avatar split clause 829,2110,1698,1658,1562,1558,1554]
2243. 8 | ~3 | 15 | 17 | 20 [avatar split clause 828,2110,1698,1658,1562,1584]
2244. 9 | ~3 | 15 | 17 | 20 [avatar split clause 827,2110,1698,1658,1562,1589]
2245. ~1 | ~2 | 10 | 15 | 17 | 20 [avatar split clause 826,2110,1698,1658,1594,1558,1554]
2246. 1 | 2 | 10 | 15 | 17 | 20 [avatar split clause 825,2110,1698,1658,1594,1558,1554]
2247. 8 | 10 | 15 | 17 | 20 [avatar split clause 824,2110,1698,1658,1594,1584]
2248. 9 | 10 | 15 | 17 | 20 [avatar split clause 823,2110,1698,1658,1594,1589]
2249. ~1 | ~2 | 11 | 15 | 17 | 20 [avatar split clause 822,2110,1698,1658,1602,1558,1554]
2250. 1 | 2 | 11 | 15 | 17 | 20 [avatar split clause 821,2110,1698,1658,1602,1558,1554]
2251. 8 | 11 | 15 | 17 | 20 [avatar split clause 820,2110,1698,1658,1602,1584]
2252. 9 | 11 | 15 | 17 | 20 [avatar split clause 819,2110,1698,1658,1602,1589]
2253. ~1 | ~2 | 12 | 15 | 17 | 20 [avatar split clause 818,2110,1698,1658,1610,1558,1554]
2254. 1 | 2 | 12 | 15 | 17 | 20 [avatar split clause 817,2110,1698,1658,1610,1558,1554]
2255. 8 | 12 | 15 | 17 | 20 [avatar split clause 816,2110,1698,1658,1610,1584]
2256. 9 | 12 | 15 | 17 | 20 [avatar split clause 815,2110,1698,1658,1610,1589]
2257. ~1 | ~2 | ~3 | 16 | 17 | 20 [avatar split clause 814,2110,1698,1678,1562,1558,1554]
2258. 1 | 2 | ~3 | 16 | 17 | 20 [avatar split clause 813,2110,1698,1678,1562,1558,1554]
2259. 8 | ~3 | 16 | 17 | 20 [avatar split clause 812,2110,1698,1678,1562,1584]
2260. 9 | ~3 | 16 | 17 | 20 [avatar split clause 811,2110,1698,1678,1562,1589]
2261. ~1 | ~2 | 10 | 16 | 17 | 20 [avatar split clause 810,2110,1698,1678,1594,1558,1554]
2262. 1 | 2 | 10 | 16 | 17 | 20 [avatar split clause 809,2110,1698,1678,1594,1558,1554]
2263. 8 | 10 | 16 | 17 | 20 [avatar split clause 808,2110,1698,1678,1594,1584]
2264. 9 | 10 | 16 | 17 | 20 [avatar split clause 807,2110,1698,1678,1594,1589]
2265. ~1 | ~2 | 11 | 16 | 17 | 20 [avatar split clause 806,2110,1698,1678,1602,1558,1554]
2266. 1 | 2 | 11 | 16 | 17 | 20 [avatar split clause 805,2110,1698,1678,1602,1558,1554]
2267. 8 | 11 | 16 | 17 | 20 [avatar split clause 804,2110,1698,1678,1602,1584]
2268. 9 | 11 | 16 | 17 | 20 [avatar split clause 803,2110,1698,1678,1602,1589]
2269. ~1 | ~2 | 12 | 16 | 17 | 20 [avatar split clause 802,2110,1698,1678,1610,1558,1554]
2270. 1 | 2 | 12 | 16 | 17 | 20 [avatar split clause 801,2110,1698,1678,1610,1558,1554]
2271. 8 | 12 | 16 | 17 | 20 [avatar split clause 800,2110,1698,1678,1610,1584]
2272. 9 | 12 | 16 | 17 | 20 [avatar split clause 799,2110,1698,1678,1610,1589]
2273. ~1 | ~2 | ~3 | ~4 | 18 | 20 [avatar split clause 798,2110,1782,1566,1562,1558,1554]
2274. 1 | 2 | ~3 | ~4 | 18 | 20 [avatar split clause 797,2110,1782,1566,1562,1558,1554]
2275. 8 | ~3 | ~4 | 18 | 20 [avatar split clause 796,2110,1782,1566,1562,1584]
2276. 9 | ~3 | ~4 | 18 | 20 [avatar split clause 795,2110,1782,1566,1562,1589]
2277. ~1 | ~2 | 10 | ~4 | 18 | 20 [avatar split clause 794,2110,1782,1566,1594,1558,1554]
2278. 1 | 2 | 10 | ~4 | 18 | 20 [avatar split clause 793,2110,1782,1566,1594,1558,1554]
2279. 8 | 10 | ~4 | 18 | 20 [avatar split clause 792,2110,1782,1566,1594,1584]
2280. 9 | 10 | ~4 | 18 | 20 [avatar split clause 791,2110,1782,1566,1594,1589]
2281. ~1 | ~2 | 11 | ~4 | 18 | 20 [avatar split clause 790,2110,1782,1566,1602,1558,1554]
2282. 1 | 2 | 11 | ~4 | 18 | 20 [avatar split clause 789,2110,1782,1566,1602,1558,1554]
2283. 8 | 11 | ~4 | 18 | 20 [avatar split clause 788,2110,1782,1566,1602,1584]
2284. 9 | 11 | ~4 | 18 | 20 [avatar split clause 787,2110,1782,1566,1602,1589]
2285. ~1 | ~2 | 12 | ~4 | 18 | 20 [avatar split clause 786,2110,1782,1566,1610,1558,1554]
2286. 1 | 2 | 12 | ~4 | 18 | 20 [avatar split clause 785,2110,1782,1566,1610,1558,1554]
2287. 8 | 12 | ~4 | 18 | 20 [avatar split clause 784,2110,1782,1566,1610,1584]
2288. 9 | 12 | ~4 | 18 | 20 [avatar split clause 783,2110,1782,1566,1610,1589]
2289. ~1 | ~2 | ~3 | 13 | 18 | 20 [avatar split clause 782,2110,1782,1618,1562,1558,1554]
2290. 1 | 2 | ~3 | 13 | 18 | 20 [avatar split clause 781,2110,1782,1618,1562,1558,1554]
2291. 8 | ~3 | 13 | 18 | 20 [avatar split clause 780,2110,1782,1618,1562,1584]
2292. 9 | ~3 | 13 | 18 | 20 [avatar split clause 779,2110,1782,1618,1562,1589]
2293. ~1 | ~2 | 10 | 13 | 18 | 20 [avatar split clause 778,2110,1782,1618,1594,1558,1554]
2294. 1 | 2 | 10 | 13 | 18 | 20 [avatar split clause 777,2110,1782,1618,1594,1558,1554]
2295. 8 | 10 | 13 | 18 | 20 [avatar split clause 776,2110,1782,1618,1594,1584]
2296. 9 | 10 | 13 | 18 | 20 [avatar split clause 775,2110,1782,1618,1594,1589]
2297. ~1 | ~2 | 11 | 13 | 18 | 20 [avatar split clause 774,2110,1782,1618,1602,1558,1554]
2298. 1 | 2 | 11 | 13 | 18 | 20 [avatar split clause 773,2110,1782,1618,1602,1558,1554]
2299. 8 | 11 | 13 | 18 | 20 [avatar split clause 772,2110,1782,1618,1602,1584]
2300. 9 | 11 | 13 | 18 | 20 [avatar split clause 771,2110,1782,1618,1602,1589]
2301. ~1 | ~2 | 12 | 13 | 18 | 20 [avatar split clause 770,2110,1782,1618,1610,1558,1554]
2302. 1 | 2 | 12 | 13 | 18 | 20 [avatar split clause 769,2110,1782,1618,1610,1558,1554]
2303. 8 | 12 | 13 | 18 | 20 [avatar split clause 768,2110,1782,1618,1610,1584]
2304. 9 | 12 | 13 | 18 | 20 [avatar split clause 767,2110,1782,1618,1610,1589]
2305. ~1 | ~2 | ~3 | 14 | 18 | 20 [avatar split clause 766,2110,1782,1638,1562,1558,1554]
2306. 1 | 2 | ~3 | 14 | 18 | 20 [avatar split clause 765,2110,1782,1638,1562,1558,1554]
2307. 8 | ~3 | 14 | 18 | 20 [avatar split clause 764,2110,1782,1638,1562,1584]
2308. 9 | ~3 | 14 | 18 | 20 [avatar split clause 763,2110,1782,1638,1562,1589]
2309. ~1 | ~2 | 10 | 14 | 18 | 20 [avatar split clause 762,2110,1782,1638,1594,1558,1554]
2310. 1 | 2 | 10 | 14 | 18 | 20 [avatar split clause 761,2110,1782,1638,1594,1558,1554]
2311. 8 | 10 | 14 | 18 | 20 [avatar split clause 760,2110,1782,1638,1594,1584]
2312. 9 | 10 | 14 | 18 | 20 [avatar split clause 759,2110,1782,1638,1594,1589]
2313. ~1 | ~2 | 11 | 14 | 18 | 20 [avatar split clause 758,2110,1782,1638,1602,1558,1554]
2314. 1 | 2 | 11 | 14 | 18 | 20 [avatar split clause 757,2110,1782,1638,1602,1558,1554]
2315. 8 | 11 | 14 | 18 | 20 [avatar split clause 756,2110,1782,1638,1602,1584]
2316. 9 | 11 | 14 | 18 | 20 [avatar split clause 755,2110,1782,1638,1602,1589]
2317. ~1 | ~2 | 12 | 14 | 18 | 20 [avatar split clause 754,2110,1782,1638,1610,1558,1554]
2318. 1 | 2 | 12 | 14 | 18 | 20 [avatar split clause 753,2110,1782,1638,1610,1558,1554]
2319. 8 | 12 | 14 | 18 | 20 [avatar split clause 752,2110,1782,1638,1610,1584]
2320. 9 | 12 | 14 | 18 | 20 [avatar split clause 751,2110,1782,1638,1610,1589]
2321. ~1 | ~2 | ~3 | 15 | 18 | 20 [avatar split clause 750,2110,1782,1658,1562,1558,1554]
2322. 1 | 2 | ~3 | 15 | 18 | 20 [avatar split clause 749,2110,1782,1658,1562,1558,1554]
2323. 8 | ~3 | 15 | 18 | 20 [avatar split clause 748,2110,1782,1658,1562,1584]
2324. 9 | ~3 | 15 | 18 | 20 [avatar split clause 747,2110,1782,1658,1562,1589]
2325. ~1 | ~2 | 10 | 15 | 18 | 20 [avatar split clause 746,2110,1782,1658,1594,1558,1554]
2326. 1 | 2 | 10 | 15 | 18 | 20 [avatar split clause 745,2110,1782,1658,1594,1558,1554]
2327. 8 | 10 | 15 | 18 | 20 [avatar split clause 744,2110,1782,1658,1594,1584]
2328. 9 | 10 | 15 | 18 | 20 [avatar split clause 743,2110,1782,1658,1594,1589]
2329. ~1 | ~2 | 11 | 15 | 18 | 20 [avatar split clause 742,2110,1782,1658,1602,1558,1554]
2330. 1 | 2 | 11 | 15 | 18 | 20 [avatar split clause 741,2110,1782,1658,1602,1558,1554]
2331. 8 | 11 | 15 | 18 | 20 [avatar split clause 740,2110,1782,1658,1602,1584]
2332. 9 | 11 | 15 | 18 | 20 [avatar split clause 739,2110,1782,1658,1602,1589]
2333. ~1 | ~2 | 12 | 15 | 18 | 20 [avatar split clause 738,2110,1782,1658,1610,1558,1554]
2334. 1 | 2 | 12 | 15 | 18 | 20 [avatar split clause 737,2110,1782,1658,1610,1558,1554]
2335. 8 | 12 | 15 | 18 | 20 [avatar split clause 736,2110,1782,1658,1610,1584]
2336. 9 | 12 | 15 | 18 | 20 [avatar split clause 735,2110,1782,1658,1610,1589]
2337. ~1 | ~2 | ~3 | 16 | 18 | 20 [avatar split clause 734,2110,1782,1678,1562,1558,1554]
2338. 1 | 2 | ~3 | 16 | 18 | 20 [avatar split clause 733,2110,1782,1678,1562,1558,1554]
2339. 8 | ~3 | 16 | 18 | 20 [avatar split clause 732,2110,1782,1678,1562,1584]
2340. 9 | ~3 | 16 | 18 | 20 [avatar split clause 731,2110,1782,1678,1562,1589]
2341. ~1 | ~2 | 10 | 16 | 18 | 20 [avatar split clause 730,2110,1782,1678,1594,1558,1554]
2342. 1 | 2 | 10 | 16 | 18 | 20 [avatar split clause 729,2110,1782,1678,1594,1558,1554]
2343. 8 | 10 | 16 | 18 | 20 [avatar split clause 728,2110,1782,1678,1594,1584]
2344. 9 | 10 | 16 | 18 | 20 [avatar split clause 727,2110,1782,1678,1594,1589]
2345. ~1 | ~2 | 11 | 16 | 18 | 20 [avatar split clause 726,2110,1782,1678,1602,1558,1554]
2346. 1 | 2 | 11 | 16 | 18 | 20 [avatar split clause 725,2110,1782,1678,1602,1558,1554]
2347. 8 | 11 | 16 | 18 | 20 [avatar split clause 724,2110,1782,1678,1602,1584]
2348. 9 | 11 | 16 | 18 | 20 [avatar split clause 723,2110,1782,1678,1602,1589]
2349. ~1 | ~2 | 12 | 16 | 18 | 20 [avatar split clause 722,2110,1782,1678,1610,1558,1554]
2350. 1 | 2 | 12 | 16 | 18 | 20 [avatar split clause 721,2110,1782,1678,1610,1558,1554]
2351. 8 | 12 | 16 | 18 | 20 [avatar split clause 720,2110,1782,1678,1610,1584]
2352. 9 | 12 | 16 | 18 | 20 [avatar split clause 719,2110,1782,1678,1610,1589]
2354. 21 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
2356. m1_subset_1(sK2,u1_struct_0(sK0)) <- (21) [avatar component clause 2354] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2357. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 718,2354,1574,1570,1566,1562,1558,1554]
2358. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 717,2354,1574,1570,1566,1562,1558,1554]
2359. 8 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 716,2354,1574,1570,1566,1562,1584]
2360. 9 | ~3 | ~4 | ~5 | ~6 | 21 [avatar split clause 715,2354,1574,1570,1566,1562,1589]
2361. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 714,2354,1574,1570,1566,1594,1558,1554]
2362. 1 | 2 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 713,2354,1574,1570,1566,1594,1558,1554]
2363. 8 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 712,2354,1574,1570,1566,1594,1584]
2364. 9 | 10 | ~4 | ~5 | ~6 | 21 [avatar split clause 711,2354,1574,1570,1566,1594,1589]
2365. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 710,2354,1574,1570,1566,1602,1558,1554]
2366. 1 | 2 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 709,2354,1574,1570,1566,1602,1558,1554]
2367. 8 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 708,2354,1574,1570,1566,1602,1584]
2368. 9 | 11 | ~4 | ~5 | ~6 | 21 [avatar split clause 707,2354,1574,1570,1566,1602,1589]
2369. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 706,2354,1574,1570,1566,1610,1558,1554]
2370. 1 | 2 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 705,2354,1574,1570,1566,1610,1558,1554]
2371. 8 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 704,2354,1574,1570,1566,1610,1584]
2372. 9 | 12 | ~4 | ~5 | ~6 | 21 [avatar split clause 703,2354,1574,1570,1566,1610,1589]
2373. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 702,2354,1574,1570,1618,1562,1558,1554]
2374. 1 | 2 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 701,2354,1574,1570,1618,1562,1558,1554]
2375. 8 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 700,2354,1574,1570,1618,1562,1584]
2376. 9 | ~3 | 13 | ~5 | ~6 | 21 [avatar split clause 699,2354,1574,1570,1618,1562,1589]
2377. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 698,2354,1574,1570,1618,1594,1558,1554]
2378. 1 | 2 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 697,2354,1574,1570,1618,1594,1558,1554]
2379. 8 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 696,2354,1574,1570,1618,1594,1584]
2380. 9 | 10 | 13 | ~5 | ~6 | 21 [avatar split clause 695,2354,1574,1570,1618,1594,1589]
2381. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 694,2354,1574,1570,1618,1602,1558,1554]
2382. 1 | 2 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 693,2354,1574,1570,1618,1602,1558,1554]
2383. 8 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 692,2354,1574,1570,1618,1602,1584]
2384. 9 | 11 | 13 | ~5 | ~6 | 21 [avatar split clause 691,2354,1574,1570,1618,1602,1589]
2385. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 690,2354,1574,1570,1618,1610,1558,1554]
2386. 1 | 2 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 689,2354,1574,1570,1618,1610,1558,1554]
2387. 8 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 688,2354,1574,1570,1618,1610,1584]
2388. 9 | 12 | 13 | ~5 | ~6 | 21 [avatar split clause 687,2354,1574,1570,1618,1610,1589]
2389. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 686,2354,1574,1570,1638,1562,1558,1554]
2390. 1 | 2 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 685,2354,1574,1570,1638,1562,1558,1554]
2391. 8 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 684,2354,1574,1570,1638,1562,1584]
2392. 9 | ~3 | 14 | ~5 | ~6 | 21 [avatar split clause 683,2354,1574,1570,1638,1562,1589]
2393. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 682,2354,1574,1570,1638,1594,1558,1554]
2394. 1 | 2 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 681,2354,1574,1570,1638,1594,1558,1554]
2395. 8 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 680,2354,1574,1570,1638,1594,1584]
2396. 9 | 10 | 14 | ~5 | ~6 | 21 [avatar split clause 679,2354,1574,1570,1638,1594,1589]
2397. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 678,2354,1574,1570,1638,1602,1558,1554]
2398. 1 | 2 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 677,2354,1574,1570,1638,1602,1558,1554]
2399. 8 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 676,2354,1574,1570,1638,1602,1584]
2400. 9 | 11 | 14 | ~5 | ~6 | 21 [avatar split clause 675,2354,1574,1570,1638,1602,1589]
2401. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 674,2354,1574,1570,1638,1610,1558,1554]
2402. 1 | 2 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 673,2354,1574,1570,1638,1610,1558,1554]
2403. 8 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 672,2354,1574,1570,1638,1610,1584]
2404. 9 | 12 | 14 | ~5 | ~6 | 21 [avatar split clause 671,2354,1574,1570,1638,1610,1589]
2405. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 670,2354,1574,1570,1658,1562,1558,1554]
2406. 1 | 2 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 669,2354,1574,1570,1658,1562,1558,1554]
2407. 8 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 668,2354,1574,1570,1658,1562,1584]
2408. 9 | ~3 | 15 | ~5 | ~6 | 21 [avatar split clause 667,2354,1574,1570,1658,1562,1589]
2409. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 666,2354,1574,1570,1658,1594,1558,1554]
2410. 1 | 2 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 665,2354,1574,1570,1658,1594,1558,1554]
2411. 8 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 664,2354,1574,1570,1658,1594,1584]
2412. 9 | 10 | 15 | ~5 | ~6 | 21 [avatar split clause 663,2354,1574,1570,1658,1594,1589]
2413. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 662,2354,1574,1570,1658,1602,1558,1554]
2414. 1 | 2 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 661,2354,1574,1570,1658,1602,1558,1554]
2415. 8 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 660,2354,1574,1570,1658,1602,1584]
2416. 9 | 11 | 15 | ~5 | ~6 | 21 [avatar split clause 659,2354,1574,1570,1658,1602,1589]
2417. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 658,2354,1574,1570,1658,1610,1558,1554]
2418. 1 | 2 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 657,2354,1574,1570,1658,1610,1558,1554]
2419. 8 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 656,2354,1574,1570,1658,1610,1584]
2420. 9 | 12 | 15 | ~5 | ~6 | 21 [avatar split clause 655,2354,1574,1570,1658,1610,1589]
2421. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 654,2354,1574,1570,1678,1562,1558,1554]
2422. 1 | 2 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 653,2354,1574,1570,1678,1562,1558,1554]
2423. 8 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 652,2354,1574,1570,1678,1562,1584]
2424. 9 | ~3 | 16 | ~5 | ~6 | 21 [avatar split clause 651,2354,1574,1570,1678,1562,1589]
2425. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 650,2354,1574,1570,1678,1594,1558,1554]
2426. 1 | 2 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 649,2354,1574,1570,1678,1594,1558,1554]
2427. 8 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 648,2354,1574,1570,1678,1594,1584]
2428. 9 | 10 | 16 | ~5 | ~6 | 21 [avatar split clause 647,2354,1574,1570,1678,1594,1589]
2429. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 646,2354,1574,1570,1678,1602,1558,1554]
2430. 1 | 2 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 645,2354,1574,1570,1678,1602,1558,1554]
2431. 8 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 644,2354,1574,1570,1678,1602,1584]
2432. 9 | 11 | 16 | ~5 | ~6 | 21 [avatar split clause 643,2354,1574,1570,1678,1602,1589]
2433. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 642,2354,1574,1570,1678,1610,1558,1554]
2434. 1 | 2 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 641,2354,1574,1570,1678,1610,1558,1554]
2435. 8 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 640,2354,1574,1570,1678,1610,1584]
2436. 9 | 12 | 16 | ~5 | ~6 | 21 [avatar split clause 639,2354,1574,1570,1678,1610,1589]
2437. ~1 | ~2 | ~3 | ~4 | 17 | 21 [avatar split clause 638,2354,1698,1566,1562,1558,1554]
2438. 1 | 2 | ~3 | ~4 | 17 | 21 [avatar split clause 637,2354,1698,1566,1562,1558,1554]
2439. 8 | ~3 | ~4 | 17 | 21 [avatar split clause 636,2354,1698,1566,1562,1584]
2440. 9 | ~3 | ~4 | 17 | 21 [avatar split clause 635,2354,1698,1566,1562,1589]
2441. ~1 | ~2 | 10 | ~4 | 17 | 21 [avatar split clause 634,2354,1698,1566,1594,1558,1554]
2442. 1 | 2 | 10 | ~4 | 17 | 21 [avatar split clause 633,2354,1698,1566,1594,1558,1554]
2443. 8 | 10 | ~4 | 17 | 21 [avatar split clause 632,2354,1698,1566,1594,1584]
2444. 9 | 10 | ~4 | 17 | 21 [avatar split clause 631,2354,1698,1566,1594,1589]
2445. ~1 | ~2 | 11 | ~4 | 17 | 21 [avatar split clause 630,2354,1698,1566,1602,1558,1554]
2446. 1 | 2 | 11 | ~4 | 17 | 21 [avatar split clause 629,2354,1698,1566,1602,1558,1554]
2447. 8 | 11 | ~4 | 17 | 21 [avatar split clause 628,2354,1698,1566,1602,1584]
2448. 9 | 11 | ~4 | 17 | 21 [avatar split clause 627,2354,1698,1566,1602,1589]
2449. ~1 | ~2 | 12 | ~4 | 17 | 21 [avatar split clause 626,2354,1698,1566,1610,1558,1554]
2450. 1 | 2 | 12 | ~4 | 17 | 21 [avatar split clause 625,2354,1698,1566,1610,1558,1554]
2451. 8 | 12 | ~4 | 17 | 21 [avatar split clause 624,2354,1698,1566,1610,1584]
2452. 9 | 12 | ~4 | 17 | 21 [avatar split clause 623,2354,1698,1566,1610,1589]
2453. ~1 | ~2 | ~3 | 13 | 17 | 21 [avatar split clause 622,2354,1698,1618,1562,1558,1554]
2454. 1 | 2 | ~3 | 13 | 17 | 21 [avatar split clause 621,2354,1698,1618,1562,1558,1554]
2455. 8 | ~3 | 13 | 17 | 21 [avatar split clause 620,2354,1698,1618,1562,1584]
2456. 9 | ~3 | 13 | 17 | 21 [avatar split clause 619,2354,1698,1618,1562,1589]
2457. ~1 | ~2 | 10 | 13 | 17 | 21 [avatar split clause 618,2354,1698,1618,1594,1558,1554]
2458. 1 | 2 | 10 | 13 | 17 | 21 [avatar split clause 617,2354,1698,1618,1594,1558,1554]
2459. 8 | 10 | 13 | 17 | 21 [avatar split clause 616,2354,1698,1618,1594,1584]
2460. 9 | 10 | 13 | 17 | 21 [avatar split clause 615,2354,1698,1618,1594,1589]
2461. ~1 | ~2 | 11 | 13 | 17 | 21 [avatar split clause 614,2354,1698,1618,1602,1558,1554]
2462. 1 | 2 | 11 | 13 | 17 | 21 [avatar split clause 613,2354,1698,1618,1602,1558,1554]
2463. 8 | 11 | 13 | 17 | 21 [avatar split clause 612,2354,1698,1618,1602,1584]
2464. 9 | 11 | 13 | 17 | 21 [avatar split clause 611,2354,1698,1618,1602,1589]
2465. ~1 | ~2 | 12 | 13 | 17 | 21 [avatar split clause 610,2354,1698,1618,1610,1558,1554]
2466. 1 | 2 | 12 | 13 | 17 | 21 [avatar split clause 609,2354,1698,1618,1610,1558,1554]
2467. 8 | 12 | 13 | 17 | 21 [avatar split clause 608,2354,1698,1618,1610,1584]
2468. 9 | 12 | 13 | 17 | 21 [avatar split clause 607,2354,1698,1618,1610,1589]
2469. ~1 | ~2 | ~3 | 14 | 17 | 21 [avatar split clause 606,2354,1698,1638,1562,1558,1554]
2470. 1 | 2 | ~3 | 14 | 17 | 21 [avatar split clause 605,2354,1698,1638,1562,1558,1554]
2471. 8 | ~3 | 14 | 17 | 21 [avatar split clause 604,2354,1698,1638,1562,1584]
2472. 9 | ~3 | 14 | 17 | 21 [avatar split clause 603,2354,1698,1638,1562,1589]
2473. ~1 | ~2 | 10 | 14 | 17 | 21 [avatar split clause 602,2354,1698,1638,1594,1558,1554]
2474. 1 | 2 | 10 | 14 | 17 | 21 [avatar split clause 601,2354,1698,1638,1594,1558,1554]
2475. 8 | 10 | 14 | 17 | 21 [avatar split clause 600,2354,1698,1638,1594,1584]
2476. 9 | 10 | 14 | 17 | 21 [avatar split clause 599,2354,1698,1638,1594,1589]
2477. ~1 | ~2 | 11 | 14 | 17 | 21 [avatar split clause 598,2354,1698,1638,1602,1558,1554]
2478. 1 | 2 | 11 | 14 | 17 | 21 [avatar split clause 597,2354,1698,1638,1602,1558,1554]
2479. 8 | 11 | 14 | 17 | 21 [avatar split clause 596,2354,1698,1638,1602,1584]
2480. 9 | 11 | 14 | 17 | 21 [avatar split clause 595,2354,1698,1638,1602,1589]
2481. ~1 | ~2 | 12 | 14 | 17 | 21 [avatar split clause 594,2354,1698,1638,1610,1558,1554]
2482. 1 | 2 | 12 | 14 | 17 | 21 [avatar split clause 593,2354,1698,1638,1610,1558,1554]
2483. 8 | 12 | 14 | 17 | 21 [avatar split clause 592,2354,1698,1638,1610,1584]
2484. 9 | 12 | 14 | 17 | 21 [avatar split clause 591,2354,1698,1638,1610,1589]
2485. ~1 | ~2 | ~3 | 15 | 17 | 21 [avatar split clause 590,2354,1698,1658,1562,1558,1554]
2486. 1 | 2 | ~3 | 15 | 17 | 21 [avatar split clause 589,2354,1698,1658,1562,1558,1554]
2487. 8 | ~3 | 15 | 17 | 21 [avatar split clause 588,2354,1698,1658,1562,1584]
2488. 9 | ~3 | 15 | 17 | 21 [avatar split clause 587,2354,1698,1658,1562,1589]
2489. ~1 | ~2 | 10 | 15 | 17 | 21 [avatar split clause 586,2354,1698,1658,1594,1558,1554]
2490. 1 | 2 | 10 | 15 | 17 | 21 [avatar split clause 585,2354,1698,1658,1594,1558,1554]
2491. 8 | 10 | 15 | 17 | 21 [avatar split clause 584,2354,1698,1658,1594,1584]
2492. 9 | 10 | 15 | 17 | 21 [avatar split clause 583,2354,1698,1658,1594,1589]
2493. ~1 | ~2 | 11 | 15 | 17 | 21 [avatar split clause 582,2354,1698,1658,1602,1558,1554]
2494. 1 | 2 | 11 | 15 | 17 | 21 [avatar split clause 581,2354,1698,1658,1602,1558,1554]
2495. 8 | 11 | 15 | 17 | 21 [avatar split clause 580,2354,1698,1658,1602,1584]
2496. 9 | 11 | 15 | 17 | 21 [avatar split clause 579,2354,1698,1658,1602,1589]
2497. ~1 | ~2 | 12 | 15 | 17 | 21 [avatar split clause 578,2354,1698,1658,1610,1558,1554]
2498. 1 | 2 | 12 | 15 | 17 | 21 [avatar split clause 577,2354,1698,1658,1610,1558,1554]
2499. 8 | 12 | 15 | 17 | 21 [avatar split clause 576,2354,1698,1658,1610,1584]
2500. 9 | 12 | 15 | 17 | 21 [avatar split clause 575,2354,1698,1658,1610,1589]
2501. ~1 | ~2 | ~3 | 16 | 17 | 21 [avatar split clause 574,2354,1698,1678,1562,1558,1554]
2502. 1 | 2 | ~3 | 16 | 17 | 21 [avatar split clause 573,2354,1698,1678,1562,1558,1554]
2503. 8 | ~3 | 16 | 17 | 21 [avatar split clause 572,2354,1698,1678,1562,1584]
2504. 9 | ~3 | 16 | 17 | 21 [avatar split clause 571,2354,1698,1678,1562,1589]
2505. ~1 | ~2 | 10 | 16 | 17 | 21 [avatar split clause 570,2354,1698,1678,1594,1558,1554]
2506. 1 | 2 | 10 | 16 | 17 | 21 [avatar split clause 569,2354,1698,1678,1594,1558,1554]
2507. 8 | 10 | 16 | 17 | 21 [avatar split clause 568,2354,1698,1678,1594,1584]
2508. 9 | 10 | 16 | 17 | 21 [avatar split clause 567,2354,1698,1678,1594,1589]
2509. ~1 | ~2 | 11 | 16 | 17 | 21 [avatar split clause 566,2354,1698,1678,1602,1558,1554]
2510. 1 | 2 | 11 | 16 | 17 | 21 [avatar split clause 565,2354,1698,1678,1602,1558,1554]
2511. 8 | 11 | 16 | 17 | 21 [avatar split clause 564,2354,1698,1678,1602,1584]
2512. 9 | 11 | 16 | 17 | 21 [avatar split clause 563,2354,1698,1678,1602,1589]
2513. ~1 | ~2 | 12 | 16 | 17 | 21 [avatar split clause 562,2354,1698,1678,1610,1558,1554]
2514. 1 | 2 | 12 | 16 | 17 | 21 [avatar split clause 561,2354,1698,1678,1610,1558,1554]
2515. 8 | 12 | 16 | 17 | 21 [avatar split clause 560,2354,1698,1678,1610,1584]
2516. 9 | 12 | 16 | 17 | 21 [avatar split clause 559,2354,1698,1678,1610,1589]
2517. ~1 | ~2 | ~3 | ~4 | 18 | 21 [avatar split clause 558,2354,1782,1566,1562,1558,1554]
2518. 1 | 2 | ~3 | ~4 | 18 | 21 [avatar split clause 557,2354,1782,1566,1562,1558,1554]
2519. 8 | ~3 | ~4 | 18 | 21 [avatar split clause 556,2354,1782,1566,1562,1584]
2520. 9 | ~3 | ~4 | 18 | 21 [avatar split clause 555,2354,1782,1566,1562,1589]
2521. ~1 | ~2 | 10 | ~4 | 18 | 21 [avatar split clause 554,2354,1782,1566,1594,1558,1554]
2522. 1 | 2 | 10 | ~4 | 18 | 21 [avatar split clause 553,2354,1782,1566,1594,1558,1554]
2523. 8 | 10 | ~4 | 18 | 21 [avatar split clause 552,2354,1782,1566,1594,1584]
2524. 9 | 10 | ~4 | 18 | 21 [avatar split clause 551,2354,1782,1566,1594,1589]
2525. ~1 | ~2 | 11 | ~4 | 18 | 21 [avatar split clause 550,2354,1782,1566,1602,1558,1554]
2526. 1 | 2 | 11 | ~4 | 18 | 21 [avatar split clause 549,2354,1782,1566,1602,1558,1554]
2527. 8 | 11 | ~4 | 18 | 21 [avatar split clause 548,2354,1782,1566,1602,1584]
2528. 9 | 11 | ~4 | 18 | 21 [avatar split clause 547,2354,1782,1566,1602,1589]
2529. ~1 | ~2 | 12 | ~4 | 18 | 21 [avatar split clause 546,2354,1782,1566,1610,1558,1554]
2530. 1 | 2 | 12 | ~4 | 18 | 21 [avatar split clause 545,2354,1782,1566,1610,1558,1554]
2531. 8 | 12 | ~4 | 18 | 21 [avatar split clause 544,2354,1782,1566,1610,1584]
2532. 9 | 12 | ~4 | 18 | 21 [avatar split clause 543,2354,1782,1566,1610,1589]
2533. ~1 | ~2 | ~3 | 13 | 18 | 21 [avatar split clause 542,2354,1782,1618,1562,1558,1554]
2534. 1 | 2 | ~3 | 13 | 18 | 21 [avatar split clause 541,2354,1782,1618,1562,1558,1554]
2535. 8 | ~3 | 13 | 18 | 21 [avatar split clause 540,2354,1782,1618,1562,1584]
2536. 9 | ~3 | 13 | 18 | 21 [avatar split clause 539,2354,1782,1618,1562,1589]
2537. ~1 | ~2 | 10 | 13 | 18 | 21 [avatar split clause 538,2354,1782,1618,1594,1558,1554]
2538. 1 | 2 | 10 | 13 | 18 | 21 [avatar split clause 537,2354,1782,1618,1594,1558,1554]
2539. 8 | 10 | 13 | 18 | 21 [avatar split clause 536,2354,1782,1618,1594,1584]
2540. 9 | 10 | 13 | 18 | 21 [avatar split clause 535,2354,1782,1618,1594,1589]
2541. ~1 | ~2 | 11 | 13 | 18 | 21 [avatar split clause 534,2354,1782,1618,1602,1558,1554]
2542. 1 | 2 | 11 | 13 | 18 | 21 [avatar split clause 533,2354,1782,1618,1602,1558,1554]
2543. 8 | 11 | 13 | 18 | 21 [avatar split clause 532,2354,1782,1618,1602,1584]
2544. 9 | 11 | 13 | 18 | 21 [avatar split clause 531,2354,1782,1618,1602,1589]
2545. ~1 | ~2 | 12 | 13 | 18 | 21 [avatar split clause 530,2354,1782,1618,1610,1558,1554]
2546. 1 | 2 | 12 | 13 | 18 | 21 [avatar split clause 529,2354,1782,1618,1610,1558,1554]
2547. 8 | 12 | 13 | 18 | 21 [avatar split clause 528,2354,1782,1618,1610,1584]
2548. 9 | 12 | 13 | 18 | 21 [avatar split clause 527,2354,1782,1618,1610,1589]
2549. ~1 | ~2 | ~3 | 14 | 18 | 21 [avatar split clause 526,2354,1782,1638,1562,1558,1554]
2550. 1 | 2 | ~3 | 14 | 18 | 21 [avatar split clause 525,2354,1782,1638,1562,1558,1554]
2551. 8 | ~3 | 14 | 18 | 21 [avatar split clause 524,2354,1782,1638,1562,1584]
2552. 9 | ~3 | 14 | 18 | 21 [avatar split clause 523,2354,1782,1638,1562,1589]
2553. ~1 | ~2 | 10 | 14 | 18 | 21 [avatar split clause 522,2354,1782,1638,1594,1558,1554]
2554. 1 | 2 | 10 | 14 | 18 | 21 [avatar split clause 521,2354,1782,1638,1594,1558,1554]
2555. 8 | 10 | 14 | 18 | 21 [avatar split clause 520,2354,1782,1638,1594,1584]
2556. 9 | 10 | 14 | 18 | 21 [avatar split clause 519,2354,1782,1638,1594,1589]
2557. ~1 | ~2 | 11 | 14 | 18 | 21 [avatar split clause 518,2354,1782,1638,1602,1558,1554]
2558. 1 | 2 | 11 | 14 | 18 | 21 [avatar split clause 517,2354,1782,1638,1602,1558,1554]
2559. 8 | 11 | 14 | 18 | 21 [avatar split clause 516,2354,1782,1638,1602,1584]
2560. 9 | 11 | 14 | 18 | 21 [avatar split clause 515,2354,1782,1638,1602,1589]
2561. ~1 | ~2 | 12 | 14 | 18 | 21 [avatar split clause 514,2354,1782,1638,1610,1558,1554]
2562. 1 | 2 | 12 | 14 | 18 | 21 [avatar split clause 513,2354,1782,1638,1610,1558,1554]
2563. 8 | 12 | 14 | 18 | 21 [avatar split clause 512,2354,1782,1638,1610,1584]
2564. 9 | 12 | 14 | 18 | 21 [avatar split clause 511,2354,1782,1638,1610,1589]
2565. ~1 | ~2 | ~3 | 15 | 18 | 21 [avatar split clause 510,2354,1782,1658,1562,1558,1554]
2566. 1 | 2 | ~3 | 15 | 18 | 21 [avatar split clause 509,2354,1782,1658,1562,1558,1554]
2567. 8 | ~3 | 15 | 18 | 21 [avatar split clause 508,2354,1782,1658,1562,1584]
2568. 9 | ~3 | 15 | 18 | 21 [avatar split clause 507,2354,1782,1658,1562,1589]
2569. ~1 | ~2 | 10 | 15 | 18 | 21 [avatar split clause 506,2354,1782,1658,1594,1558,1554]
2570. 1 | 2 | 10 | 15 | 18 | 21 [avatar split clause 505,2354,1782,1658,1594,1558,1554]
2571. 8 | 10 | 15 | 18 | 21 [avatar split clause 504,2354,1782,1658,1594,1584]
2572. 9 | 10 | 15 | 18 | 21 [avatar split clause 503,2354,1782,1658,1594,1589]
2573. ~1 | ~2 | 11 | 15 | 18 | 21 [avatar split clause 502,2354,1782,1658,1602,1558,1554]
2574. 1 | 2 | 11 | 15 | 18 | 21 [avatar split clause 501,2354,1782,1658,1602,1558,1554]
2575. 8 | 11 | 15 | 18 | 21 [avatar split clause 500,2354,1782,1658,1602,1584]
2576. 9 | 11 | 15 | 18 | 21 [avatar split clause 499,2354,1782,1658,1602,1589]
2577. ~1 | ~2 | 12 | 15 | 18 | 21 [avatar split clause 498,2354,1782,1658,1610,1558,1554]
2578. 1 | 2 | 12 | 15 | 18 | 21 [avatar split clause 497,2354,1782,1658,1610,1558,1554]
2579. 8 | 12 | 15 | 18 | 21 [avatar split clause 496,2354,1782,1658,1610,1584]
2580. 9 | 12 | 15 | 18 | 21 [avatar split clause 495,2354,1782,1658,1610,1589]
2581. ~1 | ~2 | ~3 | 16 | 18 | 21 [avatar split clause 494,2354,1782,1678,1562,1558,1554]
2582. 1 | 2 | ~3 | 16 | 18 | 21 [avatar split clause 493,2354,1782,1678,1562,1558,1554]
2583. 8 | ~3 | 16 | 18 | 21 [avatar split clause 492,2354,1782,1678,1562,1584]
2584. 9 | ~3 | 16 | 18 | 21 [avatar split clause 491,2354,1782,1678,1562,1589]
2585. ~1 | ~2 | 10 | 16 | 18 | 21 [avatar split clause 490,2354,1782,1678,1594,1558,1554]
2586. 1 | 2 | 10 | 16 | 18 | 21 [avatar split clause 489,2354,1782,1678,1594,1558,1554]
2587. 8 | 10 | 16 | 18 | 21 [avatar split clause 488,2354,1782,1678,1594,1584]
2588. 9 | 10 | 16 | 18 | 21 [avatar split clause 487,2354,1782,1678,1594,1589]
2589. ~1 | ~2 | 11 | 16 | 18 | 21 [avatar split clause 486,2354,1782,1678,1602,1558,1554]
2590. 1 | 2 | 11 | 16 | 18 | 21 [avatar split clause 485,2354,1782,1678,1602,1558,1554]
2591. 8 | 11 | 16 | 18 | 21 [avatar split clause 484,2354,1782,1678,1602,1584]
2592. 9 | 11 | 16 | 18 | 21 [avatar split clause 483,2354,1782,1678,1602,1589]
2593. ~1 | ~2 | 12 | 16 | 18 | 21 [avatar split clause 482,2354,1782,1678,1610,1558,1554]
2594. 1 | 2 | 12 | 16 | 18 | 21 [avatar split clause 481,2354,1782,1678,1610,1558,1554]
2595. 8 | 12 | 16 | 18 | 21 [avatar split clause 480,2354,1782,1678,1610,1584]
2596. 9 | 12 | 16 | 18 | 21 [avatar split clause 479,2354,1782,1678,1610,1589]
2598. 22 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
2601. ~1 | ~2 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 478,2598,1574,1570,1566,1562,1558,1554]
2602. 1 | 2 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 477,2598,1574,1570,1566,1562,1558,1554]
2603. 8 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 476,2598,1574,1570,1566,1562,1584]
2604. 9 | ~3 | ~4 | ~5 | ~6 | 22 [avatar split clause 475,2598,1574,1570,1566,1562,1589]
2605. ~1 | ~2 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 474,2598,1574,1570,1566,1594,1558,1554]
2606. 1 | 2 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 473,2598,1574,1570,1566,1594,1558,1554]
2607. 8 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 472,2598,1574,1570,1566,1594,1584]
2608. 9 | 10 | ~4 | ~5 | ~6 | 22 [avatar split clause 471,2598,1574,1570,1566,1594,1589]
2609. ~1 | ~2 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 470,2598,1574,1570,1566,1602,1558,1554]
2610. 1 | 2 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 469,2598,1574,1570,1566,1602,1558,1554]
2611. 8 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 468,2598,1574,1570,1566,1602,1584]
2612. 9 | 11 | ~4 | ~5 | ~6 | 22 [avatar split clause 467,2598,1574,1570,1566,1602,1589]
2613. ~1 | ~2 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 466,2598,1574,1570,1566,1610,1558,1554]
2614. 1 | 2 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 465,2598,1574,1570,1566,1610,1558,1554]
2615. 8 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 464,2598,1574,1570,1566,1610,1584]
2616. 9 | 12 | ~4 | ~5 | ~6 | 22 [avatar split clause 463,2598,1574,1570,1566,1610,1589]
2617. ~1 | ~2 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 462,2598,1574,1570,1618,1562,1558,1554]
2618. 1 | 2 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 461,2598,1574,1570,1618,1562,1558,1554]
2619. 8 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 460,2598,1574,1570,1618,1562,1584]
2620. 9 | ~3 | 13 | ~5 | ~6 | 22 [avatar split clause 459,2598,1574,1570,1618,1562,1589]
2621. ~1 | ~2 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 458,2598,1574,1570,1618,1594,1558,1554]
2622. 1 | 2 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 457,2598,1574,1570,1618,1594,1558,1554]
2623. 8 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 456,2598,1574,1570,1618,1594,1584]
2624. 9 | 10 | 13 | ~5 | ~6 | 22 [avatar split clause 455,2598,1574,1570,1618,1594,1589]
2625. ~1 | ~2 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 454,2598,1574,1570,1618,1602,1558,1554]
2626. 1 | 2 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 453,2598,1574,1570,1618,1602,1558,1554]
2627. 8 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 452,2598,1574,1570,1618,1602,1584]
2628. 9 | 11 | 13 | ~5 | ~6 | 22 [avatar split clause 451,2598,1574,1570,1618,1602,1589]
2629. ~1 | ~2 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 450,2598,1574,1570,1618,1610,1558,1554]
2630. 1 | 2 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 449,2598,1574,1570,1618,1610,1558,1554]
2631. 8 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 448,2598,1574,1570,1618,1610,1584]
2632. 9 | 12 | 13 | ~5 | ~6 | 22 [avatar split clause 447,2598,1574,1570,1618,1610,1589]
2633. ~1 | ~2 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 446,2598,1574,1570,1638,1562,1558,1554]
2634. 1 | 2 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 445,2598,1574,1570,1638,1562,1558,1554]
2635. 8 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 444,2598,1574,1570,1638,1562,1584]
2636. 9 | ~3 | 14 | ~5 | ~6 | 22 [avatar split clause 443,2598,1574,1570,1638,1562,1589]
2637. ~1 | ~2 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 442,2598,1574,1570,1638,1594,1558,1554]
2638. 1 | 2 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 441,2598,1574,1570,1638,1594,1558,1554]
2639. 8 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 440,2598,1574,1570,1638,1594,1584]
2640. 9 | 10 | 14 | ~5 | ~6 | 22 [avatar split clause 439,2598,1574,1570,1638,1594,1589]
2641. ~1 | ~2 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 438,2598,1574,1570,1638,1602,1558,1554]
2642. 1 | 2 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 437,2598,1574,1570,1638,1602,1558,1554]
2643. 8 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 436,2598,1574,1570,1638,1602,1584]
2644. 9 | 11 | 14 | ~5 | ~6 | 22 [avatar split clause 435,2598,1574,1570,1638,1602,1589]
2645. ~1 | ~2 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 434,2598,1574,1570,1638,1610,1558,1554]
2646. 1 | 2 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 433,2598,1574,1570,1638,1610,1558,1554]
2647. 8 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 432,2598,1574,1570,1638,1610,1584]
2648. 9 | 12 | 14 | ~5 | ~6 | 22 [avatar split clause 431,2598,1574,1570,1638,1610,1589]
2649. ~1 | ~2 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 430,2598,1574,1570,1658,1562,1558,1554]
2650. 1 | 2 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 429,2598,1574,1570,1658,1562,1558,1554]
2651. 8 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 428,2598,1574,1570,1658,1562,1584]
2652. 9 | ~3 | 15 | ~5 | ~6 | 22 [avatar split clause 427,2598,1574,1570,1658,1562,1589]
2653. ~1 | ~2 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 426,2598,1574,1570,1658,1594,1558,1554]
2654. 1 | 2 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 425,2598,1574,1570,1658,1594,1558,1554]
2655. 8 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 424,2598,1574,1570,1658,1594,1584]
2656. 9 | 10 | 15 | ~5 | ~6 | 22 [avatar split clause 423,2598,1574,1570,1658,1594,1589]
2657. ~1 | ~2 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 422,2598,1574,1570,1658,1602,1558,1554]
2658. 1 | 2 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 421,2598,1574,1570,1658,1602,1558,1554]
2659. 8 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 420,2598,1574,1570,1658,1602,1584]
2660. 9 | 11 | 15 | ~5 | ~6 | 22 [avatar split clause 419,2598,1574,1570,1658,1602,1589]
2661. ~1 | ~2 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 418,2598,1574,1570,1658,1610,1558,1554]
2662. 1 | 2 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 417,2598,1574,1570,1658,1610,1558,1554]
2663. 8 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 416,2598,1574,1570,1658,1610,1584]
2664. 9 | 12 | 15 | ~5 | ~6 | 22 [avatar split clause 415,2598,1574,1570,1658,1610,1589]
2665. ~1 | ~2 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 414,2598,1574,1570,1678,1562,1558,1554]
2666. 1 | 2 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 413,2598,1574,1570,1678,1562,1558,1554]
2667. 8 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 412,2598,1574,1570,1678,1562,1584]
2668. 9 | ~3 | 16 | ~5 | ~6 | 22 [avatar split clause 411,2598,1574,1570,1678,1562,1589]
2669. ~1 | ~2 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 410,2598,1574,1570,1678,1594,1558,1554]
2670. 1 | 2 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 409,2598,1574,1570,1678,1594,1558,1554]
2671. 8 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 408,2598,1574,1570,1678,1594,1584]
2672. 9 | 10 | 16 | ~5 | ~6 | 22 [avatar split clause 407,2598,1574,1570,1678,1594,1589]
2673. ~1 | ~2 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 406,2598,1574,1570,1678,1602,1558,1554]
2674. 1 | 2 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 405,2598,1574,1570,1678,1602,1558,1554]
2675. 8 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 404,2598,1574,1570,1678,1602,1584]
2676. 9 | 11 | 16 | ~5 | ~6 | 22 [avatar split clause 403,2598,1574,1570,1678,1602,1589]
2677. ~1 | ~2 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 402,2598,1574,1570,1678,1610,1558,1554]
2678. 1 | 2 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 401,2598,1574,1570,1678,1610,1558,1554]
2679. 8 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 400,2598,1574,1570,1678,1610,1584]
2680. 9 | 12 | 16 | ~5 | ~6 | 22 [avatar split clause 399,2598,1574,1570,1678,1610,1589]
2681. ~1 | ~2 | ~3 | ~4 | 17 | 22 [avatar split clause 398,2598,1698,1566,1562,1558,1554]
2682. 1 | 2 | ~3 | ~4 | 17 | 22 [avatar split clause 397,2598,1698,1566,1562,1558,1554]
2683. 8 | ~3 | ~4 | 17 | 22 [avatar split clause 396,2598,1698,1566,1562,1584]
2684. 9 | ~3 | ~4 | 17 | 22 [avatar split clause 395,2598,1698,1566,1562,1589]
2685. ~1 | ~2 | 10 | ~4 | 17 | 22 [avatar split clause 394,2598,1698,1566,1594,1558,1554]
2686. 1 | 2 | 10 | ~4 | 17 | 22 [avatar split clause 393,2598,1698,1566,1594,1558,1554]
2687. 8 | 10 | ~4 | 17 | 22 [avatar split clause 392,2598,1698,1566,1594,1584]
2688. 9 | 10 | ~4 | 17 | 22 [avatar split clause 391,2598,1698,1566,1594,1589]
2689. ~1 | ~2 | 11 | ~4 | 17 | 22 [avatar split clause 390,2598,1698,1566,1602,1558,1554]
2690. 1 | 2 | 11 | ~4 | 17 | 22 [avatar split clause 389,2598,1698,1566,1602,1558,1554]
2691. 8 | 11 | ~4 | 17 | 22 [avatar split clause 388,2598,1698,1566,1602,1584]
2692. 9 | 11 | ~4 | 17 | 22 [avatar split clause 387,2598,1698,1566,1602,1589]
2693. ~1 | ~2 | 12 | ~4 | 17 | 22 [avatar split clause 386,2598,1698,1566,1610,1558,1554]
2694. 1 | 2 | 12 | ~4 | 17 | 22 [avatar split clause 385,2598,1698,1566,1610,1558,1554]
2695. 8 | 12 | ~4 | 17 | 22 [avatar split clause 384,2598,1698,1566,1610,1584]
2696. 9 | 12 | ~4 | 17 | 22 [avatar split clause 383,2598,1698,1566,1610,1589]
2697. ~1 | ~2 | ~3 | 13 | 17 | 22 [avatar split clause 382,2598,1698,1618,1562,1558,1554]
2698. 1 | 2 | ~3 | 13 | 17 | 22 [avatar split clause 381,2598,1698,1618,1562,1558,1554]
2699. 8 | ~3 | 13 | 17 | 22 [avatar split clause 380,2598,1698,1618,1562,1584]
2700. 9 | ~3 | 13 | 17 | 22 [avatar split clause 379,2598,1698,1618,1562,1589]
2701. ~1 | ~2 | 10 | 13 | 17 | 22 [avatar split clause 378,2598,1698,1618,1594,1558,1554]
2702. 1 | 2 | 10 | 13 | 17 | 22 [avatar split clause 377,2598,1698,1618,1594,1558,1554]
2703. 8 | 10 | 13 | 17 | 22 [avatar split clause 376,2598,1698,1618,1594,1584]
2704. 9 | 10 | 13 | 17 | 22 [avatar split clause 375,2598,1698,1618,1594,1589]
2705. ~1 | ~2 | 11 | 13 | 17 | 22 [avatar split clause 374,2598,1698,1618,1602,1558,1554]
2706. 1 | 2 | 11 | 13 | 17 | 22 [avatar split clause 373,2598,1698,1618,1602,1558,1554]
2707. 8 | 11 | 13 | 17 | 22 [avatar split clause 372,2598,1698,1618,1602,1584]
2708. 9 | 11 | 13 | 17 | 22 [avatar split clause 371,2598,1698,1618,1602,1589]
2709. ~1 | ~2 | 12 | 13 | 17 | 22 [avatar split clause 370,2598,1698,1618,1610,1558,1554]
2710. 1 | 2 | 12 | 13 | 17 | 22 [avatar split clause 369,2598,1698,1618,1610,1558,1554]
2711. 8 | 12 | 13 | 17 | 22 [avatar split clause 368,2598,1698,1618,1610,1584]
2712. 9 | 12 | 13 | 17 | 22 [avatar split clause 367,2598,1698,1618,1610,1589]
2713. ~1 | ~2 | ~3 | 14 | 17 | 22 [avatar split clause 366,2598,1698,1638,1562,1558,1554]
2714. 1 | 2 | ~3 | 14 | 17 | 22 [avatar split clause 365,2598,1698,1638,1562,1558,1554]
2715. 8 | ~3 | 14 | 17 | 22 [avatar split clause 364,2598,1698,1638,1562,1584]
2716. 9 | ~3 | 14 | 17 | 22 [avatar split clause 363,2598,1698,1638,1562,1589]
2717. ~1 | ~2 | 10 | 14 | 17 | 22 [avatar split clause 362,2598,1698,1638,1594,1558,1554]
2718. 1 | 2 | 10 | 14 | 17 | 22 [avatar split clause 361,2598,1698,1638,1594,1558,1554]
2719. 8 | 10 | 14 | 17 | 22 [avatar split clause 360,2598,1698,1638,1594,1584]
2720. 9 | 10 | 14 | 17 | 22 [avatar split clause 359,2598,1698,1638,1594,1589]
2721. ~1 | ~2 | 11 | 14 | 17 | 22 [avatar split clause 358,2598,1698,1638,1602,1558,1554]
2722. 1 | 2 | 11 | 14 | 17 | 22 [avatar split clause 357,2598,1698,1638,1602,1558,1554]
2723. 8 | 11 | 14 | 17 | 22 [avatar split clause 356,2598,1698,1638,1602,1584]
2724. 9 | 11 | 14 | 17 | 22 [avatar split clause 355,2598,1698,1638,1602,1589]
2725. ~1 | ~2 | 12 | 14 | 17 | 22 [avatar split clause 354,2598,1698,1638,1610,1558,1554]
2726. 1 | 2 | 12 | 14 | 17 | 22 [avatar split clause 353,2598,1698,1638,1610,1558,1554]
2727. 8 | 12 | 14 | 17 | 22 [avatar split clause 352,2598,1698,1638,1610,1584]
2728. 9 | 12 | 14 | 17 | 22 [avatar split clause 351,2598,1698,1638,1610,1589]
2729. ~1 | ~2 | ~3 | 15 | 17 | 22 [avatar split clause 350,2598,1698,1658,1562,1558,1554]
2730. 1 | 2 | ~3 | 15 | 17 | 22 [avatar split clause 349,2598,1698,1658,1562,1558,1554]
2731. 8 | ~3 | 15 | 17 | 22 [avatar split clause 348,2598,1698,1658,1562,1584]
2732. 9 | ~3 | 15 | 17 | 22 [avatar split clause 347,2598,1698,1658,1562,1589]
2733. ~1 | ~2 | 10 | 15 | 17 | 22 [avatar split clause 346,2598,1698,1658,1594,1558,1554]
2734. 1 | 2 | 10 | 15 | 17 | 22 [avatar split clause 345,2598,1698,1658,1594,1558,1554]
2735. 8 | 10 | 15 | 17 | 22 [avatar split clause 344,2598,1698,1658,1594,1584]
2736. 9 | 10 | 15 | 17 | 22 [avatar split clause 343,2598,1698,1658,1594,1589]
2737. ~1 | ~2 | 11 | 15 | 17 | 22 [avatar split clause 342,2598,1698,1658,1602,1558,1554]
2738. 1 | 2 | 11 | 15 | 17 | 22 [avatar split clause 341,2598,1698,1658,1602,1558,1554]
2739. 8 | 11 | 15 | 17 | 22 [avatar split clause 340,2598,1698,1658,1602,1584]
2740. 9 | 11 | 15 | 17 | 22 [avatar split clause 339,2598,1698,1658,1602,1589]
2741. ~1 | ~2 | 12 | 15 | 17 | 22 [avatar split clause 338,2598,1698,1658,1610,1558,1554]
2742. 1 | 2 | 12 | 15 | 17 | 22 [avatar split clause 337,2598,1698,1658,1610,1558,1554]
2743. 8 | 12 | 15 | 17 | 22 [avatar split clause 336,2598,1698,1658,1610,1584]
2744. 9 | 12 | 15 | 17 | 22 [avatar split clause 335,2598,1698,1658,1610,1589]
2745. ~1 | ~2 | ~3 | 16 | 17 | 22 [avatar split clause 334,2598,1698,1678,1562,1558,1554]
2746. 1 | 2 | ~3 | 16 | 17 | 22 [avatar split clause 333,2598,1698,1678,1562,1558,1554]
2747. 8 | ~3 | 16 | 17 | 22 [avatar split clause 332,2598,1698,1678,1562,1584]
2748. 9 | ~3 | 16 | 17 | 22 [avatar split clause 331,2598,1698,1678,1562,1589]
2749. ~1 | ~2 | 10 | 16 | 17 | 22 [avatar split clause 330,2598,1698,1678,1594,1558,1554]
2750. 1 | 2 | 10 | 16 | 17 | 22 [avatar split clause 329,2598,1698,1678,1594,1558,1554]
2751. 8 | 10 | 16 | 17 | 22 [avatar split clause 328,2598,1698,1678,1594,1584]
2752. 9 | 10 | 16 | 17 | 22 [avatar split clause 327,2598,1698,1678,1594,1589]
2753. ~1 | ~2 | 11 | 16 | 17 | 22 [avatar split clause 326,2598,1698,1678,1602,1558,1554]
2754. 1 | 2 | 11 | 16 | 17 | 22 [avatar split clause 325,2598,1698,1678,1602,1558,1554]
2755. 8 | 11 | 16 | 17 | 22 [avatar split clause 324,2598,1698,1678,1602,1584]
2756. 9 | 11 | 16 | 17 | 22 [avatar split clause 323,2598,1698,1678,1602,1589]
2757. ~1 | ~2 | 12 | 16 | 17 | 22 [avatar split clause 322,2598,1698,1678,1610,1558,1554]
2758. 1 | 2 | 12 | 16 | 17 | 22 [avatar split clause 321,2598,1698,1678,1610,1558,1554]
2759. 8 | 12 | 16 | 17 | 22 [avatar split clause 320,2598,1698,1678,1610,1584]
2760. 9 | 12 | 16 | 17 | 22 [avatar split clause 319,2598,1698,1678,1610,1589]
2761. ~1 | ~2 | ~3 | ~4 | 18 | 22 [avatar split clause 318,2598,1782,1566,1562,1558,1554]
2762. 1 | 2 | ~3 | ~4 | 18 | 22 [avatar split clause 317,2598,1782,1566,1562,1558,1554]
2763. 8 | ~3 | ~4 | 18 | 22 [avatar split clause 316,2598,1782,1566,1562,1584]
2764. 9 | ~3 | ~4 | 18 | 22 [avatar split clause 315,2598,1782,1566,1562,1589]
2765. ~1 | ~2 | 10 | ~4 | 18 | 22 [avatar split clause 314,2598,1782,1566,1594,1558,1554]
2766. 1 | 2 | 10 | ~4 | 18 | 22 [avatar split clause 313,2598,1782,1566,1594,1558,1554]
2767. 8 | 10 | ~4 | 18 | 22 [avatar split clause 312,2598,1782,1566,1594,1584]
2768. 9 | 10 | ~4 | 18 | 22 [avatar split clause 311,2598,1782,1566,1594,1589]
2769. ~1 | ~2 | 11 | ~4 | 18 | 22 [avatar split clause 310,2598,1782,1566,1602,1558,1554]
2770. 1 | 2 | 11 | ~4 | 18 | 22 [avatar split clause 309,2598,1782,1566,1602,1558,1554]
2771. 8 | 11 | ~4 | 18 | 22 [avatar split clause 308,2598,1782,1566,1602,1584]
2772. 9 | 11 | ~4 | 18 | 22 [avatar split clause 307,2598,1782,1566,1602,1589]
2773. ~1 | ~2 | 12 | ~4 | 18 | 22 [avatar split clause 306,2598,1782,1566,1610,1558,1554]
2774. 1 | 2 | 12 | ~4 | 18 | 22 [avatar split clause 305,2598,1782,1566,1610,1558,1554]
2775. 8 | 12 | ~4 | 18 | 22 [avatar split clause 304,2598,1782,1566,1610,1584]
2776. 9 | 12 | ~4 | 18 | 22 [avatar split clause 303,2598,1782,1566,1610,1589]
2777. ~1 | ~2 | ~3 | 13 | 18 | 22 [avatar split clause 302,2598,1782,1618,1562,1558,1554]
2778. 1 | 2 | ~3 | 13 | 18 | 22 [avatar split clause 301,2598,1782,1618,1562,1558,1554]
2779. 8 | ~3 | 13 | 18 | 22 [avatar split clause 300,2598,1782,1618,1562,1584]
2780. 9 | ~3 | 13 | 18 | 22 [avatar split clause 299,2598,1782,1618,1562,1589]
2781. ~1 | ~2 | 10 | 13 | 18 | 22 [avatar split clause 298,2598,1782,1618,1594,1558,1554]
2782. 1 | 2 | 10 | 13 | 18 | 22 [avatar split clause 297,2598,1782,1618,1594,1558,1554]
2783. 8 | 10 | 13 | 18 | 22 [avatar split clause 296,2598,1782,1618,1594,1584]
2784. 9 | 10 | 13 | 18 | 22 [avatar split clause 295,2598,1782,1618,1594,1589]
2785. ~1 | ~2 | 11 | 13 | 18 | 22 [avatar split clause 294,2598,1782,1618,1602,1558,1554]
2786. 1 | 2 | 11 | 13 | 18 | 22 [avatar split clause 293,2598,1782,1618,1602,1558,1554]
2787. 8 | 11 | 13 | 18 | 22 [avatar split clause 292,2598,1782,1618,1602,1584]
2788. 9 | 11 | 13 | 18 | 22 [avatar split clause 291,2598,1782,1618,1602,1589]
2789. ~1 | ~2 | 12 | 13 | 18 | 22 [avatar split clause 290,2598,1782,1618,1610,1558,1554]
2790. 1 | 2 | 12 | 13 | 18 | 22 [avatar split clause 289,2598,1782,1618,1610,1558,1554]
2791. 8 | 12 | 13 | 18 | 22 [avatar split clause 288,2598,1782,1618,1610,1584]
2792. 9 | 12 | 13 | 18 | 22 [avatar split clause 287,2598,1782,1618,1610,1589]
2793. ~1 | ~2 | ~3 | 14 | 18 | 22 [avatar split clause 286,2598,1782,1638,1562,1558,1554]
2794. 1 | 2 | ~3 | 14 | 18 | 22 [avatar split clause 285,2598,1782,1638,1562,1558,1554]
2795. 8 | ~3 | 14 | 18 | 22 [avatar split clause 284,2598,1782,1638,1562,1584]
2796. 9 | ~3 | 14 | 18 | 22 [avatar split clause 283,2598,1782,1638,1562,1589]
2797. ~1 | ~2 | 10 | 14 | 18 | 22 [avatar split clause 282,2598,1782,1638,1594,1558,1554]
2798. 1 | 2 | 10 | 14 | 18 | 22 [avatar split clause 281,2598,1782,1638,1594,1558,1554]
2799. 8 | 10 | 14 | 18 | 22 [avatar split clause 280,2598,1782,1638,1594,1584]
2800. 9 | 10 | 14 | 18 | 22 [avatar split clause 279,2598,1782,1638,1594,1589]
2801. ~1 | ~2 | 11 | 14 | 18 | 22 [avatar split clause 278,2598,1782,1638,1602,1558,1554]
2802. 1 | 2 | 11 | 14 | 18 | 22 [avatar split clause 277,2598,1782,1638,1602,1558,1554]
2803. 8 | 11 | 14 | 18 | 22 [avatar split clause 276,2598,1782,1638,1602,1584]
2804. 9 | 11 | 14 | 18 | 22 [avatar split clause 275,2598,1782,1638,1602,1589]
2805. ~1 | ~2 | 12 | 14 | 18 | 22 [avatar split clause 274,2598,1782,1638,1610,1558,1554]
2806. 1 | 2 | 12 | 14 | 18 | 22 [avatar split clause 273,2598,1782,1638,1610,1558,1554]
2807. 8 | 12 | 14 | 18 | 22 [avatar split clause 272,2598,1782,1638,1610,1584]
2808. 9 | 12 | 14 | 18 | 22 [avatar split clause 271,2598,1782,1638,1610,1589]
2809. ~1 | ~2 | ~3 | 15 | 18 | 22 [avatar split clause 270,2598,1782,1658,1562,1558,1554]
2810. 1 | 2 | ~3 | 15 | 18 | 22 [avatar split clause 269,2598,1782,1658,1562,1558,1554]
2811. 8 | ~3 | 15 | 18 | 22 [avatar split clause 268,2598,1782,1658,1562,1584]
2812. 9 | ~3 | 15 | 18 | 22 [avatar split clause 267,2598,1782,1658,1562,1589]
2813. ~1 | ~2 | 10 | 15 | 18 | 22 [avatar split clause 266,2598,1782,1658,1594,1558,1554]
2814. 1 | 2 | 10 | 15 | 18 | 22 [avatar split clause 265,2598,1782,1658,1594,1558,1554]
2815. 8 | 10 | 15 | 18 | 22 [avatar split clause 264,2598,1782,1658,1594,1584]
2816. 9 | 10 | 15 | 18 | 22 [avatar split clause 263,2598,1782,1658,1594,1589]
2817. ~1 | ~2 | 11 | 15 | 18 | 22 [avatar split clause 262,2598,1782,1658,1602,1558,1554]
2818. 1 | 2 | 11 | 15 | 18 | 22 [avatar split clause 261,2598,1782,1658,1602,1558,1554]
2819. 8 | 11 | 15 | 18 | 22 [avatar split clause 260,2598,1782,1658,1602,1584]
2820. 9 | 11 | 15 | 18 | 22 [avatar split clause 259,2598,1782,1658,1602,1589]
2821. ~1 | ~2 | 12 | 15 | 18 | 22 [avatar split clause 258,2598,1782,1658,1610,1558,1554]
2822. 1 | 2 | 12 | 15 | 18 | 22 [avatar split clause 257,2598,1782,1658,1610,1558,1554]
2823. 8 | 12 | 15 | 18 | 22 [avatar split clause 256,2598,1782,1658,1610,1584]
2824. 9 | 12 | 15 | 18 | 22 [avatar split clause 255,2598,1782,1658,1610,1589]
2825. ~1 | ~2 | ~3 | 16 | 18 | 22 [avatar split clause 254,2598,1782,1678,1562,1558,1554]
2826. 1 | 2 | ~3 | 16 | 18 | 22 [avatar split clause 253,2598,1782,1678,1562,1558,1554]
2827. 8 | ~3 | 16 | 18 | 22 [avatar split clause 252,2598,1782,1678,1562,1584]
2828. 9 | ~3 | 16 | 18 | 22 [avatar split clause 251,2598,1782,1678,1562,1589]
2829. ~1 | ~2 | 10 | 16 | 18 | 22 [avatar split clause 250,2598,1782,1678,1594,1558,1554]
2830. 1 | 2 | 10 | 16 | 18 | 22 [avatar split clause 249,2598,1782,1678,1594,1558,1554]
2831. 8 | 10 | 16 | 18 | 22 [avatar split clause 248,2598,1782,1678,1594,1584]
2832. 9 | 10 | 16 | 18 | 22 [avatar split clause 247,2598,1782,1678,1594,1589]
2833. ~1 | ~2 | 11 | 16 | 18 | 22 [avatar split clause 246,2598,1782,1678,1602,1558,1554]
2834. 1 | 2 | 11 | 16 | 18 | 22 [avatar split clause 245,2598,1782,1678,1602,1558,1554]
2835. 8 | 11 | 16 | 18 | 22 [avatar split clause 244,2598,1782,1678,1602,1584]
2836. 9 | 11 | 16 | 18 | 22 [avatar split clause 243,2598,1782,1678,1602,1589]
2837. ~1 | ~2 | 12 | 16 | 18 | 22 [avatar split clause 242,2598,1782,1678,1610,1558,1554]
2838. 1 | 2 | 12 | 16 | 18 | 22 [avatar split clause 241,2598,1782,1678,1610,1558,1554]
2839. 8 | 12 | 16 | 18 | 22 [avatar split clause 240,2598,1782,1678,1610,1584]
2840. 9 | 12 | 16 | 18 | 22 [avatar split clause 239,2598,1782,1678,1610,1589]
2842. 23 <=> l2_bcialg_1(sK0) [avatar definition]
2844. l2_bcialg_1(sK0) <- (23) [avatar component clause 2842] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2845. 23 [avatar split clause 236,2842]
2847. 24 <=> v8_bcialg_1(sK0) [avatar definition]
2849. v8_bcialg_1(sK0) <- (24) [avatar component clause 2847] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2850. 24 [avatar split clause 235,2847]
2852. 25 <=> v7_bcialg_1(sK0) [avatar definition]
2854. v7_bcialg_1(sK0) <- (25) [avatar component clause 2852] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2855. 25 [avatar split clause 234,2852]
2857. 26 <=> v5_bcialg_1(sK0) [avatar definition]
2859. v5_bcialg_1(sK0) <- (26) [avatar component clause 2857] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2860. 26 [avatar split clause 233,2857]
2862. 27 <=> v4_bcialg_1(sK0) [avatar definition]
2864. v4_bcialg_1(sK0) <- (27) [avatar component clause 2862] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2865. 27 [avatar split clause 232,2862]
2867. 28 <=> v3_bcialg_1(sK0) [avatar definition]
2869. v3_bcialg_1(sK0) <- (28) [avatar component clause 2867] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2870. 28 [avatar split clause 231,2867]
2872. 29 <=> v2_struct_0(sK0) [avatar definition]
2874. ~v2_struct_0(sK0) <- (~29) [avatar component clause 2872] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:14,thDist:-14}
2875. ~29 [avatar split clause 230,2872]
3306. v1_bcialg_3(sK0) | ~l2_bcialg_1(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) [resolution 1496,237] {a:1,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3309. v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23) [subsumption resolution 3306,2844] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3310. v1_bcialg_3(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24) [subsumption resolution 3309,2849] {a:1,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3311. v1_bcialg_3(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24, 25) [subsumption resolution 3310,2854] {a:1,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3312. v1_bcialg_3(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24, 25, 26) [subsumption resolution 3311,2859] {a:1,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3313. v1_bcialg_3(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24, 25, 26, 27) [subsumption resolution 3312,2864] {a:1,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3314. v1_bcialg_3(sK0) | v2_struct_0(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24, 25, 26, 27, 28) [subsumption resolution 3313,2869] {a:1,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3315. v1_bcialg_3(sK0) | v3_bciideal(sK16(sK0),sK0) <- (23, 24, 25, 26, 27, 28, ~29) [subsumption resolution 3314,2874] {a:1,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:16,thDist:-16}
3317. 90 <=> v3_bciideal(sK16(sK0),sK0) [avatar definition]
3319. v3_bciideal(sK16(sK0),sK0) <- (90) [avatar component clause 3317] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
3321. 91 <=> v1_bcialg_3(sK0) [avatar definition]
3323. v1_bcialg_3(sK0) <- (91) [avatar component clause 3321] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:16,thDist:-16}
3324. 90 | 91 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 [avatar split clause 3315,2872,2867,2862,2857,2852,2847,2842,3321,3317]
3335. ~v3_bciideal(sK16(X0),X0) | v1_bcialg_3(X0) | ~l2_bcialg_1(X0) | ~v8_bcialg_1(X0) | ~v7_bcialg_1(X0) | ~v5_bcialg_1(X0) | ~v4_bcialg_1(X0) | ~v3_bcialg_1(X0) | v2_struct_0(X0) [subsumption resolution 1497,1496] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
3337. v1_bcialg_3(sK0) | ~l2_bcialg_1(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (90) [resolution 3335,3319] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3338. v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 90) [subsumption resolution 3337,2844] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3339. v1_bcialg_3(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 90) [subsumption resolution 3338,2849] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3340. v1_bcialg_3(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 90) [subsumption resolution 3339,2854] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3341. v1_bcialg_3(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 90) [subsumption resolution 3340,2859] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3342. v1_bcialg_3(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27, 90) [subsumption resolution 3341,2864] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3343. v1_bcialg_3(sK0) | v2_struct_0(sK0) <- (23, 24, 25, 26, 27, 28, 90) [subsumption resolution 3342,2869] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3344. v1_bcialg_3(sK0) <- (23, 24, 25, 26, 27, 28, ~29, 90) [subsumption resolution 3343,2874] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:18,thDist:-18}
3345. 91 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~90 [avatar split clause 3344,3317,2872,2867,2862,2857,2852,2847,2842,3321]
3362. sK9 != sK9 | sK10 = k1_bcialg_1(sK0,sK10,sK9) | ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (10) [superposition 1550,1596] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3365. sK10 = k1_bcialg_1(sK0,sK10,sK9) | ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (10) [trivial inequality removal 3362] {a:1,w:30,wCS:3000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3385. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,sK5) | ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5) [superposition 1572,1551] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3390. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5) [trivial inequality removal 3385] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3529. sK11 != sK11 | k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1) [superposition 1548,1555] {a:1,w:37,wCS:3700,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3534. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1) [trivial inequality removal 3529] {a:1,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3550. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8) [subsumption resolution 3534,1586] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3551. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9) [subsumption resolution 3550,1591] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3552. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23) [subsumption resolution 3551,2844] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3553. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23, 91) [subsumption resolution 3552,3323] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3554. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23, 24, 91) [subsumption resolution 3553,2849] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3555. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23, 24, 25, 91) [subsumption resolution 3554,2854] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3556. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23, 24, 25, 26, 91) [subsumption resolution 3555,2859] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3557. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (1, 8, 9, 23, 24, 25, 26, 27, 91) [subsumption resolution 3556,2864] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3558. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) | v2_struct_0(sK0) <- (1, 8, 9, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 3557,2869] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3559. k4_struct_0(sK0) = k1_bcialg_1(sK0,sK12,k1_bcialg_1(sK0,sK12,sK11)) <- (1, 8, 9, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 3558,2874] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3560. 2 | ~1 | ~8 | ~9 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar split clause 3559,3321,2872,2867,2862,2857,2852,2847,2842,1589,1584,1554,1558]
3583. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23) [subsumption resolution 3390,2844] {a:1,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3584. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23, 91) [subsumption resolution 3583,3323] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3585. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23, 24, 91) [subsumption resolution 3584,2849] {a:1,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3586. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23, 24, 25, 91) [subsumption resolution 3585,2854] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3587. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23, 24, 25, 26, 91) [subsumption resolution 3586,2859] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3588. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~5, 23, 24, 25, 26, 27, 91) [subsumption resolution 3587,2864] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3589. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~5, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 3588,2869] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3590. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) <- (~5, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 3589,2874] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3591. ~18 | ~17 | 5 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar split clause 3590,3321,2872,2867,2862,2857,2852,2847,2842,1570,1698,1782]
3695. k1_bcialg_1(sK0,sK4,sK5) != k1_bcialg_1(sK0,sK4,sK5) | ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6) [superposition 1576,1552] {a:1,w:33,wCS:3300,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3709. ~m1_subset_1(sK5,u1_struct_0(sK0)) | ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6) [trivial inequality removal 3695] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3723. ~m1_subset_1(sK4,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17) [subsumption resolution 3709,1700] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3724. ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18) [subsumption resolution 3723,1784] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3725. ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23) [subsumption resolution 3724,2844] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3726. ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23, 91) [subsumption resolution 3725,3323] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3727. ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23, 24, 91) [subsumption resolution 3726,2849] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3728. ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23, 24, 25, 91) [subsumption resolution 3727,2854] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3729. ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23, 24, 25, 26, 91) [subsumption resolution 3728,2859] {a:1,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3730. ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~6, 17, 18, 23, 24, 25, 26, 27, 91) [subsumption resolution 3729,2864] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3731. v2_struct_0(sK0) <- (~6, 17, 18, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 3730,2869] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3732. $false <- (~6, 17, 18, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 3731,2874] {a:1,w:0,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
3733. 6 | ~17 | ~18 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar contradiction clause 3732]
3747. k4_struct_0(sK0) != k4_struct_0(sK0) | sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2) [superposition 1549,1559] {a:1,w:35,wCS:3500,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3755. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~m1_subset_1(sK12,u1_struct_0(sK0)) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2) [trivial inequality removal 3747] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3756. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~m1_subset_1(sK11,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8) [subsumption resolution 3755,1586] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3757. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9) [subsumption resolution 3756,1591] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3758. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23) [subsumption resolution 3757,2844] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3759. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23, 91) [subsumption resolution 3758,3323] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3760. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23, 24, 91) [subsumption resolution 3759,2849] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3761. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23, 24, 25, 91) [subsumption resolution 3760,2854] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3762. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23, 24, 25, 26, 91) [subsumption resolution 3761,2859] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3763. sK11 = k1_bcialg_1(sK0,sK11,sK12) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (2, 8, 9, 23, 24, 25, 26, 27, 91) [subsumption resolution 3762,2864] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3764. sK11 = k1_bcialg_1(sK0,sK11,sK12) | v2_struct_0(sK0) <- (2, 8, 9, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 3763,2869] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3765. sK11 = k1_bcialg_1(sK0,sK11,sK12) <- (2, 8, 9, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 3764,2874] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3766. 1 | ~2 | ~8 | ~9 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar split clause 3765,3321,2872,2867,2862,2857,2852,2847,2842,1589,1584,1558,1554]
3899. k1_bcialg_1(sK0,sK7,sK6) != k1_bcialg_1(sK0,sK7,sK6) | ~r1_bcialg_1(sK0,sK7,sK8) | ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4) [superposition 1568,1547] {a:1,w:41,wCS:4100,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3913. ~r1_bcialg_1(sK0,sK7,sK8) | ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4) [trivial inequality removal 3899] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3978. ~m1_subset_1(sK8,u1_struct_0(sK0)) | ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13) [subsumption resolution 3913,1620] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3979. ~m1_subset_1(sK7,u1_struct_0(sK0)) | ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14) [subsumption resolution 3978,1640] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3980. ~m1_subset_1(sK6,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15) [subsumption resolution 3979,1660] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3981. ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16) [subsumption resolution 3980,1680] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3982. ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23) [subsumption resolution 3981,2844] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3983. ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 91) [subsumption resolution 3982,3323] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3984. ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 24, 91) [subsumption resolution 3983,2849] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3985. ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 24, 25, 91) [subsumption resolution 3984,2854] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3986. ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 24, 25, 26, 91) [subsumption resolution 3985,2859] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3987. ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 24, 25, 26, 27, 91) [subsumption resolution 3986,2864] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3988. v2_struct_0(sK0) <- (~4, 13, 14, 15, 16, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 3987,2869] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3989. $false <- (~4, 13, 14, 15, 16, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 3988,2874] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
3990. 4 | ~13 | ~14 | ~15 | ~16 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar contradiction clause 3989]
4096. k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) != k1_bcialg_1(sK0,k1_bcialg_1(sK0,sK3,sK2),k1_bcialg_1(sK0,sK1,sK2)) | ~r1_bcialg_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7) [superposition 1580,1547] {a:1,w:53,wCS:5300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4097. ~r1_bcialg_1(sK0,sK1,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7) [trivial inequality removal 4096] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4098. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19) [subsumption resolution 4097,1868] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4099. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20) [subsumption resolution 4098,2112] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4138. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10) [subsumption resolution 3365,1564] {a:1,w:24,wCS:2400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4139. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23) [subsumption resolution 4138,2844] {a:1,w:22,wCS:2200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4140. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23, 91) [subsumption resolution 4139,3323] {a:1,w:20,wCS:2000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4141. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23, 24, 91) [subsumption resolution 4140,2849] {a:1,w:18,wCS:1800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4142. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23, 24, 25, 91) [subsumption resolution 4141,2854] {a:1,w:16,wCS:1600,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4143. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23, 24, 25, 26, 91) [subsumption resolution 4142,2859] {a:1,w:14,wCS:1400,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4144. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~3, 10, 23, 24, 25, 26, 27, 91) [subsumption resolution 4143,2864] {a:1,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4145. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~3, 10, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 4144,2869] {a:1,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4146. ~m1_subset_1(sK10,u1_struct_0(sK0)) | ~m1_subset_1(sK9,u1_struct_0(sK0)) <- (~3, 10, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 4145,2874] {a:1,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:15,thDist:-15}
4147. ~12 | ~11 | 3 | ~10 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar split clause 4146,3321,2872,2867,2862,2857,2852,2847,2842,1594,1562,1602,1610]
4205. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l2_bcialg_1(sK0) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21) [subsumption resolution 4099,2356] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4206. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v1_bcialg_3(sK0) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23) [subsumption resolution 4205,2844] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4207. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v8_bcialg_1(sK0) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 91) [subsumption resolution 4206,3323] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4208. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v7_bcialg_1(sK0) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 24, 91) [subsumption resolution 4207,2849] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4209. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v5_bcialg_1(sK0) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 24, 25, 91) [subsumption resolution 4208,2854] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4210. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v4_bcialg_1(sK0) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 24, 25, 26, 91) [subsumption resolution 4209,2859] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4211. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~v3_bcialg_1(sK0) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 24, 25, 26, 27, 91) [subsumption resolution 4210,2864] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4212. ~m1_subset_1(sK1,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (~7, 19, 20, 21, 23, 24, 25, 26, 27, 28, 91) [subsumption resolution 4211,2869] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4213. ~m1_subset_1(sK1,u1_struct_0(sK0)) <- (~7, 19, 20, 21, 23, 24, 25, 26, 27, 28, ~29, 91) [subsumption resolution 4212,2874] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
4214. ~22 | 7 | ~19 | ~20 | ~21 | ~23 | ~24 | ~25 | ~26 | ~27 | ~28 | 29 | ~91 [avatar split clause 4213,3321,2872,2867,2862,2857,2852,2847,2842,2354,2110,1866,1578,2598]
4215. $false [avatar sat refutation 1581,1582,1587,1592,1597,1598,1599,1600,1605,1606,1607,1608,1613,1614,1615,1616,1621,1622,1623,1624,1625,1626,1627,1628,1629,1630,1631,1632,1633,1634,1635,1636,1641,1642,1643,1644,1645,1646,1647,1648,1649,1650,1651,1652,1653,1654,1655,1656,1661,1662,1663,1664,1665,1666,1667,1668,1669,1670,1671,1672,1673,1674,1675,1676,1681,1682,1683,1684,1685,1686,1687,1688,1689,1690,1691,1692,1693,1694,1695,1696,1701,1702,1703,1704,1705,1706,1707,1708,1709,1710,1711,1712,1713,1714,1715,1716,1717,1718,1719,1720,1721,1722,1723,1724,1725,1726,1727,1728,1729,1730,1731,1732,1733,1734,1735,1736,1737,1738,1739,1740,1741,1742,1743,1744,1745,1746,1747,1748,1749,1750,1751,1752,1753,1754,1755,1756,1757,1758,1759,1760,1761,1762,1763,1764,1765,1766,1767,1768,1769,1770,1771,1772,1773,1774,1775,1776,1777,1778,1779,1780,1785,1786,1787,1788,1789,1790,1791,1792,1793,1794,1795,1796,1797,1798,1799,1800,1801,1802,1803,1804,1805,1806,1807,1808,1809,1810,1811,1812,1813,1814,1815,1816,1817,1818,1819,1820,1821,1822,1823,1824,1825,1826,1827,1828,1829,1830,1831,1832,1833,1834,1835,1836,1837,1838,1839,1840,1841,1842,1843,1844,1845,1846,1847,1848,1849,1850,1851,1852,1853,1854,1855,1856,1857,1858,1859,1860,1861,1862,1863,1864,1869,1870,1871,1872,1873,1874,1875,1876,1877,1878,1879,1880,1881,1882,1883,1884,1885,1886,1887,1888,1889,1890,1891,1892,1893,1894,1895,1896,1897,1898,1899,1900,1901,1902,1903,1904,1905,1906,1907,1908,1909,1910,1911,1912,1913,1914,1915,1916,1917,1918,1919,1920,1921,1922,1923,1924,1925,1926,1927,1928,1929,1930,1931,1932,1933,1934,1935,1936,1937,1938,1939,1940,1941,1942,1943,1944,1945,1946,1947,1948,1949,1950,1951,1952,1953,1954,1955,1956,1957,1958,1959,1960,1961,1962,1963,1964,1965,1966,1967,1968,1969,1970,1971,1972,1973,1974,1975,1976,1977,1978,1979,1980,1981,1982,1983,1984,1985,1986,1987,1988,1989,1990,1991,1992,1993,1994,1995,1996,1997,1998,1999,2000,2001,2002,2003,2004,2005,2006,2007,2008,2009,2010,2011,2012,2013,2014,2015,2016,2017,2018,2019,2020,2021,2022,2023,2024,2025,2026,2027,2028,2029,2030,2031,2032,2033,2034,2035,2036,2037,2038,2039,2040,2041,2042,2043,2044,2045,2046,2047,2048,2049,2050,2051,2052,2053,2054,2055,2056,2057,2058,2059,2060,2061,2062,2063,2064,2065,2066,2067,2068,2069,2070,2071,2072,2073,2074,2075,2076,2077,2078,2079,2080,2081,2082,2083,2084,2085,2086,2087,2088,2089,2090,2091,2092,2093,2094,2095,2096,2097,2098,2099,2100,2101,2102,2103,2104,2105,2106,2107,2108,2113,2114,2115,2116,2117,2118,2119,2120,2121,2122,2123,2124,2125,2126,2127,2128,2129,2130,2131,2132,2133,2134,2135,2136,2137,2138,2139,2140,2141,2142,2143,2144,2145,2146,2147,2148,2149,2150,2151,2152,2153,2154,2155,2156,2157,2158,2159,2160,2161,2162,2163,2164,2165,2166,2167,2168,2169,2170,2171,2172,2173,2174,2175,2176,2177,2178,2179,2180,2181,2182,2183,2184,2185,2186,2187,2188,2189,2190,2191,2192,2193,2194,2195,2196,2197,2198,2199,2200,2201,2202,2203,2204,2205,2206,2207,2208,2209,2210,2211,2212,2213,2214,2215,2216,2217,2218,2219,2220,2221,2222,2223,2224,2225,2226,2227,2228,2229,2230,2231,2232,2233,2234,2235,2236,2237,2238,2239,2240,2241,2242,2243,2244,2245,2246,2247,2248,2249,2250,2251,2252,2253,2254,2255,2256,2257,2258,2259,2260,2261,2262,2263,2264,2265,2266,2267,2268,2269,2270,2271,2272,2273,2274,2275,2276,2277,2278,2279,2280,2281,2282,2283,2284,2285,2286,2287,2288,2289,2290,2291,2292,2293,2294,2295,2296,2297,2298,2299,2300,2301,2302,2303,2304,2305,2306,2307,2308,2309,2310,2311,2312,2313,2314,2315,2316,2317,2318,2319,2320,2321,2322,2323,2324,2325,2326,2327,2328,2329,2330,2331,2332,2333,2334,2335,2336,2337,2338,2339,2340,2341,2342,2343,2344,2345,2346,2347,2348,2349,2350,2351,2352,2357,2358,2359,2360,2361,2362,2363,2364,2365,2366,2367,2368,2369,2370,2371,2372,2373,2374,2375,2376,2377,2378,2379,2380,2381,2382,2383,2384,2385,2386,2387,2388,2389,2390,2391,2392,2393,2394,2395,2396,2397,2398,2399,2400,2401,2402,2403,2404,2405,2406,2407,2408,2409,2410,2411,2412,2413,2414,2415,2416,2417,2418,2419,2420,2421,2422,2423,2424,2425,2426,2427,2428,2429,2430,2431,2432,2433,2434,2435,2436,2437,2438,2439,2440,2441,2442,2443,2444,2445,2446,2447,2448,2449,2450,2451,2452,2453,2454,2455,2456,2457,2458,2459,2460,2461,2462,2463,2464,2465,2466,2467,2468,2469,2470,2471,2472,2473,2474,2475,2476,2477,2478,2479,2480,2481,2482,2483,2484,2485,2486,2487,2488,2489,2490,2491,2492,2493,2494,2495,2496,2497,2498,2499,2500,2501,2502,2503,2504,2505,2506,2507,2508,2509,2510,2511,2512,2513,2514,2515,2516,2517,2518,2519,2520,2521,2522,2523,2524,2525,2526,2527,2528,2529,2530,2531,2532,2533,2534,2535,2536,2537,2538,2539,2540,2541,2542,2543,2544,2545,2546,2547,2548,2549,2550,2551,2552,2553,2554,2555,2556,2557,2558,2559,2560,2561,2562,2563,2564,2565,2566,2567,2568,2569,2570,2571,2572,2573,2574,2575,2576,2577,2578,2579,2580,2581,2582,2583,2584,2585,2586,2587,2588,2589,2590,2591,2592,2593,2594,2595,2596,2601,2602,2603,2604,2605,2606,2607,2608,2609,2610,2611,2612,2613,2614,2615,2616,2617,2618,2619,2620,2621,2622,2623,2624,2625,2626,2627,2628,2629,2630,2631,2632,2633,2634,2635,2636,2637,2638,2639,2640,2641,2642,2643,2644,2645,2646,2647,2648,2649,2650,2651,2652,2653,2654,2655,2656,2657,2658,2659,2660,2661,2662,2663,2664,2665,2666,2667,2668,2669,2670,2671,2672,2673,2674,2675,2676,2677,2678,2679,2680,2681,2682,2683,2684,2685,2686,2687,2688,2689,2690,2691,2692,2693,2694,2695,2696,2697,2698,2699,2700,2701,2702,2703,2704,2705,2706,2707,2708,2709,2710,2711,2712,2713,2714,2715,2716,2717,2718,2719,2720,2721,2722,2723,2724,2725,2726,2727,2728,2729,2730,2731,2732,2733,2734,2735,2736,2737,2738,2739,2740,2741,2742,2743,2744,2745,2746,2747,2748,2749,2750,2751,2752,2753,2754,2755,2756,2757,2758,2759,2760,2761,2762,2763,2764,2765,2766,2767,2768,2769,2770,2771,2772,2773,2774,2775,2776,2777,2778,2779,2780,2781,2782,2783,2784,2785,2786,2787,2788,2789,2790,2791,2792,2793,2794,2795,2796,2797,2798,2799,2800,2801,2802,2803,2804,2805,2806,2807,2808,2809,2810,2811,2812,2813,2814,2815,2816,2817,2818,2819,2820,2821,2822,2823,2824,2825,2826,2827,2828,2829,2830,2831,2832,2833,2834,2835,2836,2837,2838,2839,2840,2845,2850,2855,2860,2865,2870,2875,3324,3345,3560,3591,3733,3766,3990,4147,4214] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:26693,thDist:-26693}
% SZS output end Proof for bciideal__t49_bciideal
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 72
% >>> Preprocessing
% Introduced skolems: 29
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 1314
% Generated clauses: 2319
% Active clauses: 187
% Passive clauses: 275
% Final active clauses: 143
% Final passive clauses: 78
% Main loop iterations started: 193
% 
% >>> Simplifying Inferences
% Duplicate literals: 334
% Trivial inequalities: 15
% Fw subsumption resolutions: 536
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 32
% Forward subsumptions: 48
% 
% >>> Generating Inferences
% Binary resolution: 127
% Unit resulting resolution: 5
% Forward superposition: 34
% Backward superposition: 48
% Self superposition: 6
% 
% >>> AVATAR
% Split clauses: 1253
% Split components: 7131
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 1343
% SAT solver unit clauses: 25
% SAT solver binary clauses: 27
% 
% Memory used [KB]: 9594
% Time elapsed: 0.038 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% preprocessing: 0.001 s
% SAT solver time: 0.003 s
% sat proof minimization: 0.005 s
% congruence closure: 0.002 s
% binary resolution index maintenance: 0.001 s
% splitting component index usage: 0.001 s
% forward demodulation: 0.002 s
% unit resulting resolution: 0.001 s
% superposition: 0.001 s
% other: 0.038 s ( own 0.020 s ) 

