i: [1,32,0,2,1,1,0,0] t33_funct_1
a: [3,60,0,2,1,32]
s: 60
i: [1,31,0,2,1,1,0,0] t33_funct_1
a: [3,65,0,2,1,31]
s: 65
i: [1,30,0,2,1,1,0,0] t33_funct_1
a: [3,70,0,2,1,30]
s: 70
i: [1,33,0,12,2,1,0,0] t33_funct_1
a: [3,51,0,6,1,33]
s: 51
i: [1,44,0,2,1,0,0,1] rc1_funct_1
a: [3,75,0,2,1,44]
s: 75
i: [1,43,0,2,1,0,0,1] rc1_funct_1
a: [3,80,0,2,1,43]
s: 80
i: [1,47,0,2,1,0,0,1] rc2_funct_1
a: [3,85,0,2,1,47]
s: 85
i: [1,46,0,2,1,0,0,1] rc2_funct_1
a: [3,90,0,2,1,46]
s: 90
i: [1,45,0,2,1,0,0,1] rc2_funct_1
a: [3,95,0,2,1,45]
s: 95
i: [1,34,0,5,2,0,0,2] dt_k2_relat_1
s: 34
i: [1,35,0,7,2,0,0,1] involutiveness_k2_relat_1
s: 35
i: [1,40,0,7,3,0,0,1] dt_k2_funct_1
s: 40
i: [1,41,0,7,3,0,0,1] dt_k2_funct_1
s: 41
i: [1,36,0,8,2,0,0,1] t20_relat_1
s: 36
i: [1,37,0,8,2,0,0,1] t20_relat_1
s: 37
i: [1,38,0,9,4,0,0,1] fc5_funct_1
s: 38
d: [2,98,1,11,3,0,44,36,35]
d: [2,99,1,8,2,0,21,98,34]
s: 99
i: [1,39,0,9,4,0,0,1] fc5_funct_1
s: 39
d: [2,101,1,13,5,0,44,39,35]
d: [2,102,1,10,4,0,21,101,34]
s: 102
i: [1,42,0,11,4,0,0,1] d5_funct_1
s: 42
d: [2,106,1,13,5,0,44,35,42]
d: [2,116,1,11,4,0,18,106]
s: 116
d: [2,108,1,14,5,0,44,37,42]
d: [2,114,1,12,4,0,18,108]
s: 114
d: [2,107,1,14,5,0,44,36,42]
d: [2,115,1,12,4,0,18,107]
s: 115
d: [2,135,2,0,0,4,52,65,60,70,51,115]
e: 135
d: [2,136,2,11,4,1,44,51,115]
d: [2,137,2,6,3,1,19,136]
d: [2,139,2,4,2,2,21,137,70]
d: [2,140,2,2,1,3,21,139,60]
d: [2,141,2,0,0,4,21,140,65]
e: 141
r: 51
a: [3,50,0,6,1,33]
s: 50
a: [3,55,0,6,1,33]
s: 55
d: [2,143,2,0,0,4,52,65,60,70,55,114]
e: 143
d: [2,144,2,11,4,1,44,55,114]
d: [2,145,2,6,3,1,19,144]
d: [2,147,2,4,2,2,21,145,70]
d: [2,148,2,2,1,3,21,147,60]
d: [2,149,2,0,0,4,21,148,65]
e: 149
f: 149,141,30,31,32,33
% Refutation found. Thanks to Tanya!
% SZS status Theorem for funct_1__t33_funct_1
% SZS output start Proof for funct_1__t33_funct_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [input t33_funct_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_funct_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_funct_1(X0))))) [negated conjecture 1]
3. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v2_funct_1(X0) => k2_funct_1(X0) = k2_relat_1(X0))) [input d5_funct_1]
12. ! [X0] : (v1_relat_1(X0) => (k9_xtuple_0(X0) = k10_xtuple_0(k2_relat_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0)))) [input t20_relat_1]
13. ? [X0] : (((k9_xtuple_0(X0) != k10_xtuple_0(k2_funct_1(X0)) | k10_xtuple_0(X0) != k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0)) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
14. ? [X0] : ((k9_xtuple_0(X0) != k10_xtuple_0(k2_funct_1(X0)) | k10_xtuple_0(X0) != k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 13]
17. ! [X0] : ((k9_xtuple_0(X0) = k10_xtuple_0(k2_relat_1(X0)) & k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 12]
22. ! [X0] : ((k2_funct_1(X0) = k2_relat_1(X0) | ~v2_funct_1(X0)) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 3]
23. ! [X0] : (k2_funct_1(X0) = k2_relat_1(X0) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 22]
24. ? [X0] : ((k9_xtuple_0(X0) != k10_xtuple_0(k2_funct_1(X0)) | k10_xtuple_0(X0) != k9_xtuple_0(k2_funct_1(X0))) & v2_funct_1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((k9_xtuple_0(sK0) != k10_xtuple_0(k2_funct_1(sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
25. (k9_xtuple_0(sK0) != k10_xtuple_0(k2_funct_1(sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0))) & v2_funct_1(sK0) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 14,24]
30. v1_relat_1(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
31. v1_funct_1(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
32. v2_funct_1(sK0) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
33. k9_xtuple_0(sK0) != k10_xtuple_0(k2_funct_1(sK0)) | k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0)) [cnf transformation 25] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
36. k10_xtuple_0(X0) = k9_xtuple_0(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. k9_xtuple_0(X0) = k10_xtuple_0(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
42. k2_funct_1(X0) = k2_relat_1(X0) | ~v2_funct_1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
49. 1 <=> k10_xtuple_0(sK0) = k9_xtuple_0(k2_funct_1(sK0)) [avatar definition]
51. k10_xtuple_0(sK0) != k9_xtuple_0(k2_funct_1(sK0)) <- (~1) [avatar component clause 49] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
53. 2 <=> k9_xtuple_0(sK0) = k10_xtuple_0(k2_funct_1(sK0)) [avatar definition]
55. k9_xtuple_0(sK0) != k10_xtuple_0(k2_funct_1(sK0)) <- (~2) [avatar component clause 53] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
56. ~1 | ~2 [avatar split clause 33,53,49]
58. 3 <=> v2_funct_1(sK0) [avatar definition]
60. v2_funct_1(sK0) <- (3) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
61. 3 [avatar split clause 32,58]
63. 4 <=> v1_funct_1(sK0) [avatar definition]
65. v1_funct_1(sK0) <- (4) [avatar component clause 63] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. 4 [avatar split clause 31,63]
68. 5 <=> v1_relat_1(sK0) [avatar definition]
70. v1_relat_1(sK0) <- (5) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. 5 [avatar split clause 30,68]
107. k10_xtuple_0(X2) = k9_xtuple_0(k2_funct_1(X2)) | ~v1_relat_1(X2) | ~v2_funct_1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) [superposition 36,42] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
108. k9_xtuple_0(X3) = k10_xtuple_0(k2_funct_1(X3)) | ~v1_relat_1(X3) | ~v2_funct_1(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) [superposition 37,42] {a:1,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
114. k9_xtuple_0(X3) = k10_xtuple_0(k2_funct_1(X3)) | ~v1_relat_1(X3) | ~v2_funct_1(X3) | ~v1_funct_1(X3) [duplicate literal removal 108] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
115. k10_xtuple_0(X2) = k9_xtuple_0(k2_funct_1(X2)) | ~v1_relat_1(X2) | ~v2_funct_1(X2) | ~v1_funct_1(X2) [duplicate literal removal 107] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
136. k10_xtuple_0(sK0) != k10_xtuple_0(sK0) | ~v1_relat_1(sK0) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~1) [superposition 51,115] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
137. ~v1_relat_1(sK0) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~1) [trivial inequality removal 136] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
139. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~1, 5) [subsumption resolution 137,70] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
140. ~v1_funct_1(sK0) <- (~1, 3, 5) [subsumption resolution 139,60] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. $false <- (~1, 3, 4, 5) [subsumption resolution 140,65] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
142. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 141]
144. k9_xtuple_0(sK0) != k9_xtuple_0(sK0) | ~v1_relat_1(sK0) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~2) [superposition 55,114] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
145. ~v1_relat_1(sK0) | ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~2) [trivial inequality removal 144] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
147. ~v2_funct_1(sK0) | ~v1_funct_1(sK0) <- (~2, 5) [subsumption resolution 145,70] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
148. ~v1_funct_1(sK0) <- (~2, 3, 5) [subsumption resolution 147,60] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
149. $false <- (~2, 3, 4, 5) [subsumption resolution 148,65] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
150. 2 | ~3 | ~4 | ~5 [avatar contradiction clause 149]
151. $false [avatar sat refutation 56,61,66,71,142,150] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:56,thDist:-56}
% SZS output end Proof for funct_1__t33_funct_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 18
% Generated clauses: 79
% Active clauses: 23
% Passive clauses: 31
% Final active clauses: 22
% Final passive clauses: 6
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 10
% Trivial inequalities: 2
% Fw subsumption resolutions: 16
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Unit resulting resolution: 2
% Forward superposition: 7
% Backward superposition: 16
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

