i: [1,82,0,3,1,1,0,0] t9_arytm_1
a: [3,128,0,3,1,82]
s: 128
i: [1,81,0,3,1,1,0,0] t9_arytm_1
a: [3,133,0,3,1,81]
s: 133
i: [1,80,0,3,1,1,0,0] t9_arytm_1
a: [3,138,0,3,1,80]
s: 138
i: [1,83,0,5,1,1,0,0] t9_arytm_1
a: [3,123,0,5,1,83]
s: 123
i: [1,102,0,3,1,0,0,255] reflexivity_r1_tarski
s: 102
i: [1,84,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,143,0,3,1,84]
s: 143
i: [1,85,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,148,0,3,1,85]
s: 148
i: [1,114,0,11,3,0,0,255] d8_arytm_2
a: [3,158,0,3,1,114]
s: 158
i: [1,101,0,4,1,0,0,1] existence_m1_subset_1
s: 101
i: [1,86,0,5,1,0,0,255] dt_k1_arytm_2
a: [3,153,0,5,1,86]
s: 153
i: [1,96,0,6,2,0,0,255] t1_subset
s: 96
i: [1,103,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 103
i: [1,87,0,7,2,0,0,255] dt_k4_arytm_2
s: 87
i: [1,98,0,7,2,0,0,255] t3_subset
s: 98
i: [1,97,0,8,3,0,0,255] t2_subset
s: 97
i: [1,88,0,9,2,0,0,255] dt_k3_arytm_2
s: 88
d: [2,172,1,5,2,1,40,97,133]
a: [3,185,1,2,1,172]
s: 185
i: [1,100,0,9,3,0,0,255] t5_subset
s: 100
d: [2,175,1,5,2,1,40,97,148]
a: [3,199,1,3,1,175]
s: 199
i: [1,99,0,10,3,0,0,255] t4_subset
s: 99
i: [1,110,0,11,3,0,0,1] dt_k1_arytm_1
s: 110
d: [2,236,1,9,3,1,44,110,123]
d: [2,237,1,6,2,2,21,236,158]
d: [2,238,1,3,1,3,21,237,133]
d: [2,239,1,0,0,4,21,238,138]
e: 239
r: 158
a: [3,157,0,3,1,114]
s: 157
a: [3,161,0,8,2,114]
s: 161
d: [2,214,1,6,2,1,40,100,153]
a: [3,223,1,3,1,214]
s: 223
i: [1,115,0,11,3,0,0,255] d8_arytm_2
a: [3,165,0,8,2,115]
s: 165
i: [1,111,0,11,3,0,0,2] dt_k7_arytm_2
s: 111
d: [2,225,2,3,1,1,40,199,103]
a: [3,230,2,3,1,225]
s: 230
i: [1,113,0,12,4,0,0,1] connectedness_r1_arytm_2
s: 113
d: [2,254,1,9,3,1,40,113,138]
s: 254
i: [1,108,0,13,3,0,0,255] commutativity_k5_arytm_2
s: 108
i: [1,109,0,13,3,0,0,255] dt_k5_arytm_2
s: 109
d: [2,258,2,6,2,2,40,254,157]
a: [3,268,2,3,1,258]
s: 268
i: [1,112,0,13,3,0,0,2] commutativity_k7_arytm_2
s: 112
d: [2,261,2,6,2,1,40,254,138]
d: [2,264,2,3,1,1,18,261]
a: [3,277,2,3,1,264]
s: 277
i: [1,107,0,14,4,0,0,255] existence_m2_subset_1
s: 107
i: [1,116,0,17,5,0,0,1] d1_arytm_1
a: [3,169,0,14,4,116]
s: 169
d: [2,262,2,6,2,2,40,254,133]
d: [2,279,2,3,1,3,21,262,128]
a: [3,283,2,3,1,279]
s: 283
i: [1,104,0,15,5,0,0,255] dt_m2_subset_1
s: 104
r: 223
d: [2,263,2,8,2,1,40,254,101]
a: [3,288,2,4,1,263]
s: 288
i: [1,105,0,15,5,0,0,255] redefinition_m2_subset_1
s: 105
i: [1,106,0,15,5,0,0,255] redefinition_m2_subset_1
s: 106
d: [2,255,1,9,3,1,40,113,133]
s: 255
i: [1,117,0,19,5,0,0,1] d1_arytm_1
s: 117
d: [2,364,2,6,2,2,40,255,157]
a: [3,374,2,3,1,364]
s: 374
i: [1,119,0,21,5,0,0,1] d1_arytm_1
s: 119
d: [2,394,0,16,4,0,21,119,110]
s: 394
d: [2,368,2,6,2,1,40,255,133]
d: [2,370,2,3,1,1,18,368]
a: [3,383,2,3,1,370]
s: 383
i: [1,91,0,22,5,0,0,255] d8_arytm_2
s: 91
d: [2,369,2,8,2,1,40,255,101]
a: [3,388,2,4,1,369]
s: 388
i: [1,118,0,23,5,0,0,1] d1_arytm_1
s: 118
d: [2,427,0,18,4,0,21,118,111]
s: 427
d: [2,395,1,14,4,1,44,394,123]
d: [2,410,1,11,3,2,21,395,283]
d: [2,411,1,8,2,3,21,410,133]
d: [2,412,1,5,1,4,21,411,138]
a: [3,416,1,5,1,412]
s: 416
d: [2,241,1,6,2,2,21,236,133]
d: [2,242,1,3,1,3,21,241,138]
d: [2,446,2,6,2,2,44,161,416]
d: [2,455,2,3,1,3,21,446,133]
f: 455,412,279,264,242,114,80,81,82,83
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t9_arytm_1
% SZS output start Proof for arytm_1__t9_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(k1_arytm_1(X0,X1) = k11_arytm_3 & ~r1_arytm_2(X0,X1)))) [input t9_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ~(k1_arytm_1(X0,X1) = k11_arytm_3 & ~r1_arytm_2(X0,X1)))) [negated conjecture 1]
6. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1))) [input connectedness_r1_arytm_2]
7. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2)) & (r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0)))))) [input d1_arytm_1]
8. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (~(k7_arytm_2(X0,X1) != k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) & k11_arytm_3 != X0 & k11_arytm_3 != X1) & (k11_arytm_3 = X0 => k7_arytm_2(X0,X1) = X1) & (k11_arytm_3 = X1 => k7_arytm_2(X0,X1) = X0)))) [input d8_arytm_2]
10. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2)) [input dt_k1_arytm_1]
39. ? [X0] : (? [X1] : ((k1_arytm_1(X0,X1) = k11_arytm_3 & ~r1_arytm_2(X0,X1)) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
40. ? [X0] : (? [X1] : (k1_arytm_1(X0,X1) = k11_arytm_3 & ~r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 39]
43. ! [X0] : (! [X1] : (((k7_arytm_2(X0,X1) = k4_arytm_2(k5_arytm_2(k3_arytm_2(X0),k3_arytm_2(X1))) | k11_arytm_3 = X0 | k11_arytm_3 = X1) & (k7_arytm_2(X0,X1) = X1 | k11_arytm_3 != X0) & (k7_arytm_2(X0,X1) = X0 | k11_arytm_3 != X1)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 8]
44. ! [X0] : (! [X1] : (! [X2] : ((((k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2) | r1_arytm_2(X1,X0)) & ((k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 7]
63. ! [X0,X1] : (m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 10]
64. ! [X0,X1] : (m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 63]
69. ! [X0,X1] : ((r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1)) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 6]
70. ! [X0,X1] : (r1_arytm_2(X1,X0) | r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 69]
71. ? [X0] : (? [X1] : (k1_arytm_1(X0,X1) = k11_arytm_3 & ~r1_arytm_2(X0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (k11_arytm_3 = k1_arytm_1(sK0,X1) & ~r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
72. ? [X1] : (k11_arytm_3 = k1_arytm_1(sK0,X1) & ~r1_arytm_2(sK0,X1) & m1_subset_1(X1,k2_arytm_2)) => (k11_arytm_3 = k1_arytm_1(sK0,sK1) & ~r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
73. (k11_arytm_3 = k1_arytm_1(sK0,sK1) & ~r1_arytm_2(sK0,sK1) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 40,72,71]
74. ! [X0] : (! [X1] : (! [X2] : (((((k1_arytm_1(X0,X1) = X2 | k11_arytm_3 != X2) & (k11_arytm_3 = X2 | k1_arytm_1(X0,X1) != X2)) | r1_arytm_2(X1,X0)) & (((k1_arytm_1(X0,X1) = X2 | k7_arytm_2(X2,X1) != X0) & (k7_arytm_2(X2,X1) = X0 | k1_arytm_1(X0,X1) != X2)) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [nnf transformation 44]
80. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 73] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
81. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 73] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
82. ~r1_arytm_2(sK0,sK1) [cnf transformation 73] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
83. k11_arytm_3 = k1_arytm_1(sK0,sK1) [cnf transformation 73] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
90. k7_arytm_2(X0,X1) = X1 | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 43] {a:0,w:14,wCS:5600,sine:255,thAx:0,allAx:1,thDist:-1}
92. k7_arytm_2(X2,X1) = X0 | k1_arytm_1(X0,X1) != X2 | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 74] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:1,thDist:-1}
110. m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 64] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
113. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(X0,X1) | r1_arytm_2(X1,X0) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 70] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
114. k7_arytm_2(k11_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 90] {a:0,w:11,wCS:4400,sine:255,thAx:0,allAx:1,thDist:-1}
119. k7_arytm_2(k1_arytm_1(X0,X1),X1) = X0 | ~r1_arytm_2(X1,X0) | ~m1_subset_1(k1_arytm_1(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [equality resolution 92] {a:0,w:21,wCS:8400,sine:1,thAx:0,allAx:1,thDist:-1}
121. 1 <=> k11_arytm_3 = k1_arytm_1(sK0,sK1) [avatar definition]
123. k11_arytm_3 = k1_arytm_1(sK0,sK1) <- (1) [avatar component clause 121] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
124. 1 [avatar split clause 83,121]
126. 2 <=> r1_arytm_2(sK0,sK1) [avatar definition]
128. ~r1_arytm_2(sK0,sK1) <- (~2) [avatar component clause 126] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
129. ~2 [avatar split clause 82,126]
131. 3 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
133. m1_subset_1(sK1,k2_arytm_2) <- (3) [avatar component clause 131] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. 3 [avatar split clause 81,131]
136. 4 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
138. m1_subset_1(sK0,k2_arytm_2) <- (4) [avatar component clause 136] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
139. 4 [avatar split clause 80,136]
156. 8 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
160. 9 <=> ! [X1] : (k7_arytm_2(k11_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
161. k7_arytm_2(k11_arytm_3,X1) = X1 | ~m1_subset_1(X1,k2_arytm_2) <- (9) [avatar component clause 160] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
162. ~8 | 9 [avatar split clause 114,160,156]
236. m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (1) [superposition 110,123] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
241. m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (1, 3) [subsumption resolution 236,133] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
242. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (1, 3, 4) [subsumption resolution 241,138] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
243. 8 | ~1 | ~3 | ~4 [avatar split clause 242,136,131,121,156]
254. ~m1_subset_1(X9,k2_arytm_2) | r1_arytm_2(sK0,X9) | r1_arytm_2(X9,sK0) <- (4) [resolution 113,138] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
261. r1_arytm_2(sK0,sK0) | r1_arytm_2(sK0,sK0) <- (4) [resolution 254,138] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
262. r1_arytm_2(sK0,sK1) | r1_arytm_2(sK1,sK0) <- (3, 4) [resolution 254,133] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
264. r1_arytm_2(sK0,sK0) <- (4) [duplicate literal removal 261] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
275. 25 <=> r1_arytm_2(sK0,sK0) [avatar definition]
278. 25 | ~4 [avatar split clause 264,136,275]
279. r1_arytm_2(sK1,sK0) <- (~2, 3, 4) [subsumption resolution 262,128] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
281. 26 <=> r1_arytm_2(sK1,sK0) [avatar definition]
283. r1_arytm_2(sK1,sK0) <- (26) [avatar component clause 281] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
284. 26 | 2 | ~3 | ~4 [avatar split clause 279,136,131,126,281]
394. k7_arytm_2(k1_arytm_1(X0,X1),X1) = X0 | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [subsumption resolution 119,110] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
395. sK0 = k7_arytm_2(k11_arytm_3,sK1) | ~r1_arytm_2(sK1,sK0) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (1) [superposition 394,123] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
410. sK0 = k7_arytm_2(k11_arytm_3,sK1) | ~m1_subset_1(sK1,k2_arytm_2) | ~m1_subset_1(sK0,k2_arytm_2) <- (1, 26) [subsumption resolution 395,283] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
411. sK0 = k7_arytm_2(k11_arytm_3,sK1) | ~m1_subset_1(sK0,k2_arytm_2) <- (1, 3, 26) [subsumption resolution 410,133] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
412. sK0 = k7_arytm_2(k11_arytm_3,sK1) <- (1, 3, 4, 26) [subsumption resolution 411,138] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
414. 40 <=> sK0 = k7_arytm_2(k11_arytm_3,sK1) [avatar definition]
416. sK0 = k7_arytm_2(k11_arytm_3,sK1) <- (40) [avatar component clause 414] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
417. 40 | ~1 | ~3 | ~4 | ~26 [avatar split clause 412,281,136,131,121,414]
446. sK0 = sK1 | ~m1_subset_1(sK1,k2_arytm_2) <- (9, 40) [superposition 161,416] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
451. 41 <=> sK0 = sK1 [avatar definition]
455. sK0 = sK1 <- (3, 9, 40) [subsumption resolution 446,133] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
456. 41 | ~3 | ~9 | ~40 [avatar split clause 455,414,160,131,451]
457. sK0 != sK1 | r1_arytm_2(sK0,sK1) | ~r1_arytm_2(sK0,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
458. $false [avatar sat refutation 124,129,134,139,162,243,278,284,417,456,457] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:216,thDist:-216}
% SZS output end Proof for arytm_1__t9_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% >>> Preprocessing
% Introduced skolems: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 211
% Active clauses: 49
% Passive clauses: 89
% Final active clauses: 47
% Final passive clauses: 37
% Main loop iterations started: 52
% 
% >>> Simplifying Inferences
% Duplicate literals: 44
% Fw subsumption resolutions: 28
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 48
% Forward superposition: 22
% Backward superposition: 17
% 
% >>> AVATAR
% Split clauses: 20
% Split components: 45
% Sat splitting refutations: 1
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 7
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

