i: [1,96,0,2,1,1,0,0] t12_parsp_2
a: [3,183,0,2,1,96]
s: 183
i: [1,95,0,2,1,1,0,0] t12_parsp_2
a: [3,188,0,2,1,95]
s: 188
i: [1,94,0,2,1,1,0,0] t12_parsp_2
a: [3,193,0,2,1,94]
s: 193
i: [1,93,0,2,1,1,0,0] t12_parsp_2
a: [3,198,0,2,1,93]
s: 198
i: [1,100,0,9,3,1,0,0] t12_parsp_2
a: [3,159,0,3,1,100]
s: 159
i: [1,99,0,4,1,1,0,0] t12_parsp_2
a: [3,168,0,4,1,99]
s: 168
i: [1,98,0,4,1,1,0,0] t12_parsp_2
a: [3,173,0,4,1,98]
s: 173
i: [1,97,0,4,1,1,0,0] t12_parsp_2
a: [3,178,0,4,1,97]
s: 178
i: [1,101,0,5,1,1,0,0] t12_parsp_2
a: [3,150,0,5,1,101]
s: 150
d: [2,237,0,5,1,2,22,150,159]
a: [3,241,0,5,1,237]
s: 241
i: [1,102,0,2,1,0,0,255] fc1_xboole_0
a: [3,203,0,2,1,102]
s: 203
i: [1,131,0,2,1,0,0,255] rc2_xboole_0
a: [3,208,0,2,1,131]
s: 208
i: [1,132,0,2,1,0,0,2] existence_l1_struct_0
a: [3,213,0,2,1,132]
s: 213
i: [1,133,0,2,1,0,0,255] rc1_xboole_0
a: [3,218,0,2,1,133]
s: 218
i: [1,134,0,2,1,0,0,1] existence_l1_parsp_1
a: [3,223,0,2,1,134]
s: 223
i: [1,136,0,2,1,0,0,1] rc2_parsp_1
a: [3,228,0,2,1,136]
s: 228
i: [1,135,0,2,1,0,0,1] rc2_parsp_1
a: [3,233,0,2,1,135]
s: 233
i: [1,117,0,4,2,0,0,1] dt_l1_parsp_1
s: 117
d: [2,243,1,2,1,1,40,117,183]
a: [3,249,1,2,1,243]
s: 249
i: [1,130,0,4,1,0,0,1] existence_m1_subset_1
s: 130
i: [1,103,0,6,3,0,0,255] cc7_struct_0
s: 103
d: [2,244,1,2,1,1,40,117,223]
a: [3,254,1,2,1,244]
s: 254
i: [1,104,0,6,3,0,0,255] cc5_struct_0
s: 104
d: [2,245,1,2,1,1,40,117,233]
a: [3,259,1,2,1,245]
s: 259
i: [1,106,0,6,3,0,0,255] cc17_struct_0
s: 106
i: [1,107,0,6,3,0,0,255] cc2_struct_0
s: 107
d: [2,261,1,4,2,1,40,104,213]
a: [3,267,1,2,1,261]
s: 267
i: [1,108,0,6,3,0,0,255] cc6_struct_0
s: 108
d: [2,263,2,4,2,1,40,104,254]
a: [3,276,2,2,1,263]
s: 276
i: [1,109,0,6,3,0,0,255] cc1_struct_0
s: 109
i: [1,112,0,6,3,0,0,255] cc4_struct_0
s: 112
i: [1,110,0,7,3,0,0,255] cc8_struct_0
s: 110
i: [1,113,0,7,3,0,0,255] cc11_struct_0
s: 113
i: [1,114,0,7,3,0,0,255] cc11_struct_0
s: 114
i: [1,115,0,7,3,0,0,255] cc9_struct_0
s: 115
i: [1,118,0,7,3,0,0,255] fc9_struct_0
s: 118
i: [1,119,0,7,3,0,0,255] fc6_struct_0
s: 119
i: [1,126,0,7,3,0,0,255] fc2_struct_0
s: 126
i: [1,127,0,7,3,0,0,255] fc8_struct_0
s: 127
i: [1,128,0,7,3,0,0,255] fc7_struct_0
s: 128
i: [1,129,0,7,3,0,0,255] fc1_struct_0
s: 129
i: [1,116,0,9,4,0,0,255] cc10_struct_0
s: 116
i: [1,140,0,28,8,0,0,1] t25_parsp_1
d: [2,145,0,20,6,0,18,140]
s: 145
i: [1,138,0,28,8,0,0,1] t25_parsp_1
d: [2,146,0,20,6,0,18,138]
s: 146
i: [1,142,0,28,8,0,0,1] t25_parsp_1
d: [2,143,0,24,7,0,18,142]
s: 143
i: [1,141,0,28,8,0,0,1] t25_parsp_1
d: [2,144,0,24,7,0,18,141]
s: 144
i: [1,120,0,31,9,0,0,1] d2_parsp_2
s: 120
i: [1,121,0,31,9,0,0,1] d2_parsp_2
s: 121
d: [2,311,1,39,13,0,40,121,146]
d: [2,316,1,17,6,0,18,311]
s: 316
d: [2,313,1,43,14,0,40,121,143]
d: [2,314,1,21,7,0,18,313]
s: 314
d: [2,312,1,43,14,0,40,121,144]
d: [2,315,1,21,7,0,18,312]
s: 315
d: [2,309,1,39,13,0,40,121,145]
d: [2,318,1,21,7,0,18,309]
s: 318
d: [2,319,2,0,0,7,52,198,193,188,183,178,173,241,318]
e: 319
d: [2,320,2,16,6,1,40,318,241]
d: [2,322,2,12,5,2,21,320,173]
d: [2,323,2,8,4,3,21,322,178]
d: [2,324,2,6,3,4,21,323,183]
d: [2,325,2,4,2,5,21,324,188]
d: [2,326,2,2,1,6,21,325,193]
d: [2,327,2,0,0,7,21,326,198]
e: 327
r: 159
r: 241
a: [3,158,0,3,1,100]
s: 158
a: [3,163,0,3,1,100]
s: 163
r: 158
a: [3,154,0,3,1,100]
s: 154
s: 168
s: 150
a: [3,240,0,5,1,237]
s: 240
d: [2,332,0,5,1,2,22,150,163]
a: [3,336,0,5,1,332]
s: 336
d: [2,338,2,0,0,7,52,198,193,188,183,173,178,336,315]
e: 338
d: [2,339,2,16,6,1,40,336,315]
d: [2,341,2,12,5,2,21,339,178]
d: [2,342,2,8,4,3,21,341,173]
d: [2,343,2,6,3,4,21,342,183]
d: [2,344,2,4,2,5,21,343,188]
d: [2,345,2,2,1,6,21,344,193]
d: [2,346,2,0,0,7,21,345,198]
e: 346
r: 154
r: 163
r: 336
a: [3,155,0,3,1,100]
s: 155
r: 240
r: 173
s: 158
a: [3,162,0,3,1,100]
s: 162
s: 168
s: 150
a: [3,335,0,5,1,332]
s: 335
d: [2,348,0,5,1,2,23,240,155]
a: [3,353,0,5,1,348]
s: 353
d: [2,357,0,5,1,2,22,150,155]
a: [3,361,0,5,1,357]
s: 361
d: [2,365,2,0,0,7,52,198,193,188,183,178,168,361,314]
e: 365
d: [2,366,2,16,6,1,40,361,314]
d: [2,368,2,12,5,2,21,366,168]
d: [2,369,2,8,4,3,21,368,178]
d: [2,370,2,6,3,4,21,369,183]
d: [2,371,2,4,2,5,21,370,188]
d: [2,372,2,2,1,6,21,371,193]
d: [2,373,2,0,0,7,21,372,198]
e: 373
f: 373,357,346,332,327,237,93,94,95,96,97,98,99,100,101
% Refutation found. Thanks to Tanya!
% SZS status Theorem for parsp_2__t12_parsp_2
% SZS output start Proof for parsp_2__t12_parsp_2
1. ! [X0] : ((l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (~(X1 != X3 & X2 != X3 & X1 != X2) => r1_parsp_2(X0,X1,X2,X3)))))) [input t12_parsp_2]
2. ~! [X0] : ((l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (~(X1 != X3 & X2 != X3 & X1 != X2) => r1_parsp_2(X0,X1,X2,X3)))))) [negated conjecture 1]
14. ! [X0] : ((l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_parsp_2(X0,X1,X2,X3) <=> r1_parsp_1(X0,X1,X2,X1,X3)))))) [input d2_parsp_2]
33. ! [X0] : ((l1_parsp_1(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (~(~(X2 = X3 & X1 = X4) & ~(X2 = X4 & X1 = X3) & X3 != X4 & X1 != X2) => r1_parsp_1(X0,X1,X2,X3,X4))))))) [input t25_parsp_1]
34. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_parsp_2(X0,X1,X2,X3) & (X1 = X3 | X2 = X3 | X1 = X2)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
35. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_parsp_2(X0,X1,X2,X3) & (X1 = X3 | X2 = X3 | X1 = X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) [flattening 34]
63. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_parsp_2(X0,X1,X2,X3) <=> r1_parsp_1(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0))) [ennf transformation 14]
64. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_parsp_2(X0,X1,X2,X3) <=> r1_parsp_1(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0)) [flattening 63]
65. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r1_parsp_1(X0,X1,X2,X3,X4) | ((X2 != X3 | X1 != X4) & (X2 != X4 | X1 != X3) & X3 != X4 & X1 != X2)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0))) [ennf transformation 33]
66. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r1_parsp_1(X0,X1,X2,X3,X4) | ((X2 != X3 | X1 != X4) & (X2 != X4 | X1 != X3) & X3 != X4 & X1 != X2) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0)) [flattening 65]
75. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_parsp_2(X0,X1,X2,X3) & (X1 = X3 | X2 = X3 | X1 = X2) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_parsp_1(X0) & v1_parsp_2(X0) & v2_parsp_1(X0) & ~v2_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r1_parsp_2(sK0,X1,X2,X3) & (X1 = X3 | X2 = X3 | X1 = X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_parsp_1(sK0) & v1_parsp_2(sK0) & v2_parsp_1(sK0) & ~v2_struct_0(sK0)) [choice axiom]
76. ? [X1] : (? [X2] : (? [X3] : (~r1_parsp_2(sK0,X1,X2,X3) & (X1 = X3 | X2 = X3 | X1 = X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (~r1_parsp_2(sK0,sK1,X2,X3) & (sK1 = X3 | X2 = X3 | sK1 = X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
77. ? [X2] : (? [X3] : (~r1_parsp_2(sK0,sK1,X2,X3) & (sK1 = X3 | X2 = X3 | sK1 = X2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (~r1_parsp_2(sK0,sK1,sK2,X3) & (sK1 = X3 | sK2 = X3 | sK1 = sK2) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
78. ? [X3] : (~r1_parsp_2(sK0,sK1,sK2,X3) & (sK1 = X3 | sK2 = X3 | sK1 = sK2) & m1_subset_1(X3,u1_struct_0(sK0))) => (~r1_parsp_2(sK0,sK1,sK2,sK3) & (sK1 = sK3 | sK2 = sK3 | sK1 = sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
79. (((~r1_parsp_2(sK0,sK1,sK2,sK3) & (sK1 = sK3 | sK2 = sK3 | sK1 = sK2) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_parsp_1(sK0) & v1_parsp_2(sK0) & v2_parsp_1(sK0) & ~v2_struct_0(sK0) [skolemisation 35,78,77,76,75]
80. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_parsp_2(X0,X1,X2,X3) | ~r1_parsp_1(X0,X1,X2,X1,X3)) & (r1_parsp_1(X0,X1,X2,X1,X3) | ~r1_parsp_2(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0)) [nnf transformation 64]
93. ~v2_struct_0(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
94. v2_parsp_1(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
95. v1_parsp_2(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
96. l1_parsp_1(sK0) [cnf transformation 79] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
97. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 79] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
98. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 79] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
99. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 79] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
100. sK1 = sK3 | sK2 = sK3 | sK1 = sK2 [cnf transformation 79] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
101. ~r1_parsp_2(sK0,sK1,sK2,sK3) [cnf transformation 79] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
121. ~r1_parsp_1(X0,X1,X2,X1,X3) | r1_parsp_2(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [cnf transformation 80] {a:0,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
122. r1_parsp_1(X0,X1,X2,X3,X4) | X1 != X2 | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [cnf transformation 66] {a:0,w:31,wCS:12400,sine:1,thAx:0,allAx:1,thDist:-1}
123. r1_parsp_1(X0,X1,X2,X3,X4) | X3 != X4 | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [cnf transformation 66] {a:0,w:31,wCS:12400,sine:1,thAx:0,allAx:1,thDist:-1}
124. r1_parsp_1(X0,X1,X2,X3,X4) | X2 != X4 | X1 != X3 | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [cnf transformation 66] {a:0,w:34,wCS:13600,sine:1,thAx:0,allAx:1,thDist:-1}
139. r1_parsp_1(X0,X1,X4,X3,X4) | X1 != X3 | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [equality resolution 124] {a:0,w:31,wCS:12400,sine:1,thAx:0,allAx:1,thDist:-1}
140. r1_parsp_1(X0,X3,X4,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [equality resolution 139] {a:0,w:28,wCS:11200,sine:1,thAx:0,allAx:1,thDist:-1}
141. r1_parsp_1(X0,X1,X2,X4,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [equality resolution 123] {a:0,w:28,wCS:11200,sine:1,thAx:0,allAx:1,thDist:-1}
142. r1_parsp_1(X0,X2,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [equality resolution 122] {a:0,w:28,wCS:11200,sine:1,thAx:0,allAx:1,thDist:-1}
143. r1_parsp_1(X0,X2,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [duplicate literal removal 142] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
144. r1_parsp_1(X0,X1,X2,X4,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [duplicate literal removal 141] {a:0,w:24,wCS:9600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. r1_parsp_1(X0,X3,X4,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [duplicate literal removal 140] {a:0,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. 1 <=> r1_parsp_2(sK0,sK1,sK2,sK3) [avatar definition]
150. ~r1_parsp_2(sK0,sK1,sK2,sK3) <- (~1) [avatar component clause 148] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
151. ~1 [avatar split clause 101,148]
153. 2 <=> sK1 = sK2 [avatar definition]
155. sK1 = sK2 <- (2) [avatar component clause 153] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
157. 3 <=> sK2 = sK3 [avatar definition]
159. sK2 = sK3 <- (3) [avatar component clause 157] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
161. 4 <=> sK1 = sK3 [avatar definition]
163. sK1 = sK3 <- (4) [avatar component clause 161] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
164. 2 | 3 | 4 [avatar split clause 100,161,157,153]
166. 5 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
168. m1_subset_1(sK3,u1_struct_0(sK0)) <- (5) [avatar component clause 166] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
169. 5 [avatar split clause 99,166]
171. 6 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
173. m1_subset_1(sK2,u1_struct_0(sK0)) <- (6) [avatar component clause 171] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
174. 6 [avatar split clause 98,171]
176. 7 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
178. m1_subset_1(sK1,u1_struct_0(sK0)) <- (7) [avatar component clause 176] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
179. 7 [avatar split clause 97,176]
181. 8 <=> l1_parsp_1(sK0) [avatar definition]
183. l1_parsp_1(sK0) <- (8) [avatar component clause 181] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
184. 8 [avatar split clause 96,181]
186. 9 <=> v1_parsp_2(sK0) [avatar definition]
188. v1_parsp_2(sK0) <- (9) [avatar component clause 186] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
189. 9 [avatar split clause 95,186]
191. 10 <=> v2_parsp_1(sK0) [avatar definition]
193. v2_parsp_1(sK0) <- (10) [avatar component clause 191] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
194. 10 [avatar split clause 94,191]
196. 11 <=> v2_struct_0(sK0) [avatar definition]
198. ~v2_struct_0(sK0) <- (~11) [avatar component clause 196] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
199. ~11 [avatar split clause 93,196]
237. ~r1_parsp_2(sK0,sK1,sK2,sK2) <- (~1, 3) [forward demodulation 150,159] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
239. 19 <=> r1_parsp_2(sK0,sK1,sK2,sK2) [avatar definition]
241. ~r1_parsp_2(sK0,sK1,sK2,sK2) <- (~19) [avatar component clause 239] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
242. ~19 | 1 | ~3 [avatar split clause 237,157,148,239]
309. r1_parsp_2(X0,X1,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [resolution 121,145] {a:1,w:39,wCS:15600,sine:1,thAx:0,allAx:2,thDist:-2}
312. r1_parsp_2(X9,X10,X11,X10) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~l1_parsp_1(X9) | ~v1_parsp_2(X9) | ~v2_parsp_1(X9) | v2_struct_0(X9) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~l1_parsp_1(X9) | ~v2_parsp_1(X9) | v2_struct_0(X9) [resolution 121,144] {a:1,w:43,wCS:17200,sine:1,thAx:0,allAx:2,thDist:-2}
313. r1_parsp_2(X12,X13,X13,X14) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_parsp_1(X12) | ~v1_parsp_2(X12) | ~v2_parsp_1(X12) | v2_struct_0(X12) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_parsp_1(X12) | ~v2_parsp_1(X12) | v2_struct_0(X12) [resolution 121,143] {a:1,w:43,wCS:17200,sine:1,thAx:0,allAx:2,thDist:-2}
314. r1_parsp_2(X12,X13,X13,X14) | ~m1_subset_1(X14,u1_struct_0(X12)) | ~m1_subset_1(X13,u1_struct_0(X12)) | ~l1_parsp_1(X12) | ~v1_parsp_2(X12) | ~v2_parsp_1(X12) | v2_struct_0(X12) [duplicate literal removal 313] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
315. r1_parsp_2(X9,X10,X11,X10) | ~m1_subset_1(X10,u1_struct_0(X9)) | ~m1_subset_1(X11,u1_struct_0(X9)) | ~l1_parsp_1(X9) | ~v1_parsp_2(X9) | ~v2_parsp_1(X9) | v2_struct_0(X9) [duplicate literal removal 312] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
318. r1_parsp_2(X0,X1,X2,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_parsp_1(X0) | ~v1_parsp_2(X0) | ~v2_parsp_1(X0) | v2_struct_0(X0) [duplicate literal removal 309] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
320. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (~19) [resolution 318,241] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
322. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, ~19) [subsumption resolution 320,173] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
323. ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, ~19) [subsumption resolution 322,178] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
324. ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, 8, ~19) [subsumption resolution 323,183] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
325. ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, ~19) [subsumption resolution 324,188] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
326. v2_struct_0(sK0) <- (6, 7, 8, 9, 10, ~19) [subsumption resolution 325,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
327. $false <- (6, 7, 8, 9, 10, ~11, ~19) [subsumption resolution 326,198] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
328. ~6 | ~7 | ~8 | ~9 | ~10 | 11 | 19 [avatar contradiction clause 327]
332. ~r1_parsp_2(sK0,sK1,sK2,sK1) <- (~1, 4) [forward demodulation 150,163] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
334. 29 <=> r1_parsp_2(sK0,sK1,sK2,sK1) [avatar definition]
336. ~r1_parsp_2(sK0,sK1,sK2,sK1) <- (~29) [avatar component clause 334] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
337. ~29 | 1 | ~4 [avatar split clause 332,161,148,334]
339. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (~29) [resolution 336,315] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
341. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (7, ~29) [subsumption resolution 339,178] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
342. ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, ~29) [subsumption resolution 341,173] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
343. ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, 8, ~29) [subsumption resolution 342,183] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
344. ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (6, 7, 8, 9, ~29) [subsumption resolution 343,188] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
345. v2_struct_0(sK0) <- (6, 7, 8, 9, 10, ~29) [subsumption resolution 344,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
346. $false <- (6, 7, 8, 9, 10, ~11, ~29) [subsumption resolution 345,198] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
347. ~6 | ~7 | ~8 | ~9 | ~10 | 11 | 29 [avatar contradiction clause 346]
357. ~r1_parsp_2(sK0,sK1,sK1,sK3) <- (~1, 2) [forward demodulation 150,155] {a:0,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
359. 31 <=> r1_parsp_2(sK0,sK1,sK1,sK3) [avatar definition]
361. ~r1_parsp_2(sK0,sK1,sK1,sK3) <- (~31) [avatar component clause 359] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
362. ~31 | 1 | ~2 [avatar split clause 357,153,148,359]
366. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (~31) [resolution 361,314] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
368. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (5, ~31) [subsumption resolution 366,168] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
369. ~l1_parsp_1(sK0) | ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (5, 7, ~31) [subsumption resolution 368,178] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
370. ~v1_parsp_2(sK0) | ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (5, 7, 8, ~31) [subsumption resolution 369,183] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
371. ~v2_parsp_1(sK0) | v2_struct_0(sK0) <- (5, 7, 8, 9, ~31) [subsumption resolution 370,188] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
372. v2_struct_0(sK0) <- (5, 7, 8, 9, 10, ~31) [subsumption resolution 371,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
373. $false <- (5, 7, 8, 9, 10, ~11, ~31) [subsumption resolution 372,198] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
374. ~5 | ~7 | ~8 | ~9 | ~10 | 11 | 31 [avatar contradiction clause 373]
375. $false [avatar sat refutation 151,164,169,174,179,184,189,194,199,242,328,337,347,362,374] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:316,thDist:-316}
% SZS output end Proof for parsp_2__t12_parsp_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 32
% Introduced skolems: 10
% 
% >>> Saturation
% Initial clauses: 44
% Generated clauses: 152
% Active clauses: 56
% Passive clauses: 67
% Final active clauses: 48
% Main loop iterations started: 67
% 
% >>> Simplifying Inferences
% Duplicate literals: 46
% Fw subsumption resolutions: 18
% Fw demodulations: 7
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 25
% Unit resulting resolution: 3
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 11
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 39
% SAT solver unit clauses: 15
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% binary resolution index maintenance: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

