i: [1,37,0,2,1,1,0,0] t24_xxreal_0
a: [3,70,0,2,1,37]
s: 70
i: [1,36,0,2,1,1,0,0] t24_xxreal_0
a: [3,75,0,2,1,36]
s: 75
i: [1,35,0,2,1,1,0,0] t24_xxreal_0
a: [3,80,0,2,1,35]
s: 80
i: [1,39,0,3,1,1,0,0] t24_xxreal_0
a: [3,60,0,3,1,39]
s: 60
i: [1,38,0,3,1,1,0,0] t24_xxreal_0
a: [3,65,0,3,1,38]
s: 65
i: [1,41,0,5,1,1,0,0] t24_xxreal_0
a: [3,55,0,5,1,41]
s: 55
i: [1,51,0,2,1,0,0,1] rc1_xxreal_0
a: [3,96,0,2,1,51]
s: 96
i: [1,40,0,11,4,1,0,0] t24_xxreal_0
s: 40
d: [2,98,1,8,3,1,40,40,60]
d: [2,99,1,5,2,2,21,98,65]
d: [2,100,1,3,1,3,21,99,80]
a: [3,104,1,3,1,100]
s: 104
i: [1,46,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,87,0,5,2,46]
s: 87
i: [1,48,0,9,3,0,0,1] idempotence_k3_xxreal_0
a: [3,91,0,7,2,48]
s: 91
d: [2,106,1,10,4,1,40,87,40]
d: [2,107,1,8,3,1,18,106]
d: [2,108,1,6,2,2,21,107,70]
a: [3,112,1,3,1,108]
s: 112
i: [1,47,0,8,3,0,0,1] fc5_xxreal_0
s: 47
i: [1,42,0,9,3,0,0,1] t17_xxreal_0
s: 42
i: [1,50,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 50
r: 112
i: [1,49,0,11,3,0,0,1] commutativity_k3_xxreal_0
s: 49
a: [3,111,1,3,1,108]
s: 111
i: [1,43,0,12,4,0,0,1] d9_xxreal_0
s: 43
a: [3,116,1,3,1,108]
s: 116
i: [1,44,0,12,4,0,0,1] d9_xxreal_0
s: 44
i: [1,45,0,13,5,0,0,1] t1_xxreal_0
s: 45
r: 87
r: 116
r: 111
s: 112
a: [3,84,0,2,1,46]
s: 84
r: 91
r: 50
r: 49
r: 47
r: 45
r: 44
r: 43
r: 42
r: 40
d: [2,237,1,0,0,2,40,84,96]
e: 237
d: [2,236,1,0,0,2,40,84,70]
e: 236
d: [2,235,1,0,0,2,40,84,75]
e: 235
d: [2,234,1,0,0,2,40,84,80]
e: 234
d: [2,233,1,0,0,2,52,80,84]
e: 233
r: 84
a: [3,115,1,3,1,108]
s: 115
s: 40
s: 87
d: [2,147,1,10,4,1,44,55,43]
d: [2,158,1,8,3,2,21,147,70]
d: [2,159,1,6,2,3,21,158,75]
a: [3,163,1,3,1,159]
s: 163
s: 91
d: [2,168,1,10,4,1,44,55,44]
d: [2,182,1,8,3,2,21,168,70]
d: [2,183,1,6,2,3,21,182,75]
a: [3,187,1,3,1,183]
s: 187
s: 47
s: 42
d: [2,193,1,10,4,1,40,45,65]
d: [2,206,1,8,3,2,21,193,80]
d: [2,207,1,6,2,3,21,206,75]
a: [3,211,1,3,1,207]
s: 211
s: 50
d: [2,125,2,7,3,1,40,50,112]
d: [2,130,2,5,2,2,21,125,70]
d: [2,131,2,3,1,3,21,130,75]
a: [3,135,2,3,1,131]
s: 135
s: 49
s: 43
d: [2,263,3,8,3,1,40,135,40]
d: [2,264,3,5,2,2,21,263,211]
d: [2,265,3,3,1,3,21,264,75]
a: [3,269,3,3,1,265]
s: 269
d: [2,292,4,0,0,3,52,75,269,87]
e: 292
d: [2,293,4,0,0,2,52,75,75,269,269,50]
e: 293
d: [2,294,4,0,0,2,52,75,75,269,269,50]
e: 294
d: [2,295,4,2,1,2,40,269,87]
d: [2,301,4,0,0,3,21,295,75]
e: 301
d: [2,126,1,7,3,0,42,50]
d: [2,127,1,5,2,0,18,126]
d: [2,192,1,10,4,1,40,45,60]
d: [2,226,1,8,3,2,21,192,80]
d: [2,227,1,6,2,3,21,226,70]
d: [2,252,2,7,3,1,40,50,112]
d: [2,259,2,5,2,2,21,252,70]
d: [2,260,2,3,1,3,21,259,75]
f: 301,265,260,227,127,207,183,159,108,35,36,37,38,39,41
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t24_xxreal_0
% SZS output start Proof for xxreal_0__t24_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((! [X3] : (v1_xxreal_0(X3) => ((r1_xxreal_0(X3,X2) & r1_xxreal_0(X3,X1)) => r1_xxreal_0(X3,X0))) & r1_xxreal_0(X0,X2) & r1_xxreal_0(X0,X1)) => k3_xxreal_0(X1,X2) = X0)))) [input t24_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((! [X3] : (v1_xxreal_0(X3) => ((r1_xxreal_0(X3,X2) & r1_xxreal_0(X3,X1)) => r1_xxreal_0(X3,X0))) & r1_xxreal_0(X0,X2) & r1_xxreal_0(X0,X1)) => k3_xxreal_0(X1,X2) = X0)))) [negated conjecture 1]
4. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
5. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((~r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X1) & (r1_xxreal_0(X0,X1) => k3_xxreal_0(X0,X1) = X0)))) [input d9_xxreal_0]
12. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
13. ? [X0] : (? [X1] : (? [X2] : ((k3_xxreal_0(X1,X2) != X0 & (! [X3] : ((r1_xxreal_0(X3,X0) | (~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,X1))) | ~v1_xxreal_0(X3)) & r1_xxreal_0(X0,X2) & r1_xxreal_0(X0,X1))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
14. ? [X0] : (? [X1] : (? [X2] : (k3_xxreal_0(X1,X2) != X0 & ! [X3] : (r1_xxreal_0(X3,X0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,X1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(X0,X2) & r1_xxreal_0(X0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 13]
16. ! [X0] : (! [X1] : (((k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1)) & (k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 5]
17. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 12]
18. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 17]
27. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 4]
28. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 27]
29. ? [X0] : (? [X1] : (? [X2] : (k3_xxreal_0(X1,X2) != X0 & ! [X3] : (r1_xxreal_0(X3,X0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,X1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(X0,X2) & r1_xxreal_0(X0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (k3_xxreal_0(X1,X2) != sK0 & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,X1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,X2) & r1_xxreal_0(sK0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
30. ? [X1] : (? [X2] : (k3_xxreal_0(X1,X2) != sK0 & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,X1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,X2) & r1_xxreal_0(sK0,X1) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (sK0 != k3_xxreal_0(sK1,X2) & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,sK1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,X2) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
31. ? [X2] : (sK0 != k3_xxreal_0(sK1,X2) & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,X2) | ~r1_xxreal_0(X3,sK1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,X2) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(X2)) => (sK0 != k3_xxreal_0(sK1,sK2) & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,sK2) | ~r1_xxreal_0(X3,sK1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,sK2) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(sK2)) [choice axiom]
32. ((sK0 != k3_xxreal_0(sK1,sK2) & ! [X3] : (r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,sK2) | ~r1_xxreal_0(X3,sK1) | ~v1_xxreal_0(X3)) & r1_xxreal_0(sK0,sK2) & r1_xxreal_0(sK0,sK1) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 14,31,30,29]
35. v1_xxreal_0(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
36. v1_xxreal_0(sK1) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
37. v1_xxreal_0(sK2) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
38. r1_xxreal_0(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
39. r1_xxreal_0(sK0,sK2) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
40. ~r1_xxreal_0(X3,sK2) | r1_xxreal_0(X3,sK0) | ~r1_xxreal_0(X3,sK1) | ~v1_xxreal_0(X3) [cnf transformation 32] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. sK0 != k3_xxreal_0(sK1,sK2) [cnf transformation 32] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. k3_xxreal_0(X0,X1) = X0 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 16] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. k3_xxreal_0(X0,X1) = X1 | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 16] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 18] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 28] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. 1 <=> sK0 = k3_xxreal_0(sK1,sK2) [avatar definition]
55. sK0 != k3_xxreal_0(sK1,sK2) <- (~1) [avatar component clause 53] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
56. ~1 [avatar split clause 41,53]
58. 2 <=> r1_xxreal_0(sK0,sK2) [avatar definition]
60. r1_xxreal_0(sK0,sK2) <- (2) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
61. 2 [avatar split clause 39,58]
63. 3 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
65. r1_xxreal_0(sK0,sK1) <- (3) [avatar component clause 63] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66. 3 [avatar split clause 38,63]
68. 4 <=> v1_xxreal_0(sK2) [avatar definition]
70. v1_xxreal_0(sK2) <- (4) [avatar component clause 68] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
71. 4 [avatar split clause 37,68]
73. 5 <=> v1_xxreal_0(sK1) [avatar definition]
75. v1_xxreal_0(sK1) <- (5) [avatar component clause 73] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
76. 5 [avatar split clause 36,73]
78. 6 <=> v1_xxreal_0(sK0) [avatar definition]
80. v1_xxreal_0(sK0) <- (6) [avatar component clause 78] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 6 [avatar split clause 35,78]
86. 8 <=> ! [X0] : (r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0)) [avatar definition]
87. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) <- (8) [avatar component clause 86] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
106. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK2,sK1) | ~v1_xxreal_0(sK2) <- (8) [resolution 87,40] {a:1,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
107. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK2,sK1) <- (8) [duplicate literal removal 106] {a:1,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
108. r1_xxreal_0(sK2,sK0) | ~r1_xxreal_0(sK2,sK1) <- (4, 8) [subsumption resolution 107,70] {a:1,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
110. 12 <=> r1_xxreal_0(sK2,sK1) [avatar definition]
112. ~r1_xxreal_0(sK2,sK1) <- (~12) [avatar component clause 110] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
114. 13 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
117. ~12 | 13 | ~4 | ~8 [avatar split clause 108,86,68,114,110]
126. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(X0) [factoring 50] {a:1,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
127. r1_xxreal_0(X0,X0) | ~v1_xxreal_0(X0) [duplicate literal removal 126] {a:1,w:5,wCS:2000,sine:1,thAx:0,allAx:1,thDist:-1}
133. 14 <=> r1_xxreal_0(sK1,sK2) [avatar definition]
135. r1_xxreal_0(sK1,sK2) <- (14) [avatar component clause 133] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
147. sK0 != sK1 | ~r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~1) [superposition 55,43] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. sK0 != sK1 | ~r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) <- (~1, 4) [subsumption resolution 147,70] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
159. sK0 != sK1 | ~r1_xxreal_0(sK1,sK2) <- (~1, 4, 5) [subsumption resolution 158,75] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
161. 15 <=> sK0 = sK1 [avatar definition]
164. ~14 | ~15 | 1 | ~4 | ~5 [avatar split clause 159,73,68,53,161,133]
168. sK0 != sK2 | r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~1) [superposition 55,44] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. sK0 != sK2 | r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) <- (~1, 4) [subsumption resolution 168,70] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
183. sK0 != sK2 | r1_xxreal_0(sK1,sK2) <- (~1, 4, 5) [subsumption resolution 182,75] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
185. 16 <=> sK0 = sK2 [avatar definition]
188. 14 | ~16 | 1 | ~4 | ~5 [avatar split clause 183,73,68,53,185,133]
192. sK0 = sK2 | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK2) <- (2) [resolution 45,60] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
193. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (3) [resolution 45,65] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
206. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) <- (3, 6) [subsumption resolution 193,80] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
207. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) <- (3, 5, 6) [subsumption resolution 206,75] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
209. 17 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
211. ~r1_xxreal_0(sK1,sK0) <- (~17) [avatar component clause 209] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
212. ~17 | 15 | ~3 | ~5 | ~6 [avatar split clause 207,78,73,63,161,209]
224. 8 [avatar split clause 127,86]
226. sK0 = sK2 | ~r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) <- (2, 6) [subsumption resolution 192,80] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
227. sK0 = sK2 | ~r1_xxreal_0(sK2,sK0) <- (2, 4, 6) [subsumption resolution 226,70] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
228. ~13 | 16 | ~2 | ~4 | ~6 [avatar split clause 227,78,68,58,185,114]
252. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK1) <- (~12) [resolution 50,112] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
259. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK1) <- (4, ~12) [subsumption resolution 252,70] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
260. r1_xxreal_0(sK1,sK2) <- (4, 5, ~12) [subsumption resolution 259,75] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
261. 14 | ~4 | ~5 | 12 [avatar split clause 260,110,73,68,133]
263. r1_xxreal_0(sK1,sK0) | ~r1_xxreal_0(sK1,sK1) | ~v1_xxreal_0(sK1) <- (14) [resolution 135,40] {a:3,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
264. ~r1_xxreal_0(sK1,sK1) | ~v1_xxreal_0(sK1) <- (14, ~17) [subsumption resolution 263,211] {a:3,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
265. ~r1_xxreal_0(sK1,sK1) <- (5, 14, ~17) [subsumption resolution 264,75] {a:3,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
267. 19 <=> r1_xxreal_0(sK1,sK1) [avatar definition]
269. ~r1_xxreal_0(sK1,sK1) <- (~19) [avatar component clause 267] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
270. ~19 | ~5 | ~14 | 17 [avatar split clause 265,209,133,73,267]
295. ~v1_xxreal_0(sK1) <- (8, ~19) [resolution 269,87] {a:4,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
301. $false <- (5, 8, ~19) [subsumption resolution 295,75] {a:4,w:0,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
302. ~5 | ~8 | 19 [avatar contradiction clause 301]
303. $false [avatar sat refutation 56,61,66,71,76,81,117,164,188,212,224,228,261,270,302] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:289,thDist:-289}
% SZS output end Proof for xxreal_0__t24_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 17
% Generated clauses: 198
% Active clauses: 37
% Passive clauses: 62
% Final active clauses: 23
% Final passive clauses: 24
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 77
% Fw subsumption resolutions: 33
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Equational tautologies: 2
% Forward subsumptions: 17
% Backward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 4
% Factoring: 2
% Forward superposition: 11
% Backward superposition: 24
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 16
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 35
% SAT solver unit clauses: 8
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

