i: [1,43,0,2,1,1,0,0] t36_projpl_1
a: [3,101,0,2,1,43]
s: 101
i: [1,42,0,2,1,1,0,0] t36_projpl_1
a: [3,106,0,2,1,42]
s: 106
i: [1,41,0,2,1,1,0,0] t36_projpl_1
a: [3,111,0,2,1,41]
s: 111
i: [1,40,0,2,1,1,0,0] t36_projpl_1
a: [3,116,0,2,1,40]
s: 116
i: [1,39,0,2,1,1,0,0] t36_projpl_1
a: [3,121,0,2,1,39]
s: 121
i: [1,38,0,2,1,1,0,0] t36_projpl_1
a: [3,126,0,2,1,38]
s: 126
i: [1,37,0,2,1,1,0,0] t36_projpl_1
a: [3,131,0,2,1,37]
s: 131
i: [1,49,0,3,1,1,0,0] t36_projpl_1
a: [3,71,0,3,1,49]
s: 71
i: [1,48,0,3,1,1,0,0] t36_projpl_1
a: [3,76,0,3,1,48]
s: 76
i: [1,46,0,4,1,1,0,0] t36_projpl_1
a: [3,86,0,4,1,46]
s: 86
i: [1,45,0,4,1,1,0,0] t36_projpl_1
a: [3,91,0,4,1,45]
s: 91
i: [1,44,0,4,1,1,0,0] t36_projpl_1
a: [3,96,0,4,1,44]
s: 96
i: [1,50,0,7,1,1,0,0] t36_projpl_1
a: [3,66,0,7,1,50]
s: 66
i: [1,47,0,7,1,1,0,0] t36_projpl_1
a: [3,81,0,7,1,47]
s: 81
i: [1,60,0,2,1,0,0,3] rc2_xboole_0
a: [3,136,0,2,1,60]
s: 136
i: [1,61,0,2,1,0,0,3] rc1_xboole_0
a: [3,141,0,2,1,61]
s: 141
i: [1,62,0,2,1,0,0,1] existence_l1_incsp_1
a: [3,146,0,2,1,62]
s: 146
i: [1,58,0,4,1,0,0,1] existence_m1_subset_1
s: 58
i: [1,51,0,5,2,0,0,2] dt_u2_incsp_1
s: 51
i: [1,52,0,5,2,0,0,3] dt_u1_incsp_1
s: 52
i: [1,59,0,27,9,0,0,2] dt_k1_projpl_1
s: 59
i: [1,54,0,34,11,0,0,1] t16_projpl_1
s: 54
i: [1,55,0,34,11,0,0,1] t16_projpl_1
s: 55
i: [1,56,0,36,11,0,0,1] t16_projpl_1
s: 56
i: [1,57,0,41,13,0,0,1] t16_projpl_1
s: 57
i: [1,53,0,46,14,0,0,1] t35_projpl_1
s: 53
d: [2,148,1,34,11,1,44,66,56]
d: [2,162,1,31,10,2,21,148,76]
d: [2,163,1,27,9,3,21,162,91]
d: [2,164,1,23,8,4,21,163,96]
d: [2,165,1,21,7,5,21,164,101]
d: [2,166,1,19,6,6,21,165,111]
d: [2,167,1,17,5,7,21,166,116]
d: [2,168,1,15,4,8,21,167,121]
d: [2,169,1,13,3,9,21,168,126]
d: [2,170,1,11,2,10,21,169,131]
a: [3,177,1,7,1,170]
s: 177
d: [2,225,1,39,13,1,40,53,81]
d: [2,236,1,36,12,2,21,225,71]
d: [2,237,1,32,11,3,21,236,86]
d: [2,238,1,28,10,4,21,237,91]
d: [2,239,1,24,9,5,21,238,96]
d: [2,240,1,22,8,6,21,239,101]
d: [2,241,1,20,7,7,21,240,106]
d: [2,242,1,18,6,8,21,241,111]
d: [2,243,1,16,5,9,21,242,116]
d: [2,244,1,14,4,10,21,243,121]
d: [2,245,1,12,3,11,21,244,126]
d: [2,246,1,10,2,12,21,245,131]
a: [3,250,1,7,1,246]
s: 250
r: 177
a: [3,253,1,3,1,246]
s: 253
a: [3,173,1,4,1,170]
s: 173
d: [2,335,2,0,0,1,40,173,58]
e: 335
d: [2,329,2,0,0,1,52,58,173]
e: 329
d: [2,330,2,0,0,8,52,131,126,121,116,111,101,86,86,173,59]
e: 330
d: [2,331,2,0,0,9,52,131,126,121,116,111,101,86,91,173,59]
e: 331
d: [2,332,2,0,0,9,52,131,126,121,116,111,101,86,96,173,59]
e: 332
d: [2,333,2,0,0,8,52,131,126,121,116,111,101,86,58,173,59]
e: 333
d: [2,152,1,34,11,1,44,66,56]
d: [2,179,1,31,10,2,21,152,76]
d: [2,180,1,27,9,3,21,179,96]
d: [2,181,1,23,8,4,21,180,91]
d: [2,182,1,21,7,5,21,181,101]
d: [2,183,1,19,6,6,21,182,111]
d: [2,184,1,17,5,7,21,183,116]
d: [2,185,1,15,4,8,21,184,121]
d: [2,186,1,13,3,9,21,185,126]
d: [2,187,1,11,2,10,21,186,131]
d: [2,266,2,39,13,1,40,250,53]
d: [2,270,2,36,12,2,21,266,76]
d: [2,318,2,32,11,3,21,270,96]
d: [2,319,2,28,10,4,21,318,86]
d: [2,320,2,24,9,5,21,319,91]
d: [2,321,2,22,8,6,21,320,101]
d: [2,322,2,20,7,7,21,321,106]
d: [2,323,2,18,6,8,21,322,111]
d: [2,324,2,16,5,9,21,323,116]
d: [2,325,2,14,4,10,21,324,121]
d: [2,326,2,12,3,11,21,325,126]
d: [2,327,2,10,2,12,21,326,131]
f: 329,327,246,187,37,38,39,40,41,42,43,44,45,46,47,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for projpl_1__t36_projpl_1
% SZS output start Proof for projpl_1__t36_projpl_1
1. ! [X0] : ((l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ~(~r1_incsp_1(X0,X3,k1_projpl_1(X0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3))))))) [input t36_projpl_1]
2. ~! [X0] : ((l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ~(~r1_incsp_1(X0,X3,k1_projpl_1(X0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3))))))) [negated conjecture 1]
9. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
12. ! [X0] : ((l1_incsp_1(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u2_incsp_1(X0)) => (X1 != X2 => (((r1_incsp_1(X0,X2,X3) & r1_incsp_1(X0,X1,X3)) => k1_projpl_1(X0,X1,X2) = X3) & k1_projpl_1(X0,X2,X1) = k1_projpl_1(X0,X1,X2) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2)))))))) [input t16_projpl_1]
13. ! [X0] : ((l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => ! [X1] : (m1_subset_1(X1,u1_incsp_1(X0)) => ! [X2] : (m1_subset_1(X2,u1_incsp_1(X0)) => ! [X3] : (m1_subset_1(X3,u1_incsp_1(X0)) => ~(~r1_incsp_1(X0,X2,k1_projpl_1(X0,X3,X1)) & X2 != X3 & X1 != X3 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3))))))) [input t35_projpl_1]
14. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r1_incsp_1(X0,X3,k1_projpl_1(X0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3))) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & (l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0))) [ennf transformation 2]
15. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(X0,X3,k1_projpl_1(X0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) [flattening 14]
18. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_incsp_1(X0,X2,k1_projpl_1(X0,X3,X1)) | X2 = X3 | X1 = X3 | ~r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3))) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | (~l1_incsp_1(X0) | ~v5_incproj(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0))) [ennf transformation 13]
19. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (r1_incsp_1(X0,X2,k1_projpl_1(X0,X3,X1)) | X2 = X3 | X1 = X3 | ~r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3)) | ~m1_subset_1(X3,u1_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0) | ~v5_incproj(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0)) [flattening 18]
20. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((((k1_projpl_1(X0,X1,X2) = X3 | (~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3))) & k1_projpl_1(X0,X2,X1) = k1_projpl_1(X0,X1,X2) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2))) | X1 = X2) | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | (~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0))) [ennf transformation 12]
21. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k1_projpl_1(X0,X1,X2) = X3 | ~r1_incsp_1(X0,X2,X3) | ~r1_incsp_1(X0,X1,X3)) & k1_projpl_1(X0,X2,X1) = k1_projpl_1(X0,X1,X2) & r1_incsp_1(X0,X2,k1_projpl_1(X0,X1,X2)) & r1_incsp_1(X0,X1,k1_projpl_1(X0,X1,X2))) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0))) | ~m1_subset_1(X2,u1_incsp_1(X0))) | ~m1_subset_1(X1,u1_incsp_1(X0))) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0)) [flattening 20]
24. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(X0,X3,k1_projpl_1(X0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(X0))) & m1_subset_1(X2,u1_incsp_1(X0))) & m1_subset_1(X1,u1_incsp_1(X0))) & l1_incsp_1(X0) & v5_incproj(X0) & v4_incproj(X0) & v3_incproj(X0) & v2_incproj(X0) & v1_incproj(X0) & v6_incsp_1(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(sK0,X1,k1_projpl_1(sK0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v5_incproj(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0)) [choice axiom]
25. ? [X1] : (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X2,X1)) & X2 != X3 & X1 != X2 & r1_incsp_1(sK0,X1,k1_projpl_1(sK0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(X1,u1_incsp_1(sK0))) => (? [X2] : (? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X2,sK1)) & X2 != X3 & sK1 != X2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) [choice axiom]
26. ? [X2] : (? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,X2,sK1)) & X2 != X3 & sK1 != X2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,X2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(X2,u1_incsp_1(sK0))) => (? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK2,sK1)) & sK2 != X3 & sK1 != sK2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) [choice axiom]
27. ? [X3] : (~r1_incsp_1(sK0,X3,k1_projpl_1(sK0,sK2,sK1)) & sK2 != X3 & sK1 != sK2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,X3)) & m1_subset_1(X3,u1_incsp_1(sK0))) => (~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) & sK2 != sK3 & sK1 != sK2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) & m1_subset_1(sK3,u1_incsp_1(sK0))) [choice axiom]
28. (((~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) & sK2 != sK3 & sK1 != sK2 & r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) & m1_subset_1(sK3,u1_incsp_1(sK0))) & m1_subset_1(sK2,u1_incsp_1(sK0))) & m1_subset_1(sK1,u1_incsp_1(sK0))) & l1_incsp_1(sK0) & v5_incproj(sK0) & v4_incproj(sK0) & v3_incproj(sK0) & v2_incproj(sK0) & v1_incproj(sK0) & v6_incsp_1(sK0) [skolemisation 15,27,26,25,24]
29. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK4(X0),X0)) [choice axiom]
30. ! [X0] : m1_subset_1(sK4(X0),X0) [skolemisation 9,29]
37. v6_incsp_1(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
38. v1_incproj(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
39. v2_incproj(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
40. v3_incproj(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
41. v4_incproj(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
42. v5_incproj(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
43. l1_incsp_1(sK0) [cnf transformation 28] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
44. m1_subset_1(sK1,u1_incsp_1(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
45. m1_subset_1(sK2,u1_incsp_1(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
46. m1_subset_1(sK3,u1_incsp_1(sK0)) [cnf transformation 28] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
47. r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) [cnf transformation 28] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
48. sK1 != sK2 [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
49. sK2 != sK3 [cnf transformation 28] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
50. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) [cnf transformation 28] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
53. ~r1_incsp_1(X0,X1,k1_projpl_1(X0,X2,X3)) | X2 = X3 | X1 = X3 | r1_incsp_1(X0,X2,k1_projpl_1(X0,X3,X1)) | ~m1_subset_1(X3,u1_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v5_incproj(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 19] {a:0,w:46,wCS:18400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. k1_projpl_1(X0,X2,X1) = k1_projpl_1(X0,X1,X2) | X1 = X2 | ~m1_subset_1(X3,u2_incsp_1(X0)) | ~m1_subset_1(X2,u1_incsp_1(X0)) | ~m1_subset_1(X1,u1_incsp_1(X0)) | ~l1_incsp_1(X0) | ~v4_incproj(X0) | ~v3_incproj(X0) | ~v2_incproj(X0) | ~v1_incproj(X0) | ~v6_incsp_1(X0) [cnf transformation 21] {a:0,w:36,wCS:14400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. m1_subset_1(sK4(X0),X0) [cnf transformation 30] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
64. 1 <=> r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) [avatar definition]
66. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) <- (~1) [avatar component clause 64] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
67. ~1 [avatar split clause 50,64]
69. 2 <=> sK2 = sK3 [avatar definition]
71. sK2 != sK3 <- (~2) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
72. ~2 [avatar split clause 49,69]
74. 3 <=> sK1 = sK2 [avatar definition]
76. sK1 != sK2 <- (~3) [avatar component clause 74] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
77. ~3 [avatar split clause 48,74]
79. 4 <=> r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) [avatar definition]
81. r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) <- (4) [avatar component clause 79] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
82. 4 [avatar split clause 47,79]
84. 5 <=> m1_subset_1(sK3,u1_incsp_1(sK0)) [avatar definition]
86. m1_subset_1(sK3,u1_incsp_1(sK0)) <- (5) [avatar component clause 84] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
87. 5 [avatar split clause 46,84]
89. 6 <=> m1_subset_1(sK2,u1_incsp_1(sK0)) [avatar definition]
91. m1_subset_1(sK2,u1_incsp_1(sK0)) <- (6) [avatar component clause 89] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
92. 6 [avatar split clause 45,89]
94. 7 <=> m1_subset_1(sK1,u1_incsp_1(sK0)) [avatar definition]
96. m1_subset_1(sK1,u1_incsp_1(sK0)) <- (7) [avatar component clause 94] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
97. 7 [avatar split clause 44,94]
99. 8 <=> l1_incsp_1(sK0) [avatar definition]
101. l1_incsp_1(sK0) <- (8) [avatar component clause 99] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
102. 8 [avatar split clause 43,99]
104. 9 <=> v5_incproj(sK0) [avatar definition]
106. v5_incproj(sK0) <- (9) [avatar component clause 104] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
107. 9 [avatar split clause 42,104]
109. 10 <=> v4_incproj(sK0) [avatar definition]
111. v4_incproj(sK0) <- (10) [avatar component clause 109] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
112. 10 [avatar split clause 41,109]
114. 11 <=> v3_incproj(sK0) [avatar definition]
116. v3_incproj(sK0) <- (11) [avatar component clause 114] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
117. 11 [avatar split clause 40,114]
119. 12 <=> v2_incproj(sK0) [avatar definition]
121. v2_incproj(sK0) <- (12) [avatar component clause 119] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
122. 12 [avatar split clause 39,119]
124. 13 <=> v1_incproj(sK0) [avatar definition]
126. v1_incproj(sK0) <- (13) [avatar component clause 124] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
127. 13 [avatar split clause 38,124]
129. 14 <=> v6_incsp_1(sK0) [avatar definition]
131. v6_incsp_1(sK0) <- (14) [avatar component clause 129] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
132. 14 [avatar split clause 37,129]
152. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | sK1 = sK2 | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1) [superposition 66,56] {a:1,w:34,wCS:3400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
172. 18 <=> ! [X0] : ~m1_subset_1(X0,u2_incsp_1(sK0)) [avatar definition]
173. ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (18) [avatar component clause 172] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
175. 19 <=> r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) [avatar definition]
179. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3) [subsumption resolution 152,76] {a:1,w:31,wCS:3100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
180. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 7) [subsumption resolution 179,96] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
181. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7) [subsumption resolution 180,91] {a:1,w:23,wCS:2300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
182. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7, 8) [subsumption resolution 181,101] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
183. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7, 8, 10) [subsumption resolution 182,111] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
184. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7, 8, 10, 11) [subsumption resolution 183,116] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
185. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7, 8, 10, 11, 12) [subsumption resolution 184,121] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
186. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) | ~v6_incsp_1(sK0) <- (~1, ~3, 6, 7, 8, 10, 11, 12, 13) [subsumption resolution 185,126] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
187. ~r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(X0,u2_incsp_1(sK0)) <- (~1, ~3, 6, 7, 8, 10, 11, 12, 13, 14) [subsumption resolution 186,131] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
188. 18 | ~19 | 1 | 3 | ~6 | ~7 | ~8 | ~10 | ~11 | ~12 | ~13 | ~14 [avatar split clause 187,129,124,119,114,109,99,94,89,74,64,175,172]
225. sK2 = sK3 | sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (4) [resolution 53,81] {a:1,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
236. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4) [subsumption resolution 225,71] {a:1,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
237. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5) [subsumption resolution 236,86] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
238. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6) [subsumption resolution 237,91] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
239. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7) [subsumption resolution 238,96] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
240. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8) [subsumption resolution 239,101] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
241. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8, 9) [subsumption resolution 240,106] {a:1,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
242. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8, 9, 10) [subsumption resolution 241,111] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
243. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8, 9, 10, 11) [subsumption resolution 242,116] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
244. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8, 9, 10, 11, 12) [subsumption resolution 243,121] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
245. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) | ~v6_incsp_1(sK0) <- (~2, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13) [subsumption resolution 244,126] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
246. sK1 = sK3 | r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) <- (~2, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14) [subsumption resolution 245,131] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
248. 20 <=> r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) [avatar definition]
250. r1_incsp_1(sK0,sK2,k1_projpl_1(sK0,sK3,sK1)) <- (20) [avatar component clause 248] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
252. 21 <=> sK1 = sK3 [avatar definition]
255. 20 | 21 | 2 | ~4 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 [avatar split clause 246,129,124,119,114,109,104,99,94,89,84,79,69,252,248]
266. sK1 = sK3 | sK1 = sK2 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (20) [resolution 250,53] {a:2,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
270. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(sK1,u1_incsp_1(sK0)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 20) [subsumption resolution 266,76] {a:2,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
317. sK1 != sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK2,sK1)) | ~r1_incsp_1(sK0,sK1,k1_projpl_1(sK0,sK2,sK3)) [theory tautology sat conflict] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:1,thDist:-1}
318. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(sK3,u1_incsp_1(sK0)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 7, 20) [subsumption resolution 270,96] {a:2,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
319. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~m1_subset_1(sK2,u1_incsp_1(sK0)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 7, 20) [subsumption resolution 318,86] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
320. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~l1_incsp_1(sK0) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 20) [subsumption resolution 319,91] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
321. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v5_incproj(sK0) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 20) [subsumption resolution 320,101] {a:2,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
322. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v4_incproj(sK0) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 9, 20) [subsumption resolution 321,106] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
323. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v3_incproj(sK0) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 9, 10, 20) [subsumption resolution 322,111] {a:2,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
324. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v2_incproj(sK0) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 9, 10, 11, 20) [subsumption resolution 323,116] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
325. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v1_incproj(sK0) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 9, 10, 11, 12, 20) [subsumption resolution 324,121] {a:2,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
326. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) | ~v6_incsp_1(sK0) <- (~3, 5, 6, 7, 8, 9, 10, 11, 12, 13, 20) [subsumption resolution 325,126] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
327. sK1 = sK3 | r1_incsp_1(sK0,sK3,k1_projpl_1(sK0,sK1,sK2)) <- (~3, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 20) [subsumption resolution 326,131] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
328. 19 | 21 | 3 | ~5 | ~6 | ~7 | ~8 | ~9 | ~10 | ~11 | ~12 | ~13 | ~14 | ~20 [avatar split clause 327,248,129,124,119,114,109,104,99,94,89,84,74,252,175]
329. $false <- (18) [unit resulting resolution 58,173] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
337. ~18 [avatar contradiction clause 329]
348. $false [avatar sat refutation 67,72,77,82,87,92,97,102,107,112,117,122,127,132,188,255,317,328,337] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:449,thDist:-449}
% SZS output end Proof for projpl_1__t36_projpl_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 8
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 213
% Active clauses: 30
% Passive clauses: 39
% Final active clauses: 29
% Final passive clauses: 8
% Main loop iterations started: 30
% 
% >>> Simplifying Inferences
% Duplicate literals: 171
% Fw subsumption resolutions: 108
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 6
% Unit resulting resolution: 5
% Forward superposition: 9
% Backward superposition: 15
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 7
% Split components: 16
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 34
% SAT solver unit clauses: 19
% 
% Memory used [KB]: 5117
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

