i: [1,70,0,10,2,1,0,0] t134_zfmisc_1
a: [3,121,0,5,1,70]
s: 121
i: [1,69,0,5,1,1,0,0] t134_zfmisc_1
a: [3,130,0,5,1,69]
s: 130
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,135,0,2,1,71]
s: 135
i: [1,72,0,2,1,0,0,255] fc1_xboole_0
a: [3,140,0,2,1,72]
s: 140
i: [1,102,0,2,1,0,0,2] rc2_xboole_0
a: [3,150,0,2,1,102]
s: 150
i: [1,103,0,2,1,0,0,3] rc2_zfmisc_1
a: [3,155,0,2,1,103]
s: 155
i: [1,104,0,2,1,0,0,2] rc1_xboole_0
a: [3,160,0,2,1,104]
s: 160
i: [1,106,0,2,1,0,0,3] rc1_zfmisc_1
a: [3,165,0,2,1,106]
s: 165
i: [1,105,0,2,1,0,0,3] rc1_zfmisc_1
a: [3,170,0,2,1,105]
s: 170
i: [1,74,0,3,1,0,0,2] fc2_xboole_0
s: 74
i: [1,75,0,3,1,0,0,2] fc2_zfmisc_1
s: 75
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,145,0,3,1,73]
s: 145
r: 135
i: [1,76,0,4,2,0,0,3] cc2_zfmisc_1
s: 76
d: [2,173,1,2,1,1,40,76,160]
a: [3,178,1,2,1,173]
s: 178
i: [1,77,0,4,2,0,0,3] cc1_zfmisc_1
s: 77
i: [1,98,0,4,1,0,0,1] fc1_zfmisc_1
s: 98
d: [2,174,1,2,1,1,40,76,140]
a: [3,183,1,2,1,174]
s: 183
i: [1,99,0,4,1,0,0,2] fc3_xboole_0
s: 99
i: [1,108,0,4,1,0,0,2] d1_tarski
s: 108
i: [1,78,0,5,2,0,0,255] t6_boole
s: 78
i: [1,91,0,5,2,0,0,1] t7_boole
s: 91
d: [2,188,1,2,1,1,40,91,130]
a: [3,193,1,2,1,188]
s: 193
i: [1,114,0,5,1,0,0,2] d2_tarski
s: 114
i: [1,116,0,5,1,0,0,2] d2_tarski
s: 116
i: [1,101,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 101
i: [1,90,0,7,3,0,0,2] t8_boole
s: 90
d: [2,197,1,5,1,1,40,101,130]
a: [3,204,1,5,1,197]
s: 204
i: [1,100,0,7,1,0,0,2] commutativity_k2_tarski
s: 100
d: [2,198,1,4,1,0,40,101,108]
s: 198
i: [1,109,0,7,2,0,0,2] d1_tarski
s: 109
i: [1,111,0,9,2,0,0,1] d4_tarski
s: 111
d: [2,199,1,5,1,0,40,101,114]
s: 199
i: [1,112,0,9,2,0,0,1] d4_tarski
s: 112
d: [2,200,1,5,1,0,40,101,116]
s: 200
i: [1,79,0,10,1,0,0,1] d5_tarski
s: 79
i: [1,110,0,10,3,0,0,1] d4_tarski
s: 110
d: [2,239,1,7,2,1,40,110,121]
s: 239
i: [1,117,0,11,3,0,0,2] d2_tarski
s: 117
d: [2,240,2,9,3,1,40,239,110]
s: 240
i: [1,82,0,14,3,0,0,2] d1_tarski
s: 82
i: [1,83,0,14,3,0,0,2] d1_tarski
s: 83
d: [2,251,3,7,2,1,40,240,108]
s: 251
i: [1,88,0,14,3,0,0,1] d4_tarski
s: 88
d: [2,261,4,5,1,1,40,251,108]
a: [3,266,4,5,1,261]
s: 266
i: [1,87,0,16,3,0,0,1] d4_tarski
s: 87
i: [1,89,0,17,4,0,0,1] d4_tarski
s: 89
d: [2,259,4,6,1,1,40,251,114]
s: 259
i: [1,96,0,17,3,0,0,2] d2_tarski
s: 96
d: [2,284,5,5,1,1,44,259,79]
s: 284
d: [2,286,6,0,0,2,40,284,130]
e: 286
d: [2,285,6,0,0,2,52,130,284]
e: 285
r: 284
r: 259
r: 251
r: 240
r: 239
r: 121
a: [3,120,0,5,1,70]
s: 120
a: [3,125,0,5,1,70]
s: 125
d: [2,290,1,4,1,1,40,120,91]
a: [3,299,1,4,1,290]
s: 299
i: [1,97,0,17,3,0,0,2] d2_tarski
s: 97
d: [2,289,1,5,1,1,40,120,101]
a: [3,294,1,5,1,289]
s: 294
i: [1,95,0,23,4,0,0,2] d2_tarski
s: 95
d: [2,301,1,7,2,1,40,125,110]
s: 301
d: [2,306,2,9,3,1,40,301,110]
s: 306
d: [2,214,1,6,2,0,40,111,91]
s: 214
d: [2,317,2,3,1,1,40,214,120]
a: [3,325,2,3,1,317]
s: 325
d: [2,230,1,6,1,0,44,114,79]
s: 230
d: [2,185,1,7,3,0,44,78,78]
s: 185
d: [2,314,3,7,2,1,40,306,108]
s: 314
d: [2,231,1,7,1,0,44,116,79]
s: 231
d: [2,340,3,0,0,2,52,114,130,231,306]
e: 340
f: 340,285,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t134_zfmisc_1
% SZS output start Proof for zfmisc_1__t134_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => (r2_hidden(X1,k3_tarski(k3_tarski(X2))) & r2_hidden(X0,k3_tarski(k3_tarski(X2))))) [input t134_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(k4_tarski(X0,X1),X2) => (r2_hidden(X1,k3_tarski(k3_tarski(X2))) & r2_hidden(X0,k3_tarski(k3_tarski(X2))))) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
8. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
10. ! [X0] : ! [X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [input d4_tarski]
11. ! [X0] : ! [X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [input d5_tarski]
30. ~! [X0,X1,X2] : (r2_hidden(k4_tarski(X0,X1),X2) => (r2_hidden(X1,k3_tarski(k3_tarski(X2))) & r2_hidden(X0,k3_tarski(k3_tarski(X2))))) [flattening 2]
31. ! [X0,X1] : k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [flattening 11]
32. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
33. ! [X0,X1] : (k3_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> ? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)))) [flattening 10]
36. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 8]
37. ? [X0,X1,X2] : ((~r2_hidden(X1,k3_tarski(k3_tarski(X2))) | ~r2_hidden(X0,k3_tarski(k3_tarski(X2)))) & r2_hidden(k4_tarski(X0,X1),X2)) [ennf transformation 30]
44. ? [X0,X1,X2] : ((~r2_hidden(X1,k3_tarski(k3_tarski(X2))) | ~r2_hidden(X0,k3_tarski(k3_tarski(X2)))) & r2_hidden(k4_tarski(X0,X1),X2)) => ((~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) | ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2)))) & r2_hidden(k4_tarski(sK0,sK1),sK2)) [choice axiom]
45. (~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) | ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2)))) & r2_hidden(k4_tarski(sK0,sK1),sK2) [skolemisation 37,44]
46. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 32]
47. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 46]
48. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
49. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 47,48]
50. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3))) & (? [X3] : (r2_hidden(X3,X0) & r2_hidden(X2,X3)) | ~r2_hidden(X2,X1))) | k3_tarski(X0) != X1)) [nnf transformation 33]
51. ! [X0,X1] : ((k3_tarski(X0) = X1 | ? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [rectify 50]
52. ! [X1,X0] : (? [X2] : ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(X2,X3)) | ~r2_hidden(X2,X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(X2,X4)) | r2_hidden(X2,X1))) => ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK4(X0,X1),X3)) | ~r2_hidden(sK4(X0,X1),X1)) & (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK4(X0,X1),X4)) | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
53. ! [X1,X0] : (? [X4] : (r2_hidden(X4,X0) & r2_hidden(sK4(X0,X1),X4)) => (r2_hidden(sK5(X0,X1),X0) & r2_hidden(sK4(X0,X1),sK5(X0,X1)))) [choice axiom]
54. ! [X5,X0] : (? [X7] : (r2_hidden(X7,X0) & r2_hidden(X5,X7)) => (r2_hidden(sK6(X0,X5),X0) & r2_hidden(X5,sK6(X0,X5)))) [choice axiom]
55. ! [X0,X1] : ((k3_tarski(X0) = X1 | ((! [X3] : (~r2_hidden(X3,X0) | ~r2_hidden(sK4(X0,X1),X3)) | ~r2_hidden(sK4(X0,X1),X1)) & ((r2_hidden(sK5(X0,X1),X0) & r2_hidden(sK4(X0,X1),sK5(X0,X1))) | r2_hidden(sK4(X0,X1),X1)))) & (! [X5] : ((r2_hidden(X5,X1) | ! [X6] : (~r2_hidden(X6,X0) | ~r2_hidden(X5,X6))) & ((r2_hidden(sK6(X0,X5),X0) & r2_hidden(X5,sK6(X0,X5))) | ~r2_hidden(X5,X1))) | k3_tarski(X0) != X1)) [skolemisation 51,54,53,52]
56. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 36]
57. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 56]
58. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 57]
59. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK7(X0,X1,X2) != X1 & sK7(X0,X1,X2) != X0) | ~r2_hidden(sK7(X0,X1,X2),X2)) & (sK7(X0,X1,X2) = X1 | sK7(X0,X1,X2) = X0 | r2_hidden(sK7(X0,X1,X2),X2)))) [choice axiom]
60. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK7(X0,X1,X2) != X1 & sK7(X0,X1,X2) != X0) | ~r2_hidden(sK7(X0,X1,X2),X2)) & (sK7(X0,X1,X2) = X1 | sK7(X0,X1,X2) = X0 | r2_hidden(sK7(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 58,59]
69. r2_hidden(k4_tarski(sK0,sK1),sK2) [cnf transformation 45] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) | ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) [cnf transformation 45] {a:0,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. k4_tarski(X0,X1) = k2_tarski(k2_tarski(X0,X1),k1_tarski(X0)) [cnf transformation 31] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
81. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 49] {a:0,w:10,wCS:4000,sine:2,thAx:0,allAx:2,thDist:-2}
86. r2_hidden(X5,X1) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) | k3_tarski(X0) != X1 [cnf transformation 55] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:4,thDist:-4}
93. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 60] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:2,thDist:-2}
94. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 60] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:2,thDist:-2}
107. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 81] {a:0,w:7,wCS:2800,sine:2,thAx:0,allAx:2,thDist:-2}
108. r2_hidden(X3,k1_tarski(X3)) [equality resolution 107] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
110. r2_hidden(X5,k3_tarski(X0)) | ~r2_hidden(X6,X0) | ~r2_hidden(X5,X6) [equality resolution 86] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
113. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 94] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:2,thDist:-2}
114. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 113] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
115. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 93] {a:0,w:8,wCS:3200,sine:2,thAx:0,allAx:2,thDist:-2}
116. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 115] {a:0,w:5,wCS:2000,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
119. 1 <=> r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) [avatar definition]
121. ~r2_hidden(sK0,k3_tarski(k3_tarski(sK2))) <- (~1) [avatar component clause 119] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
123. 2 <=> r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) [avatar definition]
125. ~r2_hidden(sK1,k3_tarski(k3_tarski(sK2))) <- (~2) [avatar component clause 123] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
126. ~1 | ~2 [avatar split clause 70,123,119]
128. 3 <=> r2_hidden(k4_tarski(sK0,sK1),sK2) [avatar definition]
130. r2_hidden(k4_tarski(sK0,sK1),sK2) <- (3) [avatar component clause 128] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. 3 [avatar split clause 69,128]
231. r2_hidden(k2_tarski(X8,X9),k4_tarski(X8,X9)) [superposition 116,79] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
239. ~r2_hidden(X14,k3_tarski(sK2)) | ~r2_hidden(sK0,X14) <- (~1) [resolution 110,121] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
240. ~r2_hidden(sK0,X0) | ~r2_hidden(X1,sK2) | ~r2_hidden(X0,X1) <- (~1) [resolution 239,110] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
251. ~r2_hidden(k1_tarski(sK0),X7) | ~r2_hidden(X7,sK2) <- (~1) [resolution 240,108] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
259. ~r2_hidden(k2_tarski(X2,k1_tarski(sK0)),sK2) <- (~1) [resolution 251,114] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:14,thDist:-14}
284. ~r2_hidden(k4_tarski(sK0,X2),sK2) <- (~1) [superposition 259,79] {a:5,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
285. $false <- (~1, 3) [unit resulting resolution 130,284] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
288. 1 | ~3 [avatar contradiction clause 285]
301. ~r2_hidden(X0,k3_tarski(sK2)) | ~r2_hidden(sK1,X0) <- (~2) [resolution 125,110] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
306. ~r2_hidden(sK1,X0) | ~r2_hidden(X1,sK2) | ~r2_hidden(X0,X1) <- (~2) [resolution 301,110] {a:2,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
340. $false <- (~2, 3) [unit resulting resolution 114,130,231,306] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
346. 2 | ~3 [avatar contradiction clause 340]
347. $false [avatar sat refutation 126,131,288,346] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:149,thDist:-149}
% SZS output end Proof for zfmisc_1__t134_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 28
% Introduced skolems: 12
% 
% >>> Saturation
% Initial clauses: 38
% Generated clauses: 186
% Active clauses: 61
% Passive clauses: 135
% Final active clauses: 54
% Final passive clauses: 60
% Main loop iterations started: 63
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 3
% Forward subsumptions: 21
% 
% >>> Generating Inferences
% Binary resolution: 89
% Unit resulting resolution: 2
% Forward superposition: 17
% Backward superposition: 15
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 9
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

