i: [1,49,0,7,2,1,0,0] t34_int_1
a: [3,75,0,2,1,49]
s: 75
i: [1,47,0,2,1,1,0,0] t34_int_1
a: [3,82,0,2,1,47]
s: 82
a: [3,72,0,5,1,49]
s: 72
i: [1,64,0,7,3,0,0,255] reflexivity_r1_xxreal_0
a: [3,86,0,2,1,64]
s: 86
i: [1,66,0,2,1,0,0,255] rc1_xxreal_0
a: [3,94,0,2,1,66]
s: 94
d: [2,106,0,0,0,2,21,94,86]
e: 106
r: 86
s: 94
i: [1,67,0,2,1,0,0,1] rc1_xreal_0
a: [3,99,0,2,1,67]
s: 99
i: [1,68,0,2,1,0,0,1] rc2_int_1
a: [3,104,0,2,1,68]
s: 104
i: [1,50,0,4,2,0,0,255] cc4_xreal_0
s: 50
i: [1,63,0,4,2,0,0,1] cc3_int_1
s: 63
i: [1,51,0,5,2,0,0,1] dt_k2_int_1
s: 51
d: [2,108,1,2,1,1,40,50,82]
a: [3,113,1,2,1,108]
s: 113
i: [1,52,0,5,2,0,0,1] dt_k1_int_1
s: 52
d: [2,109,1,2,1,1,40,50,99]
a: [3,118,1,2,1,109]
s: 118
a: [3,89,0,5,2,64]
s: 89
i: [1,53,0,8,2,0,0,1] projectivity_k1_int_1
s: 53
d: [2,121,1,2,1,1,40,63,104]
a: [3,125,1,2,1,121]
s: 125
i: [1,54,0,8,2,0,0,1] projectivity_k2_int_1
s: 54
d: [2,131,2,2,1,1,40,125,50]
a: [3,135,2,2,1,131]
s: 135
i: [1,55,0,8,3,0,0,255] t26_int_1
s: 55
d: [2,139,0,6,2,0,21,55,63]
s: 139
d: [2,127,1,5,2,0,40,51,63]
s: 127
i: [1,56,0,8,3,0,0,255] t26_int_1
s: 56
d: [2,128,1,5,2,0,40,52,63]
s: 128
i: [1,57,0,8,3,0,0,255] t31_int_1
s: 57
d: [2,148,0,6,2,0,21,57,63]
s: 148
d: [2,142,2,5,2,0,40,127,50]
s: 142
i: [1,58,0,8,3,0,0,255] t31_int_1
s: 58
d: [2,146,2,5,2,0,40,128,50]
s: 146
i: [1,59,0,8,3,0,0,1] t25_int_1
s: 59
i: [1,60,0,8,3,0,0,1] t25_int_1
s: 60
d: [2,157,0,6,2,0,21,60,63]
s: 157
i: [1,61,0,8,3,0,0,1] t30_int_1
s: 61
d: [2,163,1,7,3,0,44,139,157]
d: [2,165,1,5,2,0,18,163]
s: 165
i: [1,62,0,8,3,0,0,1] t30_int_1
s: 62
d: [2,168,0,6,2,0,21,62,63]
s: 168
d: [2,169,1,6,2,1,44,72,168]
d: [2,178,1,4,1,2,21,169,75]
a: [3,182,1,4,1,178]
s: 182
d: [2,184,2,0,0,2,52,75,182,157]
e: 184
d: [2,185,2,5,2,1,44,182,157]
d: [2,186,2,2,1,1,19,185]
d: [2,188,2,0,0,2,21,186,75]
e: 188
r: 72
r: 75
a: [3,76,0,2,1,49]
s: 76
a: [3,71,0,5,1,49]
s: 71
i: [1,65,0,10,4,0,0,255] connectedness_r1_xxreal_0
s: 65
d: [2,194,1,8,3,1,44,58,71]
d: [2,206,1,6,2,2,21,194,76]
d: [2,207,1,4,1,3,21,206,82]
a: [3,211,1,4,1,207]
s: 211
d: [2,221,2,9,3,1,40,211,65]
d: [2,223,2,7,2,2,21,221,113]
a: [3,227,2,3,1,223]
s: 227
d: [2,233,3,0,0,2,52,82,227,146]
e: 233
d: [2,234,3,2,1,1,40,227,146]
d: [2,237,3,0,0,2,21,234,82]
e: 237
r: 227
d: [2,193,1,8,2,1,44,54,71]
d: [2,200,1,6,1,2,21,193,82]
a: [3,204,1,6,1,200]
s: 204
a: [3,226,2,3,1,223]
s: 226
d: [2,140,1,10,3,0,44,139,53]
d: [2,141,1,7,2,0,21,140,52]
s: 141
d: [2,149,1,10,3,0,44,148,54]
d: [2,150,1,7,2,0,21,149,51]
s: 150
d: [2,240,2,6,2,1,44,51,204]
a: [3,251,2,3,1,240]
s: 251
d: [2,271,3,0,0,2,52,82,251,128]
e: 271
d: [2,272,3,2,1,1,40,251,128]
d: [2,275,3,0,0,2,21,272,82]
e: 275
r: 251
d: [2,213,1,13,5,0,40,65,56]
d: [2,217,1,11,4,0,21,213,50]
d: [2,218,1,8,3,0,21,217,146]
s: 218
d: [2,277,2,0,0,3,52,82,76,211,218]
e: 277
d: [2,278,2,4,2,1,40,218,211]
d: [2,282,2,2,1,2,21,278,76]
d: [2,283,2,0,0,3,21,282,82]
e: 283
i: [1,48,0,7,2,1,0,0] t34_int_1
f: 283,207,186,178,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for int_1__t34_int_1
% SZS output start Proof for int_1__t34_int_1
1. ! [X0] : (v1_xreal_0(X0) => (k1_int_1(X0) = k2_int_1(X0) <=> v1_int_1(X0))) [input t34_int_1]
2. ~! [X0] : (v1_xreal_0(X0) => (k1_int_1(X0) = k2_int_1(X0) <=> v1_int_1(X0))) [negated conjecture 1]
3. ! [X0] : (v1_int_1(X0) => v1_xreal_0(X0)) [input cc3_int_1]
5. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
6. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
7. ! [X0] : (v1_xreal_0(X0) => v1_int_1(k1_int_1(X0))) [input dt_k1_int_1]
15. ! [X0] : (v1_xreal_0(X0) => (k1_int_1(X0) = X0 <=> v1_int_1(X0))) [input t25_int_1]
16. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(X0,k1_int_1(X0)) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(X0,k1_int_1(X0))))) [input t26_int_1]
17. ! [X0] : (v1_xreal_0(X0) => (k2_int_1(X0) = X0 <=> v1_int_1(X0))) [input t30_int_1]
18. ! [X0] : (v1_xreal_0(X0) => (~(r1_xxreal_0(k2_int_1(X0),X0) & ~v1_int_1(X0)) & ~(v1_int_1(X0) & ~r1_xxreal_0(k2_int_1(X0),X0)))) [input t31_int_1]
20. ? [X0] : ((k1_int_1(X0) = k2_int_1(X0) <~> v1_int_1(X0)) & v1_xreal_0(X0)) [ennf transformation 2]
21. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 5]
23. ! [X0] : (v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 7]
26. ! [X0] : (((~r1_xxreal_0(X0,k1_int_1(X0)) | v1_int_1(X0)) & (~v1_int_1(X0) | r1_xxreal_0(X0,k1_int_1(X0)))) | ~v1_xreal_0(X0)) [ennf transformation 16]
27. ! [X0] : (((~r1_xxreal_0(k2_int_1(X0),X0) | v1_int_1(X0)) & (~v1_int_1(X0) | r1_xxreal_0(k2_int_1(X0),X0))) | ~v1_xreal_0(X0)) [ennf transformation 18]
28. ! [X0] : ((k1_int_1(X0) = X0 <=> v1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 15]
29. ! [X0] : ((k2_int_1(X0) = X0 <=> v1_int_1(X0)) | ~v1_xreal_0(X0)) [ennf transformation 17]
30. ! [X0] : (v1_xreal_0(X0) | ~v1_int_1(X0)) [ennf transformation 3]
33. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 6]
34. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 33]
35. ? [X0] : (((~v1_int_1(X0) | k1_int_1(X0) != k2_int_1(X0)) & (v1_int_1(X0) | k1_int_1(X0) = k2_int_1(X0))) & v1_xreal_0(X0)) [nnf transformation 20]
36. ? [X0] : ((~v1_int_1(X0) | k1_int_1(X0) != k2_int_1(X0)) & (v1_int_1(X0) | k1_int_1(X0) = k2_int_1(X0)) & v1_xreal_0(X0)) [flattening 35]
37. ? [X0] : ((~v1_int_1(X0) | k1_int_1(X0) != k2_int_1(X0)) & (v1_int_1(X0) | k1_int_1(X0) = k2_int_1(X0)) & v1_xreal_0(X0)) => ((~v1_int_1(sK0) | k1_int_1(sK0) != k2_int_1(sK0)) & (v1_int_1(sK0) | k1_int_1(sK0) = k2_int_1(sK0)) & v1_xreal_0(sK0)) [choice axiom]
38. (~v1_int_1(sK0) | k1_int_1(sK0) != k2_int_1(sK0)) & (v1_int_1(sK0) | k1_int_1(sK0) = k2_int_1(sK0)) & v1_xreal_0(sK0) [skolemisation 36,37]
39. ! [X0] : (((k1_int_1(X0) = X0 | ~v1_int_1(X0)) & (v1_int_1(X0) | k1_int_1(X0) != X0)) | ~v1_xreal_0(X0)) [nnf transformation 28]
40. ! [X0] : (((k2_int_1(X0) = X0 | ~v1_int_1(X0)) & (v1_int_1(X0) | k2_int_1(X0) != X0)) | ~v1_xreal_0(X0)) [nnf transformation 29]
47. v1_xreal_0(sK0) [cnf transformation 38] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
48. v1_int_1(sK0) | k1_int_1(sK0) = k2_int_1(sK0) [cnf transformation 38] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. ~v1_int_1(sK0) | k1_int_1(sK0) != k2_int_1(sK0) [cnf transformation 38] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 21] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
52. v1_int_1(k1_int_1(X0)) | ~v1_xreal_0(X0) [cnf transformation 23] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. ~r1_xxreal_0(X0,k1_int_1(X0)) | v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 26] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
58. ~r1_xxreal_0(k2_int_1(X0),X0) | v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
60. k1_int_1(X0) = X0 | ~v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 39] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
62. k2_int_1(X0) = X0 | ~v1_int_1(X0) | ~v1_xreal_0(X0) [cnf transformation 40] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
63. ~v1_int_1(X0) | v1_xreal_0(X0) [cnf transformation 30] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
65. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 34] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
70. 1 <=> k1_int_1(sK0) = k2_int_1(sK0) [avatar definition]
71. k1_int_1(sK0) = k2_int_1(sK0) <- (1) [avatar component clause 70] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
72. k1_int_1(sK0) != k2_int_1(sK0) <- (~1) [avatar component clause 70] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
74. 2 <=> v1_int_1(sK0) [avatar definition]
75. v1_int_1(sK0) <- (2) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. ~v1_int_1(sK0) <- (~2) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
77. ~1 | ~2 [avatar split clause 49,74,70]
78. 1 | 2 [avatar split clause 48,74,70]
80. 3 <=> v1_xreal_0(sK0) [avatar definition]
82. v1_xreal_0(sK0) <- (3) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
83. 3 [avatar split clause 47,80]
128. v1_xreal_0(k1_int_1(X0)) | ~v1_xreal_0(X0) [resolution 52,63] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
146. v1_xxreal_0(k1_int_1(X0)) | ~v1_xreal_0(X0) [resolution 128,50] {a:2,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
157. k1_int_1(X0) = X0 | ~v1_int_1(X0) [subsumption resolution 60,63] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
168. k2_int_1(X0) = X0 | ~v1_int_1(X0) [subsumption resolution 62,63] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
169. sK0 != k1_int_1(sK0) | ~v1_int_1(sK0) <- (~1) [superposition 72,168] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
178. sK0 != k1_int_1(sK0) <- (~1, 2) [subsumption resolution 169,75] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
180. 13 <=> sK0 = k1_int_1(sK0) [avatar definition]
182. sK0 != k1_int_1(sK0) <- (~13) [avatar component clause 180] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
183. ~13 | 1 | ~2 [avatar split clause 178,74,70,180]
185. sK0 != sK0 | ~v1_int_1(sK0) <- (~13) [superposition 182,157] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
186. ~v1_int_1(sK0) <- (~13) [trivial inequality removal 185] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
190. ~2 | 13 [avatar split clause 186,180,74]
194. ~r1_xxreal_0(k1_int_1(sK0),sK0) | v1_int_1(sK0) | ~v1_xreal_0(sK0) <- (1) [superposition 58,71] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
206. ~r1_xxreal_0(k1_int_1(sK0),sK0) | ~v1_xreal_0(sK0) <- (1, ~2) [subsumption resolution 194,76] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
207. ~r1_xxreal_0(k1_int_1(sK0),sK0) <- (1, ~2, 3) [subsumption resolution 206,82] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
209. 15 <=> r1_xxreal_0(k1_int_1(sK0),sK0) [avatar definition]
211. ~r1_xxreal_0(k1_int_1(sK0),sK0) <- (~15) [avatar component clause 209] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
212. ~15 | ~1 | 2 | ~3 [avatar split clause 207,80,74,70,209]
213. r1_xxreal_0(k1_int_1(X0),X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(k1_int_1(X0)) | v1_int_1(X0) | ~v1_xreal_0(X0) [resolution 65,56] {a:1,w:13,wCS:5200,sine:255,thAx:0,allAx:2,thDist:-2}
217. r1_xxreal_0(k1_int_1(X0),X0) | ~v1_xxreal_0(k1_int_1(X0)) | v1_int_1(X0) | ~v1_xreal_0(X0) [subsumption resolution 213,50] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
218. r1_xxreal_0(k1_int_1(X0),X0) | v1_int_1(X0) | ~v1_xreal_0(X0) [subsumption resolution 217,146] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
278. v1_int_1(sK0) | ~v1_xreal_0(sK0) <- (~15) [resolution 218,211] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
282. ~v1_xreal_0(sK0) <- (~2, ~15) [subsumption resolution 278,76] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
283. $false <- (~2, 3, ~15) [subsumption resolution 282,82] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
284. 2 | ~3 | 15 [avatar contradiction clause 283]
285. $false [avatar sat refutation 77,78,83,183,190,212,284] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:163,thDist:-163}
% SZS output end Proof for int_1__t34_int_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 22
% Generated clauses: 157
% Active clauses: 43
% Passive clauses: 53
% Final active clauses: 38
% Final passive clauses: 5
% Main loop iterations started: 48
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Trivial inequalities: 4
% Fw subsumption resolutions: 21
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Forward subsumptions: 38
% 
% >>> Generating Inferences
% Binary resolution: 17
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 25
% Backward superposition: 32
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 8
% Split components: 17
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 31
% SAT solver unit clauses: 4
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

