i: [1,35,0,3,1,1,0,0] l18_arytm_1
a: [3,62,0,3,1,35]
s: 62
i: [1,34,0,3,1,1,0,0] l18_arytm_1
a: [3,67,0,3,1,34]
s: 67
i: [1,36,0,7,1,1,0,0] l18_arytm_1
a: [3,57,0,7,1,36]
s: 57
i: [1,37,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,72,0,3,1,37]
s: 72
i: [1,38,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,77,0,3,1,38]
s: 77
i: [1,49,0,17,5,0,0,1] d1_arytm_1
a: [3,82,0,3,1,49]
s: 82
i: [1,44,0,4,1,0,0,1] existence_m1_subset_1
s: 44
i: [1,45,0,11,3,0,0,1] dt_k1_arytm_1
s: 45
i: [1,46,0,11,3,0,0,1] dt_k7_arytm_2
s: 46
i: [1,48,0,12,4,0,0,2] connectedness_r1_arytm_2
s: 48
i: [1,47,0,13,3,0,0,1] commutativity_k7_arytm_2
s: 47
d: [2,89,1,9,3,1,40,48,67]
s: 89
i: [1,53,0,16,4,0,0,2] t19_arytm_2
s: 53
d: [2,98,2,6,2,1,40,89,67]
d: [2,101,2,3,1,1,18,98]
a: [3,105,2,3,1,101]
s: 105
d: [2,125,0,11,3,0,21,53,46]
s: 125
i: [1,50,0,19,5,0,0,1] d1_arytm_1
s: 50
d: [2,99,2,6,2,2,40,89,62]
a: [3,110,2,3,1,99]
s: 110
a: [3,85,0,14,4,49]
s: 85
d: [2,100,2,8,2,1,40,89,44]
a: [3,119,2,4,1,100]
s: 119
i: [1,52,0,21,5,0,0,1] d1_arytm_1
s: 52
d: [2,150,0,16,4,0,21,52,45]
s: 150
d: [2,132,1,16,4,2,44,57,85]
d: [2,135,1,13,3,3,21,132,62]
a: [3,143,1,5,1,135]
s: 143
d: [2,169,2,0,0,3,52,67,62,143,46]
e: 169
d: [2,170,2,6,2,1,40,143,46]
d: [2,172,2,3,1,2,21,170,62]
d: [2,173,2,0,0,3,21,172,67]
e: 173
r: 143
i: [1,51,0,23,5,0,0,1] d1_arytm_1
s: 51
a: [3,147,1,3,1,135]
s: 147
d: [2,175,0,18,4,0,21,51,125]
d: [2,176,0,13,3,0,21,175,46]
s: 176
d: [2,177,1,0,0,3,52,62,67,57,176]
e: 177
d: [2,181,1,9,3,1,44,57,176]
d: [2,186,1,6,2,1,19,181]
d: [2,191,1,3,1,2,21,186,67]
d: [2,192,1,0,0,3,21,191,62]
e: 192
f: 192,34,35,36
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__l18_arytm_1
% SZS output start Proof for arytm_1__l18_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => k1_arytm_1(k7_arytm_2(X0,X1),X1) = X0)) [input l18_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => k1_arytm_1(k7_arytm_2(X0,X1),X1) = X0)) [negated conjecture 1]
5. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2)) & (r1_arytm_2(X1,X0) => (k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0)))))) [input d1_arytm_1]
11. ! [X0,X1] : ((m1_subset_1(X1,k2_arytm_2) & m1_subset_1(X0,k2_arytm_2)) => m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2)) [input dt_k7_arytm_2]
15. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ! [X2] : (m1_subset_1(X2,k2_arytm_2) => (k7_arytm_2(X1,X2) = X0 => r1_arytm_2(X2,X0))))) [input t19_arytm_2]
16. ? [X0] : (? [X1] : (k1_arytm_1(k7_arytm_2(X0,X1),X1) != X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
17. ! [X0] : (! [X1] : (! [X2] : ((((k1_arytm_1(X0,X1) = X2 <=> k11_arytm_3 = X2) | r1_arytm_2(X1,X0)) & ((k1_arytm_1(X0,X1) = X2 <=> k7_arytm_2(X2,X1) = X0) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 5]
18. ! [X0] : (! [X1] : (! [X2] : ((r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 15]
19. ! [X0] : (! [X1] : (! [X2] : (r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 18]
22. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | (~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2))) [ennf transformation 11]
23. ! [X0,X1] : (m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 22]
28. ? [X0] : (? [X1] : (k1_arytm_1(k7_arytm_2(X0,X1),X1) != X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (sK0 != k1_arytm_1(k7_arytm_2(sK0,X1),X1) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
29. ? [X1] : (sK0 != k1_arytm_1(k7_arytm_2(sK0,X1),X1) & m1_subset_1(X1,k2_arytm_2)) => (sK0 != k1_arytm_1(k7_arytm_2(sK0,sK1),sK1) & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
30. (sK0 != k1_arytm_1(k7_arytm_2(sK0,sK1),sK1) & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 16,29,28]
31. ! [X0] : (! [X1] : (! [X2] : (((((k1_arytm_1(X0,X1) = X2 | k11_arytm_3 != X2) & (k11_arytm_3 = X2 | k1_arytm_1(X0,X1) != X2)) | r1_arytm_2(X1,X0)) & (((k1_arytm_1(X0,X1) = X2 | k7_arytm_2(X2,X1) != X0) & (k7_arytm_2(X2,X1) = X0 | k1_arytm_1(X0,X1) != X2)) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X2,k2_arytm_2)) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [nnf transformation 17]
34. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. sK0 != k1_arytm_1(k7_arytm_2(sK0,sK1),sK1) [cnf transformation 30] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. k1_arytm_1(X0,X1) = X2 | k7_arytm_2(X2,X1) != X0 | ~r1_arytm_2(X1,X0) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 31] {a:0,w:22,wCS:8800,sine:1,thAx:0,allAx:1,thDist:-1}
43. r1_arytm_2(X2,X0) | k7_arytm_2(X1,X2) != X0 | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 19] {a:0,w:17,wCS:6800,sine:2,thAx:0,allAx:1,thDist:-1}
46. m1_subset_1(k7_arytm_2(X0,X1),k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k1_arytm_1(k7_arytm_2(X2,X1),X1) = X2 | ~r1_arytm_2(X1,k7_arytm_2(X2,X1)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X2,X1),k2_arytm_2) [equality resolution 40] {a:0,w:23,wCS:9200,sine:1,thAx:0,allAx:1,thDist:-1}
53. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X1,X2),k2_arytm_2) [equality resolution 43] {a:0,w:16,wCS:6400,sine:2,thAx:0,allAx:1,thDist:-1}
55. 1 <=> sK0 = k1_arytm_1(k7_arytm_2(sK0,sK1),sK1) [avatar definition]
57. sK0 != k1_arytm_1(k7_arytm_2(sK0,sK1),sK1) <- (~1) [avatar component clause 55] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. ~1 [avatar split clause 36,55]
60. 2 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
62. m1_subset_1(sK1,k2_arytm_2) <- (2) [avatar component clause 60] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. 2 [avatar split clause 35,60]
65. 3 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
67. m1_subset_1(sK0,k2_arytm_2) <- (3) [avatar component clause 65] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
68. 3 [avatar split clause 34,65]
125. r1_arytm_2(X2,k7_arytm_2(X1,X2)) | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 53,46] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
175. k1_arytm_1(k7_arytm_2(X2,X1),X1) = X2 | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k7_arytm_2(X2,X1),k2_arytm_2) [subsumption resolution 51,125] {a:0,w:18,wCS:7200,sine:1,thAx:0,allAx:1,thDist:-1}
176. k1_arytm_1(k7_arytm_2(X2,X1),X1) = X2 | ~m1_subset_1(X2,k2_arytm_2) | ~m1_subset_1(X1,k2_arytm_2) [subsumption resolution 175,46] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
181. sK0 != sK0 | ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~1) [superposition 57,176] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
186. ~m1_subset_1(sK0,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~1) [trivial inequality removal 181] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
191. ~m1_subset_1(sK1,k2_arytm_2) <- (~1, 3) [subsumption resolution 186,67] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
192. $false <- (~1, 2, 3) [subsumption resolution 191,62] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
193. 1 | ~2 | ~3 [avatar contradiction clause 192]
194. $false [avatar sat refutation 58,63,68,193] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:90,thDist:-90}
% SZS output end Proof for arytm_1__l18_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 14
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 94
% Active clauses: 21
% Passive clauses: 39
% Final active clauses: 20
% Final passive clauses: 14
% Main loop iterations started: 25
% 
% >>> Simplifying Inferences
% Duplicate literals: 26
% Trivial inequalities: 1
% Fw subsumption resolutions: 15
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 4
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 2
% Forward superposition: 8
% Backward superposition: 11
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 9
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 17
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

