i: [1,135,0,2,1,1,0,0] t12_membered
a: [3,176,0,2,1,135]
s: 176
i: [1,134,0,2,1,1,0,0] t12_membered
a: [3,181,0,2,1,134]
s: 181
i: [1,136,0,5,2,1,0,0] t12_membered
s: 136
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,186,0,2,1,137]
s: 186
i: [1,160,0,2,1,0,0,1] rc4_ordinal1
a: [3,196,0,2,1,160]
s: 196
i: [1,161,0,2,1,0,0,255] rc1_ordinal1
a: [3,201,0,2,1,161]
s: 201
i: [1,165,0,2,1,0,0,255] rc2_ordinal1
a: [3,206,0,2,1,165]
s: 206
i: [1,164,0,2,1,0,0,255] rc2_ordinal1
a: [3,211,0,2,1,164]
s: 211
i: [1,163,0,2,1,0,0,255] rc2_ordinal1
a: [3,216,0,2,1,163]
s: 216
i: [1,162,0,2,1,0,0,255] rc2_ordinal1
a: [3,221,0,2,1,162]
s: 221
i: [1,167,0,2,1,0,0,1] rc1_membered
a: [3,226,0,2,1,167]
s: 226
i: [1,166,0,2,1,0,0,1] rc1_membered
a: [3,231,0,2,1,166]
s: 231
i: [1,169,0,2,1,0,0,1] rc5_ordinal1
a: [3,236,0,2,1,169]
s: 236
i: [1,168,0,2,1,0,0,1] rc5_ordinal1
a: [3,241,0,2,1,168]
s: 241
i: [1,170,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,246,0,2,1,170]
s: 246
i: [1,171,0,2,1,0,0,255] rc4_xreal_0
a: [3,251,0,2,1,171]
s: 251
i: [1,172,0,2,1,0,0,255] rc3_xreal_0
a: [3,256,0,2,1,172]
s: 256
i: [1,138,0,3,1,0,0,255] d2_xboole_0
a: [3,191,0,3,1,138]
s: 191
r: 186
d: [2,258,0,2,1,2,23,186,191]
a: [3,262,0,2,1,258]
s: 262
i: [1,140,0,4,2,0,0,1] cc7_ordinal1
s: 140
i: [1,141,0,4,2,0,0,255] cc3_ordinal1
s: 141
d: [2,264,1,2,1,1,40,140,251]
a: [3,269,1,2,1,264]
s: 269
i: [1,143,0,4,2,0,0,255] cc6_ordinal1
s: 143
d: [2,265,1,2,1,1,40,140,262]
a: [3,274,1,2,1,265]
s: 274
i: [1,145,0,4,2,0,0,255] cc3_nat_1
s: 145
i: [1,146,0,4,2,0,0,255] cc1_nat_1
s: 146
d: [2,276,1,2,1,1,40,141,251]
a: [3,281,1,2,1,276]
s: 281
i: [1,148,0,4,2,0,0,255] cc1_ordinal1
s: 148
d: [2,277,1,2,1,1,40,141,262]
a: [3,286,1,2,1,277]
s: 286
i: [1,149,0,4,2,0,0,255] cc1_ordinal1
s: 149
i: [1,158,0,4,1,0,0,2] existence_m1_subset_1
s: 158
d: [2,288,1,2,1,1,40,143,196]
a: [3,294,1,2,1,288]
s: 294
i: [1,142,0,5,2,0,0,255] t6_boole
s: 142
d: [2,289,1,2,1,1,40,143,236]
a: [3,299,1,2,1,289]
s: 299
i: [1,152,0,5,2,0,0,1] d1_xboole_0
s: 152
i: [1,157,0,5,2,0,0,1] t7_boole
s: 157
d: [2,304,1,2,1,1,40,145,256]
a: [3,308,1,2,1,304]
s: 308
i: [1,151,0,6,3,0,0,255] cc2_ordinal1
s: 151
d: [2,310,1,2,1,1,40,148,201]
a: [3,316,1,2,1,310]
s: 316
i: [1,153,0,6,2,0,0,1] d1_xboole_0
s: 153
i: [1,154,0,6,2,0,0,1] t1_subset
s: 154
d: [2,377,1,5,2,0,40,153,136]
d: [2,379,1,3,1,1,21,377,181]
a: [3,383,1,3,1,379]
s: 383
i: [1,159,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 159
d: [2,330,1,2,1,1,40,149,201]
a: [3,341,1,2,1,330]
s: 341
i: [1,139,0,7,3,0,0,1] cc17_membered
s: 139
i: [1,150,0,7,3,0,0,255] cc5_ordinal1
s: 150
d: [2,312,2,2,1,1,40,148,281]
a: [3,321,2,2,1,312]
s: 321
i: [1,156,0,7,3,0,0,255] t8_boole
s: 156
d: [2,323,2,2,1,1,40,286,148]
a: [3,327,2,2,1,323]
s: 327
i: [1,155,0,8,3,0,0,1] t2_subset
s: 155
d: [2,388,1,5,2,0,40,139,158]
s: 388
d: [2,329,2,2,1,1,40,149,286]
a: [3,336,2,2,1,329]
s: 336
d: [2,389,1,5,2,0,40,150,158]
s: 389
d: [2,332,2,2,1,1,40,149,281]
a: [3,346,2,2,1,332]
s: 346
d: [2,385,1,6,2,0,40,154,153]
s: 385
d: [2,386,1,6,2,0,40,159,153]
s: 386
d: [2,348,2,2,1,1,40,294,149]
a: [3,353,2,2,1,348]
s: 353
d: [2,391,1,6,2,0,40,155,158]
s: 391
d: [2,403,2,5,2,0,40,391,136]
d: [2,407,2,3,1,1,21,403,181]
a: [3,411,2,3,1,407]
s: 411
d: [2,413,3,0,0,2,52,176,411,388]
e: 413
d: [2,414,3,0,0,2,52,176,158,411,139]
e: 414
d: [2,415,3,2,1,1,40,411,388]
d: [2,418,3,0,0,2,21,415,176]
e: 418
f: 418,407,134,135
% Refutation found. Thanks to Tanya!
% SZS status Theorem for membered__t12_membered
% SZS output start Proof for membered__t12_membered
1. ! [X0] : ((v6_membered(X0) & ~v1_xboole_0(X0)) => ? [X1] : (r2_hidden(X1,X0) & v7_ordinal1(X1))) [input t12_membered]
2. ~! [X0] : ((v6_membered(X0) & ~v1_xboole_0(X0)) => ? [X1] : (r2_hidden(X1,X0) & v7_ordinal1(X1))) [negated conjecture 1]
9. ! [X0] : (v6_membered(X0) => ! [X1] : (m1_subset_1(X1,X0) => v7_ordinal1(X1))) [input cc17_membered]
38. ! [X0] : ? [X1] : m1_subset_1(X1,X0) [input existence_m1_subset_1]
52. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
57. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 52]
91. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v7_ordinal1(X1)) & (v6_membered(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
92. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v7_ordinal1(X1)) & v6_membered(X0) & ~v1_xboole_0(X0)) [flattening 91]
93. ! [X0] : (! [X1] : (v7_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v6_membered(X0)) [ennf transformation 9]
105. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 57]
106. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 105]
110. ? [X0] : (! [X1] : (~r2_hidden(X1,X0) | ~v7_ordinal1(X1)) & v6_membered(X0) & ~v1_xboole_0(X0)) => (! [X1] : (~r2_hidden(X1,sK0) | ~v7_ordinal1(X1)) & v6_membered(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
111. ! [X1] : (~r2_hidden(X1,sK0) | ~v7_ordinal1(X1)) & v6_membered(sK0) & ~v1_xboole_0(sK0) [skolemisation 92,110]
116. ! [X0] : (? [X1] : m1_subset_1(X1,X0) => m1_subset_1(sK2(X0),X0)) [choice axiom]
117. ! [X0] : m1_subset_1(sK2(X0),X0) [skolemisation 38,116]
134. ~v1_xboole_0(sK0) [cnf transformation 111] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
135. v6_membered(sK0) [cnf transformation 111] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
136. ~r2_hidden(X1,sK0) | ~v7_ordinal1(X1) [cnf transformation 111] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
139. ~m1_subset_1(X1,X0) | v7_ordinal1(X1) | ~v6_membered(X0) [cnf transformation 93] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
155. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 106] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
158. m1_subset_1(sK2(X0),X0) [cnf transformation 117] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
174. 1 <=> v6_membered(sK0) [avatar definition]
176. v6_membered(sK0) <- (1) [avatar component clause 174] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
177. 1 [avatar split clause 135,174]
179. 2 <=> v1_xboole_0(sK0) [avatar definition]
181. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 179] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
182. ~2 [avatar split clause 134,179]
388. v7_ordinal1(sK2(X0)) | ~v6_membered(X0) [resolution 139,158] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
391. r2_hidden(sK2(X0),X0) | v1_xboole_0(X0) [resolution 155,158] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
403. v1_xboole_0(sK0) | ~v7_ordinal1(sK2(sK0)) [resolution 391,136] {a:2,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
407. ~v7_ordinal1(sK2(sK0)) <- (~2) [subsumption resolution 403,181] {a:2,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
409. 37 <=> v7_ordinal1(sK2(sK0)) [avatar definition]
411. ~v7_ordinal1(sK2(sK0)) <- (~37) [avatar component clause 409] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
412. ~37 | 2 [avatar split clause 407,179,409]
415. ~v6_membered(sK0) <- (~37) [resolution 411,388] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
418. $false <- (1, ~37) [subsumption resolution 415,176] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
419. ~1 | 37 [avatar contradiction clause 418]
420. $false [avatar sat refutation 177,182,412,419] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:194,thDist:-194}
% SZS output end Proof for membered__t12_membered
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 54
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 12
% 
% >>> Saturation
% Initial clauses: 39
% Generated clauses: 133
% Active clauses: 57
% Passive clauses: 71
% Final active clauses: 56
% Final passive clauses: 12
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 13
% 
% >>> Generating Inferences
% Binary resolution: 44
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 3
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 44
% SAT solver unit clauses: 17
% SAT solver binary clauses: 24
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% superposition: 0.001 s
% other: 0.002 s ( own 0.0000 s ) 

