i: [1,64,0,2,1,1,0,0] t90_mcart_1
a: [3,117,0,2,1,64]
s: 117
i: [1,63,0,2,1,1,0,0] t90_mcart_1
a: [3,122,0,2,1,63]
s: 122
i: [1,69,0,3,1,1,0,0] t90_mcart_1
a: [3,92,0,3,1,69]
s: 92
i: [1,66,0,3,1,1,0,0] t90_mcart_1
a: [3,107,0,3,1,66]
s: 107
i: [1,65,0,3,1,1,0,0] t90_mcart_1
a: [3,112,0,3,1,65]
s: 112
i: [1,68,0,5,1,1,0,0] t90_mcart_1
a: [3,97,0,5,1,68]
s: 97
i: [1,67,0,5,1,1,0,0] t90_mcart_1
a: [3,102,0,5,1,67]
s: 102
i: [1,70,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,127,0,2,1,70]
s: 127
i: [1,71,0,2,1,0,0,255] fc1_xboole_0
a: [3,132,0,2,1,71]
s: 132
i: [1,84,0,2,1,0,0,1] rc2_xboole_0
a: [3,142,0,2,1,84]
s: 142
i: [1,85,0,2,1,0,0,1] rc1_xboole_0
a: [3,147,0,2,1,85]
s: 147
i: [1,87,0,2,1,0,0,1] rc1_relat_1
a: [3,152,0,2,1,87]
s: 152
i: [1,86,0,2,1,0,0,1] rc1_relat_1
a: [3,157,0,2,1,86]
s: 157
i: [1,88,0,2,1,0,0,255] rc2_relat_1
a: [3,162,0,2,1,88]
s: 162
i: [1,72,0,3,1,0,0,255] d2_xboole_0
a: [3,137,0,3,1,72]
s: 137
r: 132
i: [1,73,0,4,2,0,0,1] cc1_relat_1
s: 73
d: [2,165,1,2,1,1,40,73,127]
a: [3,170,1,2,1,165]
s: 170
i: [1,82,0,4,1,0,0,1] existence_m1_subset_1
s: 82
d: [2,166,1,2,1,1,40,73,147]
a: [3,175,1,2,1,166]
s: 175
i: [1,74,0,5,2,0,0,255] t6_boole
s: 74
i: [1,81,0,5,2,0,0,2] t7_boole
s: 81
d: [2,177,0,5,2,1,22,74,137]
s: 177
i: [1,78,0,6,2,0,0,1] t1_subset
s: 78
i: [1,83,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 83
i: [1,80,0,7,3,0,0,1] t8_boole
s: 80
i: [1,79,0,8,3,0,0,1] t2_subset
s: 79
d: [2,181,1,5,2,1,40,79,107]
d: [2,184,1,3,1,2,21,181,122]
a: [3,188,1,3,1,184]
s: 188
i: [1,77,0,21,6,0,0,1] t89_mcart_1
s: 77
d: [2,182,1,5,2,1,40,79,112]
d: [2,190,1,3,1,2,21,182,122]
a: [3,194,1,3,1,190]
s: 194
d: [2,210,2,0,0,6,52,117,92,188,97,102,194,77]
e: 210
d: [2,211,2,0,0,6,52,117,188,92,97,102,194,77]
e: 211
f: 211,190,184,63,64,65,66,67,68,69
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t90_mcart_1
% SZS output start Proof for mcart_1__t90_mcart_1
1. ! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => ((k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2)) => X1 = X2)))) [input t90_mcart_1]
2. ~! [X0] : ((v1_relat_1(X0) & ~v1_xboole_0(X0)) => ! [X1] : (m1_subset_1(X1,X0) => ! [X2] : (m1_subset_1(X2,X0) => ((k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2)) => X1 = X2)))) [negated conjecture 1]
20. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
24. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : ((k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) => X0 = X2)) [input t89_mcart_1]
25. ! [X0,X1] : (v1_relat_1(X1) => ! [X2] : ((k2_xtuple_0(X2) = k2_xtuple_0(X0) & k1_xtuple_0(X2) = k1_xtuple_0(X0) & r2_hidden(X0,X1) & r2_hidden(X2,X1)) => X0 = X2)) [flattening 24]
27. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 20]
33. ? [X0] : (? [X1] : (? [X2] : ((X1 != X2 & (k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2))) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & (v1_relat_1(X0) & ~v1_xboole_0(X0))) [ennf transformation 2]
34. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & v1_relat_1(X0) & ~v1_xboole_0(X0)) [flattening 33]
41. ! [X0,X1] : (! [X2] : (X0 = X2 | (k2_xtuple_0(X2) != k2_xtuple_0(X0) | k1_xtuple_0(X2) != k1_xtuple_0(X0) | ~r2_hidden(X0,X1) | ~r2_hidden(X2,X1))) | ~v1_relat_1(X1)) [ennf transformation 25]
42. ! [X0,X1] : (! [X2] : (X0 = X2 | k2_xtuple_0(X2) != k2_xtuple_0(X0) | k1_xtuple_0(X2) != k1_xtuple_0(X0) | ~r2_hidden(X0,X1) | ~r2_hidden(X2,X1)) | ~v1_relat_1(X1)) [flattening 41]
44. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 27]
45. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 44]
49. ? [X0] : (? [X1] : (? [X2] : (X1 != X2 & k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2) & m1_subset_1(X2,X0)) & m1_subset_1(X1,X0)) & v1_relat_1(X0) & ~v1_xboole_0(X0)) => (? [X1] : (? [X2] : (X1 != X2 & k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2) & m1_subset_1(X2,sK0)) & m1_subset_1(X1,sK0)) & v1_relat_1(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
50. ? [X1] : (? [X2] : (X1 != X2 & k2_xtuple_0(X1) = k2_xtuple_0(X2) & k1_xtuple_0(X1) = k1_xtuple_0(X2) & m1_subset_1(X2,sK0)) & m1_subset_1(X1,sK0)) => (? [X2] : (sK1 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(sK1) & k1_xtuple_0(X2) = k1_xtuple_0(sK1) & m1_subset_1(X2,sK0)) & m1_subset_1(sK1,sK0)) [choice axiom]
51. ? [X2] : (sK1 != X2 & k2_xtuple_0(X2) = k2_xtuple_0(sK1) & k1_xtuple_0(X2) = k1_xtuple_0(sK1) & m1_subset_1(X2,sK0)) => (sK1 != sK2 & k2_xtuple_0(sK1) = k2_xtuple_0(sK2) & k1_xtuple_0(sK1) = k1_xtuple_0(sK2) & m1_subset_1(sK2,sK0)) [choice axiom]
52. ((sK1 != sK2 & k2_xtuple_0(sK1) = k2_xtuple_0(sK2) & k1_xtuple_0(sK1) = k1_xtuple_0(sK2) & m1_subset_1(sK2,sK0)) & m1_subset_1(sK1,sK0)) & v1_relat_1(sK0) & ~v1_xboole_0(sK0) [skolemisation 34,51,50,49]
63. ~v1_xboole_0(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
64. v1_relat_1(sK0) [cnf transformation 52] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
65. m1_subset_1(sK1,sK0) [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
66. m1_subset_1(sK2,sK0) [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
67. k1_xtuple_0(sK1) = k1_xtuple_0(sK2) [cnf transformation 52] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
68. k2_xtuple_0(sK1) = k2_xtuple_0(sK2) [cnf transformation 52] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
69. sK1 != sK2 [cnf transformation 52] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
77. k2_xtuple_0(X2) != k2_xtuple_0(X0) | X0 = X2 | k1_xtuple_0(X2) != k1_xtuple_0(X0) | ~r2_hidden(X0,X1) | ~r2_hidden(X2,X1) | ~v1_relat_1(X1) [cnf transformation 42] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
79. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 45] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
90. 1 <=> sK1 = sK2 [avatar definition]
92. sK1 != sK2 <- (~1) [avatar component clause 90] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
93. ~1 [avatar split clause 69,90]
95. 2 <=> k2_xtuple_0(sK1) = k2_xtuple_0(sK2) [avatar definition]
97. k2_xtuple_0(sK1) = k2_xtuple_0(sK2) <- (2) [avatar component clause 95] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
98. 2 [avatar split clause 68,95]
100. 3 <=> k1_xtuple_0(sK1) = k1_xtuple_0(sK2) [avatar definition]
102. k1_xtuple_0(sK1) = k1_xtuple_0(sK2) <- (3) [avatar component clause 100] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 3 [avatar split clause 67,100]
105. 4 <=> m1_subset_1(sK2,sK0) [avatar definition]
107. m1_subset_1(sK2,sK0) <- (4) [avatar component clause 105] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
108. 4 [avatar split clause 66,105]
110. 5 <=> m1_subset_1(sK1,sK0) [avatar definition]
112. m1_subset_1(sK1,sK0) <- (5) [avatar component clause 110] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. 5 [avatar split clause 65,110]
115. 6 <=> v1_relat_1(sK0) [avatar definition]
117. v1_relat_1(sK0) <- (6) [avatar component clause 115] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
118. 6 [avatar split clause 64,115]
120. 7 <=> v1_xboole_0(sK0) [avatar definition]
122. ~v1_xboole_0(sK0) <- (~7) [avatar component clause 120] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
123. ~7 [avatar split clause 63,120]
181. v1_xboole_0(sK0) | r2_hidden(sK2,sK0) <- (4) [resolution 79,107] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
182. v1_xboole_0(sK0) | r2_hidden(sK1,sK0) <- (5) [resolution 79,112] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
184. r2_hidden(sK2,sK0) <- (4, ~7) [subsumption resolution 181,122] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
186. 18 <=> r2_hidden(sK2,sK0) [avatar definition]
188. r2_hidden(sK2,sK0) <- (18) [avatar component clause 186] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
189. 18 | ~4 | 7 [avatar split clause 184,120,105,186]
190. r2_hidden(sK1,sK0) <- (5, ~7) [subsumption resolution 182,122] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
192. 19 <=> r2_hidden(sK1,sK0) [avatar definition]
194. r2_hidden(sK1,sK0) <- (19) [avatar component clause 192] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
195. 19 | ~5 | 7 [avatar split clause 190,120,110,192]
211. $false <- (~1, 2, 3, 6, 18, 19) [unit resulting resolution 117,188,92,97,102,194,77] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:27,thDist:-27}
216. 1 | ~2 | ~3 | ~6 | ~18 | ~19 [avatar contradiction clause 211]
217. $false [avatar sat refutation 93,98,103,108,113,118,123,189,195,216] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:143,thDist:-143}
% SZS output end Proof for mcart_1__t90_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 23
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 26
% Generated clauses: 72
% Active clauses: 28
% Passive clauses: 34
% Final active clauses: 27
% Final passive clauses: 5
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 2
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 2
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 15
% SAT solver binary clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

