i: [1,45,0,2,1,1,0,0] t36_orders_1
a: [3,82,0,2,1,45]
s: 82
i: [1,46,0,3,1,1,0,0] t36_orders_1
a: [3,77,0,3,1,46]
s: 77
i: [1,47,0,4,1,1,0,0] t36_orders_1
a: [3,72,0,4,1,47]
s: 72
i: [1,60,0,4,1,0,0,255] fc6_relat_1
s: 60
i: [1,61,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 61
i: [1,48,0,6,3,0,0,1] d4_orders_1
s: 48
i: [1,49,0,6,3,0,0,1] d4_orders_1
s: 49
i: [1,50,0,6,3,0,0,1] d4_orders_1
s: 50
i: [1,63,0,6,2,0,0,1] dt_k2_wellord1
s: 63
i: [1,64,0,6,2,0,0,3] fc1_relat_1
s: 64
i: [1,62,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 62
d: [2,85,1,6,2,0,44,64,62]
s: 85
i: [1,56,0,8,3,0,0,1] d7_orders_1
s: 56
d: [2,90,1,5,2,1,40,56,77]
d: [2,91,1,3,1,2,21,90,82]
a: [3,95,1,3,1,91]
s: 95
i: [1,57,0,8,3,0,0,1] d7_orders_1
s: 57
i: [1,58,0,8,3,0,0,1] d7_orders_1
s: 58
d: [2,97,1,5,2,1,40,57,77]
d: [2,98,1,3,1,2,21,97,82]
a: [3,102,1,3,1,98]
s: 102
i: [1,53,0,9,3,0,0,2] l47_orders_1
s: 53
d: [2,104,1,5,2,1,40,58,77]
d: [2,105,1,3,1,2,21,104,82]
a: [3,109,1,3,1,105]
s: 109
i: [1,54,0,9,3,0,0,2] l46_orders_1
s: 54
i: [1,55,0,9,3,0,0,2] l48_orders_1
s: 55
d: [2,86,1,6,2,0,44,64,62]
s: 86
i: [1,51,0,10,5,0,0,1] d4_orders_1
s: 51
i: [1,67,0,10,4,0,0,2] fc18_relat_2
s: 67
i: [1,68,0,10,4,0,0,2] fc19_relat_2
s: 68
i: [1,52,0,11,2,0,0,255] d6_wellord1
s: 52
d: [2,130,1,14,5,0,44,67,52]
d: [2,136,1,12,4,0,18,130]
d: [2,141,1,8,3,0,21,136,60]
s: 141
i: [1,65,0,12,5,0,0,2] fc5_relat_2
s: 65
d: [2,125,1,13,2,0,44,52,61]
d: [2,138,1,9,1,0,21,125,60]
s: 138
i: [1,66,0,12,5,0,0,2] fc22_relat_2
s: 66
i: [1,59,0,14,5,0,0,1] d7_orders_1
s: 59
d: [2,128,1,13,2,0,44,61,52]
d: [2,139,1,9,1,0,21,128,60]
s: 139
d: [2,115,1,10,4,0,44,67,62]
s: 115
d: [2,152,2,13,3,0,44,55,138]
d: [2,156,2,9,2,0,21,152,60]
s: 156
d: [2,116,1,10,4,0,44,67,62]
s: 116
d: [2,129,1,16,5,0,44,68,52]
d: [2,137,1,14,4,0,18,129]
d: [2,140,1,10,3,0,21,137,60]
s: 140
d: [2,153,2,13,3,0,44,54,138]
d: [2,157,2,9,2,0,21,153,60]
s: 157
d: [2,126,1,11,2,0,44,52,62]
s: 126
d: [2,154,2,13,3,0,44,53,138]
d: [2,158,2,9,2,0,21,154,60]
s: 158
d: [2,127,1,11,2,0,44,52,62]
s: 127
d: [2,132,1,11,2,0,44,62,52]
s: 132
d: [2,133,1,11,2,0,44,62,52]
s: 133
d: [2,147,1,18,6,0,44,65,52]
d: [2,148,1,16,5,0,18,147]
d: [2,150,1,12,4,0,21,148,60]
s: 150
d: [2,170,2,19,3,0,44,126,52]
d: [2,186,2,15,2,0,21,170,60]
d: [2,187,2,11,1,0,21,186,60]
s: 187
d: [2,162,1,18,6,0,44,66,52]
d: [2,163,1,16,5,0,18,162]
d: [2,165,1,12,4,0,21,163,60]
s: 165
d: [2,111,1,21,6,0,40,51,55]
d: [2,112,1,17,5,0,21,111,63]
s: 112
d: [2,196,3,10,2,0,44,63,187]
d: [2,206,3,6,1,0,21,196,60]
s: 206
d: [2,113,1,22,7,0,40,67,51]
d: [2,118,1,18,6,0,21,113,85]
s: 118
d: [2,201,3,10,2,0,44,63,187]
d: [2,211,3,6,1,0,21,201,60]
s: 211
d: [2,119,1,22,7,0,40,68,51]
d: [2,124,1,18,6,0,21,119,85]
s: 124
d: [2,180,2,19,3,0,44,52,126]
d: [2,190,2,15,2,0,21,180,60]
d: [2,191,2,11,1,0,21,190,60]
s: 191
d: [2,192,3,14,3,0,44,141,187]
d: [2,202,3,10,2,0,21,192,60]
s: 202
d: [2,175,2,18,6,0,44,65,126]
d: [2,185,2,16,5,0,18,175]
d: [2,188,2,12,4,0,21,185,60]
s: 188
d: [2,197,3,14,3,0,44,141,187]
d: [2,207,3,10,2,0,21,197,60]
s: 207
d: [2,176,2,18,6,0,44,66,126]
d: [2,184,2,16,5,0,18,176]
d: [2,189,2,12,4,0,21,184,60]
s: 189
d: [2,142,2,20,6,0,40,141,51]
d: [2,143,2,16,5,0,21,142,63]
s: 143
d: [2,249,4,10,2,0,44,202,187]
s: 249
d: [2,212,2,18,6,0,40,112,53]
d: [2,216,2,16,5,0,18,212]
s: 216
d: [2,250,4,10,2,0,44,202,187]
s: 250
d: [2,227,2,22,7,0,44,118,52]
d: [2,230,2,20,6,0,18,227]
d: [2,233,2,16,5,0,21,230,60]
s: 233
d: [2,240,2,22,7,0,44,124,126]
d: [2,241,2,20,6,0,18,240]
d: [2,246,2,16,5,0,21,241,60]
s: 246
d: [2,193,3,15,3,0,44,55,187]
d: [2,203,3,11,2,0,21,193,60]
s: 203
d: [2,213,2,21,5,0,44,112,138]
d: [2,217,2,17,4,0,21,213,60]
s: 217
d: [2,194,3,15,3,0,44,54,187]
d: [2,204,3,11,2,0,21,194,60]
s: 204
d: [2,225,2,18,6,0,44,118,62]
s: 225
d: [2,226,2,18,6,0,44,118,62]
s: 226
d: [2,195,3,15,3,0,44,53,187]
d: [2,205,3,11,2,0,21,195,60]
s: 205
d: [2,223,2,22,9,0,40,118,66]
d: [2,232,2,18,7,0,18,223]
s: 232
d: [2,198,3,15,3,0,44,55,187]
d: [2,208,3,11,2,0,21,198,60]
s: 208
d: [2,228,2,24,7,0,44,118,126]
d: [2,229,2,22,6,0,18,228]
d: [2,234,2,18,5,0,21,229,60]
s: 234
d: [2,235,2,22,9,0,40,124,66]
d: [2,244,2,18,7,0,18,235]
s: 244
d: [2,199,3,15,3,0,44,54,187]
d: [2,209,3,11,2,0,21,199,60]
s: 209
d: [2,239,2,24,7,0,44,124,52]
d: [2,242,2,22,6,0,18,239]
d: [2,245,2,18,5,0,21,242,60]
s: 245
d: [2,200,3,15,3,0,44,53,187]
d: [2,210,3,11,2,0,21,200,60]
s: 210
d: [2,252,3,17,6,0,40,143,53]
d: [2,256,3,15,5,0,18,252]
s: 256
d: [2,264,3,17,6,0,40,216,54]
d: [2,268,3,15,5,0,18,264]
s: 268
d: [2,323,4,0,0,5,52,82,95,102,109,72,268]
e: 323
d: [2,324,4,11,4,1,40,268,72]
d: [2,332,4,8,3,2,21,324,109]
d: [2,333,4,6,2,3,21,332,82]
d: [2,334,4,3,1,4,21,333,102]
d: [2,335,4,0,0,5,21,334,95]
e: 335
f: 335,105,98,91,45,46,47
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t36_orders_1
% SZS output start Proof for orders_1__t36_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_orders_1(X0,X1) => v2_orders_1(k2_wellord1(X0,X1)))) [input t36_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_orders_1(X0,X1) => v2_orders_1(k2_wellord1(X0,X1)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d4_orders_1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r2_orders_1(X0,X1) <=> (r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d7_orders_1]
7. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k2_wellord1(X0,X1))) [input dt_k2_wellord1]
17. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_relat_2(X0,X1) => v1_relat_2(k2_wellord1(X0,X1)))) [input l46_orders_1]
18. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r8_relat_2(X0,X1) => v8_relat_2(k2_wellord1(X0,X1)))) [input l47_orders_1]
19. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_relat_2(X0,X1) => v4_relat_2(k2_wellord1(X0,X1)))) [input l48_orders_1]
21. ? [X0] : (? [X1] : (~v2_orders_1(k2_wellord1(X0,X1)) & r2_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
22. ! [X0] : ((v2_orders_1(X0) <=> (v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
24. ! [X0] : (! [X1] : (v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 18]
25. ! [X0] : (! [X1] : (v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 17]
26. ! [X0] : (! [X1] : (v4_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 19]
27. ! [X0] : (! [X1] : (r2_orders_1(X0,X1) <=> (r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 6]
28. ! [X0,X1] : (v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
38. ? [X0] : (? [X1] : (~v2_orders_1(k2_wellord1(X0,X1)) & r2_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (~v2_orders_1(k2_wellord1(sK0,X1)) & r2_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
39. ? [X1] : (~v2_orders_1(k2_wellord1(sK0,X1)) & r2_orders_1(sK0,X1)) => (~v2_orders_1(k2_wellord1(sK0,sK1)) & r2_orders_1(sK0,sK1)) [choice axiom]
40. (~v2_orders_1(k2_wellord1(sK0,sK1)) & r2_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 21,39,38]
41. ! [X0] : (((v2_orders_1(X0) | (~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 22]
42. ! [X0] : (((v2_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v2_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 41]
43. ! [X0] : (! [X1] : ((r2_orders_1(X0,X1) | (~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 27]
44. ! [X0] : (! [X1] : ((r2_orders_1(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r2_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 43]
45. v1_relat_1(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
46. r2_orders_1(sK0,sK1) [cnf transformation 40] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
47. ~v2_orders_1(k2_wellord1(sK0,sK1)) [cnf transformation 40] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
51. ~v4_relat_2(X0) | v2_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 42] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
54. v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
55. v4_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
56. ~r2_orders_1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. ~r2_orders_1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. ~r2_orders_1(X0,X1) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
70. 1 <=> v2_orders_1(k2_wellord1(sK0,sK1)) [avatar definition]
72. ~v2_orders_1(k2_wellord1(sK0,sK1)) <- (~1) [avatar component clause 70] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
73. ~1 [avatar split clause 47,70]
75. 2 <=> r2_orders_1(sK0,sK1) [avatar definition]
77. r2_orders_1(sK0,sK1) <- (2) [avatar component clause 75] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. 2 [avatar split clause 46,75]
80. 3 <=> v1_relat_1(sK0) [avatar definition]
82. v1_relat_1(sK0) <- (3) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 3 [avatar split clause 45,80]
90. r1_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 56,77] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
91. r1_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 90,82] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
93. 4 <=> r1_relat_2(sK0,sK1) [avatar definition]
95. r1_relat_2(sK0,sK1) <- (4) [avatar component clause 93] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
96. 4 | ~2 | ~3 [avatar split clause 91,80,75,93]
97. r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 57,77] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. r8_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 97,82] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. 5 <=> r8_relat_2(sK0,sK1) [avatar definition]
102. r8_relat_2(sK0,sK1) <- (5) [avatar component clause 100] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 5 | ~2 | ~3 [avatar split clause 98,80,75,100]
104. r4_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 58,77] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
105. r4_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 104,82] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
107. 6 <=> r4_relat_2(sK0,sK1) [avatar definition]
109. r4_relat_2(sK0,sK1) <- (6) [avatar component clause 107] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 6 | ~2 | ~3 [avatar split clause 105,80,75,107]
111. v2_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 51,55] {a:1,w:21,wCS:8400,sine:1,thAx:0,allAx:2,thDist:-2}
112. ~v8_relat_2(k2_wellord1(X0,X1)) | v2_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [subsumption resolution 111,63] {a:1,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
212. v2_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 112,53] {a:2,w:18,wCS:7200,sine:1,thAx:0,allAx:3,thDist:-3}
216. ~v1_relat_2(k2_wellord1(X0,X1)) | v2_orders_1(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r8_relat_2(X0,X1) [duplicate literal removal 212] {a:2,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
264. v2_orders_1(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 216,54] {a:3,w:17,wCS:6800,sine:1,thAx:0,allAx:4,thDist:-4}
268. v2_orders_1(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) [duplicate literal removal 264] {a:3,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
324. ~r4_relat_2(sK0,sK1) | ~v1_relat_1(sK0) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1) [resolution 268,72] {a:4,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
332. ~v1_relat_1(sK0) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 6) [subsumption resolution 324,109] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
333. ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3, 6) [subsumption resolution 332,82] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
334. ~r1_relat_2(sK0,sK1) <- (~1, 3, 5, 6) [subsumption resolution 333,102] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
335. $false <- (~1, 3, 4, 5, 6) [subsumption resolution 334,95] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
336. 1 | ~3 | ~4 | ~5 | ~6 [avatar contradiction clause 335]
340. $false [avatar sat refutation 73,78,83,96,103,110,336] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:62,thDist:-62}
% SZS output end Proof for orders_1__t36_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 18
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 269
% Active clauses: 55
% Passive clauses: 108
% Final active clauses: 55
% Final passive clauses: 32
% Main loop iterations started: 76
% 
% >>> Simplifying Inferences
% Duplicate literals: 42
% Fw subsumption resolutions: 65
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 62
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 1
% Forward superposition: 78
% Backward superposition: 34
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

