i: [1,50,0,2,1,1,0,0] t32_sysrel
a: [3,97,0,2,1,50]
s: 97
i: [1,51,0,24,4,1,0,0] t32_sysrel
a: [3,80,0,6,1,51]
s: 80
i: [1,52,0,3,1,0,0,2] dt_k4_relat_1
s: 52
i: [1,55,0,3,1,0,0,255] fc28_relat_1
s: 55
i: [1,68,0,3,1,0,0,1] reflexivity_r1_tarski
s: 68
i: [1,56,0,4,1,0,0,255] fc28_relat_1
s: 56
i: [1,57,0,4,1,0,0,255] fc28_relat_1
s: 57
i: [1,67,0,4,1,0,0,255] existence_m1_subset_1
s: 67
i: [1,74,0,4,1,0,0,255] rc3_relat_1
s: 74
i: [1,53,0,5,1,0,0,1] rd2_relat_1
s: 53
i: [1,54,0,5,1,0,0,1] rd1_relat_1
s: 54
i: [1,58,0,5,2,0,0,1] dt_k1_sysrel
s: 58
i: [1,64,0,5,1,0,0,1] t17_xboole_1
s: 64
i: [1,69,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 69
i: [1,75,0,5,1,0,0,255] rc3_relat_1
s: 75
i: [1,76,0,5,1,0,0,255] rc3_relat_1
s: 76
i: [1,71,0,6,2,0,0,2] fc1_relat_1
s: 71
i: [1,65,0,7,2,0,0,255] t3_subset
s: 65
d: [2,101,1,5,1,0,40,65,67]
s: 101
i: [1,66,0,7,2,0,0,255] t3_subset
s: 66
i: [1,70,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 70
d: [2,104,1,5,1,0,44,64,70]
s: 104
i: [1,63,0,8,3,0,0,255] cc2_relat_1
s: 63
d: [2,106,1,5,1,0,44,64,70]
s: 106
i: [1,59,0,9,2,0,0,1] t26_sysrel
s: 59
i: [1,60,0,10,2,0,0,1] d1_sysrel
s: 60
d: [2,103,1,6,2,0,44,71,70]
s: 103
i: [1,61,0,12,4,0,0,1] t11_relat_1
s: 61
d: [2,123,1,9,3,1,40,61,80]
d: [2,126,1,7,2,2,21,123,97]
a: [3,130,1,3,1,126]
s: 130
d: [2,138,2,0,0,2,52,97,130,58]
e: 138
d: [2,139,2,2,1,1,40,130,58]
d: [2,141,2,0,0,2,21,139,97]
e: 141
r: 130
i: [1,62,0,12,4,0,0,1] t11_relat_1
s: 62
i: [1,72,0,12,4,0,0,255] cc6_relat_1
s: 72
a: [3,129,1,3,1,126]
s: 129
i: [1,73,0,12,4,0,0,255] cc5_relat_1
s: 73
a: [3,134,1,4,1,126]
s: 134
d: [2,105,1,6,2,0,44,71,70]
s: 105
d: [2,110,1,6,2,0,40,63,67]
s: 110
d: [2,115,1,6,2,0,44,64,60]
s: 115
d: [2,155,2,0,0,2,52,97,134,115]
e: 155
d: [2,156,2,2,1,1,40,115,134]
d: [2,158,2,0,0,2,21,156,97]
e: 158
r: 80
r: 134
a: [3,79,0,6,1,51]
s: 79
a: [3,133,1,4,1,126]
s: 133
a: [3,84,0,6,1,51]
s: 84
d: [2,160,2,0,0,4,52,97,129,133,84,62]
e: 160
d: [2,161,1,9,3,1,40,84,62]
d: [2,164,1,5,2,2,21,161,133]
d: [2,165,1,3,1,3,21,164,97]
d: [2,166,1,0,0,4,21,165,129]
e: 166
r: 84
a: [3,83,0,6,1,51]
s: 83
a: [3,88,0,6,1,51]
s: 88
d: [2,172,1,8,2,1,44,88,59]
d: [2,173,1,2,1,2,21,172,79]
d: [2,174,1,0,0,3,21,173,97]
e: 174
d: [2,169,1,8,2,1,44,83,59]
d: [2,170,1,6,1,2,21,169,97]
f: 174,170,166,158,141,126,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for sysrel__t32_sysrel
% SZS output start Proof for sysrel__t32_sysrel
1. ! [X0] : (v1_relat_1(X0) => (r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)))) [input t32_sysrel]
2. ~! [X0] : (v1_relat_1(X0) => (r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) & r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) & r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)))) [negated conjecture 1]
7. ! [X0] : (v1_relat_1(X0) => k1_sysrel(X0) = k3_xboole_0(X0,k4_relat_1(k9_xtuple_0(X0)))) [input d1_sysrel]
9. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k1_sysrel(X0))) [input dt_k1_sysrel]
23. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))))) [input t11_relat_1]
24. ! [X0] : ! [X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [input t17_xboole_1]
25. ! [X0] : (v1_relat_1(X0) => k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0))) [input t26_sysrel]
27. ! [X0,X1] : r1_tarski(k3_xboole_0(X0,X1),X0) [flattening 24]
31. ? [X0] : ((~r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) | ~r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0))) & v1_relat_1(X0)) [ennf transformation 2]
32. ! [X0] : (v1_relat_1(k1_sysrel(X0)) | ~v1_relat_1(X0)) [ennf transformation 9]
33. ! [X0] : (k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0)) | ~v1_relat_1(X0)) [ennf transformation 25]
34. ! [X0] : (k1_sysrel(X0) = k3_xboole_0(X0,k4_relat_1(k9_xtuple_0(X0))) | ~v1_relat_1(X0)) [ennf transformation 7]
35. ! [X0] : (! [X1] : (((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 23]
36. ! [X0] : (! [X1] : ((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 35]
43. ? [X0] : ((~r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(X0)),k10_xtuple_0(X0)) | ~r1_tarski(k9_xtuple_0(k1_sysrel(X0)),k9_xtuple_0(X0))) & v1_relat_1(X0)) => ((~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0))) & v1_relat_1(sK0)) [choice axiom]
44. (~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0))) & v1_relat_1(sK0) [skolemisation 31,43]
50. v1_relat_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) | ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) [cnf transformation 44] {a:0,w:24,wCS:2400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
58. v1_relat_1(k1_sysrel(X0)) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
59. k9_xtuple_0(k1_sysrel(X0)) = k10_xtuple_0(k1_sysrel(X0)) | ~v1_relat_1(X0) [cnf transformation 33] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
60. k1_sysrel(X0) = k3_xboole_0(X0,k4_relat_1(k9_xtuple_0(X0))) | ~v1_relat_1(X0) [cnf transformation 34] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. r1_tarski(k3_xboole_0(X0,X1),X0) [cnf transformation 27] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
78. 1 <=> r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) [avatar definition]
79. r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) <- (1) [avatar component clause 78] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) <- (~1) [avatar component clause 78] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 2 <=> r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) [avatar definition]
83. r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) <- (2) [avatar component clause 82] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) <- (~2) [avatar component clause 82] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. 3 <=> r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) [avatar definition]
88. ~r1_tarski(k10_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) <- (~3) [avatar component clause 86] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
90. 4 <=> r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) [avatar definition]
93. ~1 | ~2 | ~3 | ~4 [avatar split clause 51,90,86,82,78]
95. 5 <=> v1_relat_1(sK0) [avatar definition]
97. v1_relat_1(sK0) <- (5) [avatar component clause 95] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
98. 5 [avatar split clause 50,95]
115. r1_tarski(k1_sysrel(X4),X4) | ~v1_relat_1(X4) [superposition 64,60] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
123. ~r1_tarski(k1_sysrel(sK0),sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(k1_sysrel(sK0)) <- (~1) [resolution 61,80] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
126. ~r1_tarski(k1_sysrel(sK0),sK0) | ~v1_relat_1(k1_sysrel(sK0)) <- (~1, 5) [subsumption resolution 123,97] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
128. 6 <=> v1_relat_1(k1_sysrel(sK0)) [avatar definition]
129. v1_relat_1(k1_sysrel(sK0)) <- (6) [avatar component clause 128] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
130. ~v1_relat_1(k1_sysrel(sK0)) <- (~6) [avatar component clause 128] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
132. 7 <=> r1_tarski(k1_sysrel(sK0),sK0) [avatar definition]
133. r1_tarski(k1_sysrel(sK0),sK0) <- (7) [avatar component clause 132] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
134. ~r1_tarski(k1_sysrel(sK0),sK0) <- (~7) [avatar component clause 132] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
135. ~6 | ~7 | 1 | ~5 [avatar split clause 126,95,78,132,128]
139. ~v1_relat_1(sK0) <- (~6) [resolution 130,58] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
141. $false <- (5, ~6) [subsumption resolution 139,97] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
142. ~5 | 6 [avatar contradiction clause 141]
156. ~v1_relat_1(sK0) <- (~7) [resolution 115,134] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. $false <- (5, ~7) [subsumption resolution 156,97] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
159. ~5 | 7 [avatar contradiction clause 158]
161. ~r1_tarski(k1_sysrel(sK0),sK0) | ~v1_relat_1(sK0) | ~v1_relat_1(k1_sysrel(sK0)) <- (~2) [resolution 84,62] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
164. ~v1_relat_1(sK0) | ~v1_relat_1(k1_sysrel(sK0)) <- (~2, 7) [subsumption resolution 161,133] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
165. ~v1_relat_1(k1_sysrel(sK0)) <- (~2, 5, 7) [subsumption resolution 164,97] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
166. $false <- (~2, 5, 6, 7) [subsumption resolution 165,129] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. 2 | ~5 | ~6 | ~7 [avatar contradiction clause 166]
169. r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (2) [superposition 83,59] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
170. r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k10_xtuple_0(sK0)) <- (2, 5) [subsumption resolution 169,97] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
171. 4 | ~2 | ~5 [avatar split clause 170,95,82,90]
172. ~r1_tarski(k9_xtuple_0(k1_sysrel(sK0)),k9_xtuple_0(sK0)) | ~v1_relat_1(sK0) <- (~3) [superposition 88,59] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
173. ~v1_relat_1(sK0) <- (1, ~3) [subsumption resolution 172,79] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
174. $false <- (1, ~3, 5) [subsumption resolution 173,97] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
175. ~1 | 3 | ~5 [avatar contradiction clause 174]
176. $false [avatar sat refutation 93,98,135,142,159,167,171,175] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:61,thDist:-61}
% SZS output end Proof for sysrel__t32_sysrel
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 98
% Active clauses: 39
% Passive clauses: 56
% Final active clauses: 35
% Final passive clauses: 14
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 17
% Fw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 11
% 
% >>> Generating Inferences
% Binary resolution: 12
% Unit resulting resolution: 3
% Forward superposition: 18
% Backward superposition: 8
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 6
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 1
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

