i: [1,119,0,4,1,1,0,0] l3_arithm
a: [3,173,0,4,1,119]
s: 173
i: [1,126,0,2,1,0,0,255] fc6_ordinal1
a: [3,203,0,2,1,126]
s: 203
i: [1,125,0,2,1,0,0,255] fc6_ordinal1
a: [3,208,0,2,1,125]
s: 208
i: [1,159,0,2,1,0,0,255] rc4_ordinal1
a: [3,221,0,2,1,159]
s: 221
i: [1,160,0,2,1,0,0,255] rc1_ordinal1
a: [3,226,0,2,1,160]
s: 226
i: [1,164,0,2,1,0,0,255] rc2_ordinal1
a: [3,231,0,2,1,164]
s: 231
i: [1,163,0,2,1,0,0,255] rc2_ordinal1
a: [3,236,0,2,1,163]
s: 236
i: [1,162,0,2,1,0,0,255] rc2_ordinal1
a: [3,241,0,2,1,162]
s: 241
i: [1,161,0,2,1,0,0,255] rc2_ordinal1
a: [3,246,0,2,1,161]
s: 246
i: [1,165,0,2,1,0,0,255] rc2_xcmplx_0
a: [3,251,0,2,1,165]
s: 251
i: [1,167,0,2,1,0,0,255] rc5_ordinal1
a: [3,256,0,2,1,167]
s: 256
i: [1,166,0,2,1,0,0,255] rc5_ordinal1
a: [3,261,0,2,1,166]
s: 261
i: [1,151,0,3,1,0,0,255] reflexivity_r1_tarski
s: 151
i: [1,120,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,178,0,3,1,120]
s: 178
i: [1,121,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,183,0,3,1,121]
s: 183
i: [1,122,0,3,1,0,0,255] t1_numerals
a: [3,188,0,3,1,122]
s: 188
i: [1,168,0,11,3,0,0,2] t11_arytm_0
a: [3,213,0,3,1,168]
s: 213
d: [2,263,0,3,1,2,22,188,183]
a: [3,267,0,3,1,263]
s: 267
i: [1,128,0,4,2,0,0,255] cc6_ordinal1
s: 128
i: [1,129,0,4,2,0,0,255] cc1_ordinal1
s: 129
i: [1,130,0,4,2,0,0,255] cc1_ordinal1
s: 130
d: [2,269,1,2,1,1,40,128,221]
a: [3,274,1,2,1,269]
s: 274
i: [1,132,0,4,2,0,0,255] cc9_ordinal1
s: 132
d: [2,270,1,2,1,1,40,128,256]
a: [3,279,1,2,1,270]
s: 279
i: [1,133,0,4,2,0,0,255] cc7_ordinal1
s: 133
i: [1,134,0,4,2,0,0,255] cc3_ordinal1
s: 134
d: [2,281,1,2,1,1,40,129,203]
a: [3,287,1,2,1,281]
s: 287
i: [1,150,0,4,1,0,0,2] existence_m1_subset_1
s: 150
d: [2,282,1,2,1,1,40,129,226]
a: [3,292,1,2,1,282]
s: 292
i: [1,123,0,4,1,0,0,255] dt_k5_numbers
a: [3,193,0,4,1,123]
s: 193
i: [1,124,0,4,1,0,0,255] dt_k6_numbers
a: [3,198,0,4,1,124]
s: 198
d: [2,294,1,2,1,1,40,130,203]
a: [3,300,1,2,1,294]
s: 300
d: [2,331,0,4,1,2,22,193,178]
a: [3,335,0,4,1,331]
s: 335
d: [2,295,1,2,1,1,40,130,226]
a: [3,305,1,2,1,295]
s: 305
d: [2,337,0,4,1,2,22,198,178]
a: [3,341,0,4,1,337]
s: 341
i: [1,135,0,5,2,0,0,255] t6_boole
s: 135
d: [2,307,2,2,1,1,40,274,130]
a: [3,312,2,2,1,307]
s: 312
i: [1,136,0,5,2,0,0,255] cc8_ordinal1
s: 136
d: [2,344,1,2,1,1,40,136,267]
a: [3,349,1,2,1,344]
s: 349
i: [1,147,0,5,2,0,0,255] t7_boole
s: 147
d: [2,343,0,5,2,1,22,135,183]
s: 343
d: [2,367,1,6,2,2,44,173,343]
s: 367
i: [1,142,0,6,3,0,0,255] cc2_ordinal1
s: 142
d: [2,308,2,2,1,1,40,274,129]
a: [3,317,2,2,1,308]
s: 317
i: [1,143,0,6,2,0,0,255] t1_subset
s: 143
i: [1,152,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 152
d: [2,319,2,2,1,1,40,279,130]
a: [3,324,2,2,1,319]
s: 324
i: [1,131,0,7,3,0,0,255] cc5_ordinal1
s: 131
d: [2,320,2,2,1,1,40,279,129]
a: [3,329,2,2,1,320]
s: 329
i: [1,137,0,7,2,0,0,1] dt_k3_arytm_0
s: 137
i: [1,145,0,7,2,0,0,255] t3_subset
s: 145
d: [2,378,1,8,3,0,40,137,131]
a: [3,382,1,2,1,378]
s: 382
i: [1,146,0,7,3,0,0,1] t8_boole
s: 146
d: [2,356,2,2,1,1,40,349,128]
a: [3,360,2,2,1,356]
s: 360
i: [1,127,0,8,3,0,0,255] cc10_ordinal1
s: 127
i: [1,138,0,8,2,0,0,1] involutiveness_k3_arytm_0
s: 138
d: [2,402,1,4,2,1,40,127,335]
a: [3,407,1,2,1,402]
s: 407
i: [1,144,0,8,3,0,0,255] t2_subset
s: 144
d: [2,388,3,2,1,1,40,360,130]
a: [3,394,3,2,1,388]
s: 394
i: [1,149,0,9,3,0,0,255] t5_subset
s: 149
i: [1,148,0,10,3,0,0,255] t4_subset
s: 148
d: [2,417,1,9,3,0,40,144,137]
a: [3,431,1,2,1,417]
s: 431
i: [1,157,0,11,3,0,0,2] dt_k1_arytm_0
s: 157
d: [2,389,3,2,1,1,40,360,129]
a: [3,399,3,2,1,389]
s: 399
i: [1,158,0,13,3,0,0,2] commutativity_k1_arytm_0
s: 158
i: [1,169,0,13,3,0,0,1] d3_arytm_0
s: 169
d: [2,414,2,10,3,2,44,367,138]
s: 414
d: [2,463,0,9,2,0,21,169,137]
s: 463
d: [2,345,1,3,1,0,40,136,150]
a: [3,354,1,3,1,345]
s: 354
i: [1,156,0,14,4,0,0,255] existence_m2_subset_1
s: 156
i: [1,141,0,15,4,0,0,1] d3_arytm_0
s: 141
d: [2,416,1,5,2,1,40,144,267]
d: [2,420,1,3,1,2,21,416,208]
a: [3,424,1,3,1,420]
s: 424
d: [2,482,2,0,0,3,52,213,335,424,148]
e: 482
r: 213
a: [3,212,0,3,1,168]
s: 212
r: 431
a: [3,432,1,2,1,417]
s: 432
d: [2,512,2,0,0,2,52,407,432,132]
e: 512
d: [2,513,2,0,0,2,52,382,432,134]
e: 513
d: [2,514,2,0,0,3,52,424,335,432,149]
e: 514
d: [2,515,2,2,1,1,40,432,134]
d: [2,521,2,0,0,2,21,515,382]
e: 521
d: [2,517,2,2,1,1,40,432,132]
d: [2,523,2,0,0,2,21,517,407]
e: 523
r: 432
a: [3,216,0,8,2,168]
s: 216
i: [1,153,0,15,5,0,0,255] dt_m2_subset_1
s: 153
s: 431
i: [1,154,0,15,5,0,0,255] redefinition_m2_subset_1
s: 154
d: [2,446,1,5,2,1,40,149,335]
a: [3,450,1,3,1,446]
s: 450
d: [2,551,2,0,0,2,40,450,424]
e: 551
d: [2,549,2,0,0,2,52,424,450]
e: 549
d: [2,550,2,0,0,3,52,267,208,450,144]
e: 550
r: 450
i: [1,155,0,15,5,0,0,255] redefinition_m2_subset_1
s: 155
d: [2,498,1,5,2,1,40,212,144]
d: [2,501,1,3,1,2,21,498,431]
a: [3,505,1,3,1,501]
s: 505
d: [2,471,2,3,1,1,40,354,128]
a: [3,475,2,3,1,471]
s: 475
d: [2,418,1,7,2,1,40,144,335]
a: [3,437,1,4,1,418]
s: 437
a: [3,440,1,3,1,418]
s: 440
d: [2,365,1,5,2,2,44,267,343]
s: 365
d: [2,374,1,5,2,0,40,131,150]
s: 374
d: [2,483,2,3,1,1,40,424,152]
a: [3,496,2,3,1,483]
s: 496
d: [2,500,1,5,2,2,44,212,343]
s: 500
d: [2,560,2,3,1,1,40,505,152]
a: [3,567,2,3,1,560]
s: 567
d: [2,364,1,6,2,2,44,341,343]
s: 364
d: [2,403,1,6,2,0,40,127,150]
s: 403
d: [2,569,3,3,1,1,40,475,130]
a: [3,574,3,3,1,569]
s: 574
d: [2,419,1,6,2,0,40,144,150]
s: 419
d: [2,570,3,3,1,1,40,475,129]
a: [3,579,3,3,1,570]
s: 579
d: [2,453,1,6,2,1,40,148,335]
s: 453
d: [2,362,1,7,3,1,44,343,343]
s: 362
d: [2,581,2,4,1,1,40,437,152]
a: [3,587,2,4,1,581]
s: 587
d: [2,401,1,7,3,0,40,127,145]
s: 401
d: [2,608,2,6,2,1,40,453,419]
d: [2,609,2,4,1,2,21,608,208]
a: [3,613,2,4,1,609]
s: 613
a: [3,428,1,7,2,417]
s: 428
d: [2,447,1,7,2,0,40,149,150]
s: 447
d: [2,616,3,6,2,1,40,613,144]
d: [2,618,3,4,1,2,21,616,431]
a: [3,622,3,4,1,618]
s: 622
d: [2,556,1,11,4,1,40,155,335]
d: [2,558,1,9,3,2,21,556,208]
d: [2,559,1,7,2,3,21,558,431]
s: 559
d: [2,444,4,4,2,2,44,394,343]
s: 444
d: [2,387,1,8,3,0,40,145,131]
s: 387
d: [2,445,1,8,3,0,40,149,145]
s: 445
d: [2,458,4,4,2,2,44,399,343]
s: 458
d: [2,454,1,8,2,0,40,148,150]
s: 454
d: [2,630,4,4,1,1,40,622,152]
a: [3,636,4,4,1,630]
s: 636
d: [2,526,1,14,4,1,44,216,158]
d: [2,537,1,11,3,1,18,526]
d: [2,538,1,8,2,2,21,537,212]
s: 538
d: [2,527,1,14,4,1,44,216,158]
d: [2,536,1,11,3,1,18,527]
d: [2,539,1,8,2,2,21,536,212]
s: 539
d: [2,649,2,11,3,2,44,538,463]
d: [2,661,2,7,2,3,21,649,173]
d: [2,662,2,4,1,3,21,661,212]
a: [3,666,2,4,1,662]
s: 666
d: [2,672,3,0,0,2,52,212,666,137]
e: 672
d: [2,673,3,3,1,1,40,666,137]
d: [2,677,3,0,0,2,21,673,212]
e: 677
d: [2,654,2,11,3,2,44,463,538]
d: [2,669,2,7,2,3,21,654,173]
d: [2,670,2,4,1,3,21,669,212]
f: 677,670,482,420,331,263,168,125,123,122,121,120,119
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arithm__l3_arithm
% SZS output start Proof for arithm__l3_arithm
1. k6_numbers = k3_arytm_0(k6_numbers) [input l3_arithm]
2. ~k6_numbers = k3_arytm_0(k6_numbers) [negated conjecture 1]
16. ! [X0,X1] : ((m1_subset_1(X1,k1_numbers) & m1_subset_1(X0,k1_numbers)) => k1_arytm_0(X0,X1) = k1_arytm_0(X1,X0)) [input commutativity_k1_arytm_0]
17. ! [X0] : (m1_subset_1(X0,k1_numbers) => ! [X1] : (m1_subset_1(X1,k1_numbers) => (k3_arytm_0(X0) = X1 <=> k6_numbers = k1_arytm_0(X0,X1)))) [input d3_arytm_0]
22. ! [X0] : (m1_subset_1(X0,k1_numbers) => m1_subset_1(k3_arytm_0(X0),k1_numbers)) [input dt_k3_arytm_0]
24. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [input dt_k5_numbers]
30. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
38. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
39. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
42. ! [X0] : (m1_subset_1(X0,k1_numbers) => ! [X1] : (m1_subset_1(X1,k1_numbers) => (k6_numbers = X1 => k1_arytm_0(X0,X1) = X0))) [input t11_arytm_0]
43. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
45. ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [input t2_subset]
47. ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [input t4_subset]
52. k6_numbers != k3_arytm_0(k6_numbers) [flattening 2]
54. ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1))) [flattening 45]
58. ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2)) [flattening 47]
76. ! [X0] : (m1_subset_1(k3_arytm_0(X0),k1_numbers) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 22]
78. ! [X0] : (! [X1] : ((k1_arytm_0(X0,X1) = X0 | k6_numbers != X1) | ~m1_subset_1(X1,k1_numbers)) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 42]
79. ! [X0] : (! [X1] : (k1_arytm_0(X0,X1) = X0 | k6_numbers != X1 | ~m1_subset_1(X1,k1_numbers)) | ~m1_subset_1(X0,k1_numbers)) [flattening 78]
80. ! [X0] : (! [X1] : ((k3_arytm_0(X0) = X1 <=> k6_numbers = k1_arytm_0(X0,X1)) | ~m1_subset_1(X1,k1_numbers)) | ~m1_subset_1(X0,k1_numbers)) [ennf transformation 17]
84. ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1)) [ennf transformation 54]
85. ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1)) [flattening 84]
89. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))) [ennf transformation 58]
90. ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)) [flattening 89]
101. ! [X0,X1] : (k1_arytm_0(X0,X1) = k1_arytm_0(X1,X0) | (~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X0,k1_numbers))) [ennf transformation 16]
102. ! [X0,X1] : (k1_arytm_0(X0,X1) = k1_arytm_0(X1,X0) | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X0,k1_numbers)) [flattening 101]
103. ! [X0] : (! [X1] : (((k3_arytm_0(X0) = X1 | k6_numbers != k1_arytm_0(X0,X1)) & (k6_numbers = k1_arytm_0(X0,X1) | k3_arytm_0(X0) != X1)) | ~m1_subset_1(X1,k1_numbers)) | ~m1_subset_1(X0,k1_numbers)) [nnf transformation 80]
119. k6_numbers != k3_arytm_0(k6_numbers) [cnf transformation 52] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
120. k4_ordinal1 = k5_numbers [cnf transformation 38] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
121. k6_numbers = k1_xboole_0 [cnf transformation 39] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
122. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 43] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
123. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [cnf transformation 24] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
125. ~v1_xboole_0(k4_ordinal1) [cnf transformation 30] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
137. m1_subset_1(k3_arytm_0(X0),k1_numbers) | ~m1_subset_1(X0,k1_numbers) [cnf transformation 76] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. k1_arytm_0(X0,X1) = X0 | k6_numbers != X1 | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X0,k1_numbers) [cnf transformation 79] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:1,thDist:-1}
140. k6_numbers = k1_arytm_0(X0,X1) | k3_arytm_0(X0) != X1 | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X0,k1_numbers) [cnf transformation 103] {a:0,w:15,wCS:6000,sine:1,thAx:0,allAx:1,thDist:-1}
144. ~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1) [cnf transformation 85] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
148. ~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1) [cnf transformation 90] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
158. k1_arytm_0(X0,X1) = k1_arytm_0(X1,X0) | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X0,k1_numbers) [cnf transformation 102] {a:0,w:13,wCS:5200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
168. k1_arytm_0(X0,k6_numbers) = X0 | ~m1_subset_1(k6_numbers,k1_numbers) | ~m1_subset_1(X0,k1_numbers) [equality resolution 139] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:1,thDist:-1}
169. k6_numbers = k1_arytm_0(X0,k3_arytm_0(X0)) | ~m1_subset_1(k3_arytm_0(X0),k1_numbers) | ~m1_subset_1(X0,k1_numbers) [equality resolution 140] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
171. 1 <=> k6_numbers = k3_arytm_0(k6_numbers) [avatar definition]
173. k6_numbers != k3_arytm_0(k6_numbers) <- (~1) [avatar component clause 171] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
174. ~1 [avatar split clause 119,171]
176. 2 <=> k4_ordinal1 = k5_numbers [avatar definition]
178. k4_ordinal1 = k5_numbers <- (2) [avatar component clause 176] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
179. 2 [avatar split clause 120,176]
181. 3 <=> k6_numbers = k1_xboole_0 [avatar definition]
183. k6_numbers = k1_xboole_0 <- (3) [avatar component clause 181] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
184. 3 [avatar split clause 121,181]
186. 4 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
188. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (4) [avatar component clause 186] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
189. 4 [avatar split clause 122,186]
191. 5 <=> m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) [avatar definition]
193. m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)) <- (5) [avatar component clause 191] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
194. 5 [avatar split clause 123,191]
206. 8 <=> v1_xboole_0(k4_ordinal1) [avatar definition]
208. ~v1_xboole_0(k4_ordinal1) <- (~8) [avatar component clause 206] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
209. ~8 [avatar split clause 125,206]
211. 9 <=> m1_subset_1(k6_numbers,k1_numbers) [avatar definition]
212. m1_subset_1(k6_numbers,k1_numbers) <- (9) [avatar component clause 211] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
213. ~m1_subset_1(k6_numbers,k1_numbers) <- (~9) [avatar component clause 211] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
215. 10 <=> ! [X0] : (k1_arytm_0(X0,k6_numbers) = X0 | ~m1_subset_1(X0,k1_numbers)) [avatar definition]
216. k1_arytm_0(X0,k6_numbers) = X0 | ~m1_subset_1(X0,k1_numbers) <- (10) [avatar component clause 215] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
217. ~9 | 10 [avatar split clause 168,215,211]
263. m1_subset_1(k6_numbers,k4_ordinal1) <- (3, 4) [forward demodulation 188,183] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
265. 20 <=> m1_subset_1(k6_numbers,k4_ordinal1) [avatar definition]
267. m1_subset_1(k6_numbers,k4_ordinal1) <- (20) [avatar component clause 265] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
268. 20 | ~3 | ~4 [avatar split clause 263,186,181,265]
331. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (2, 5) [forward demodulation 193,178] {a:0,w:4,wCS:1600,sine:255,thAx:0,allAx:1,thDist:-1}
333. 31 <=> m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) [avatar definition]
335. m1_subset_1(k4_ordinal1,k1_zfmisc_1(k1_numbers)) <- (31) [avatar component clause 333] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
336. 31 | ~2 | ~5 [avatar split clause 331,191,176,333]
416. v1_xboole_0(k4_ordinal1) | r2_hidden(k6_numbers,k4_ordinal1) <- (20) [resolution 144,267] {a:1,w:5,wCS:2000,sine:255,thAx:0,allAx:2,thDist:-2}
420. r2_hidden(k6_numbers,k4_ordinal1) <- (~8, 20) [subsumption resolution 416,208] {a:1,w:3,wCS:1200,sine:255,thAx:0,allAx:2,thDist:-2}
422. 42 <=> r2_hidden(k6_numbers,k4_ordinal1) [avatar definition]
424. r2_hidden(k6_numbers,k4_ordinal1) <- (42) [avatar component clause 422] {a:1,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
425. 42 | 8 | ~20 [avatar split clause 420,265,206,422]
463. k6_numbers = k1_arytm_0(X0,k3_arytm_0(X0)) | ~m1_subset_1(X0,k1_numbers) [subsumption resolution 169,137] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
482. $false <- (~9, 31, 42) [unit resulting resolution 213,335,424,148] {a:2,w:0,sine:255,thAx:0,allAx:5,thDist:-5}
487. 9 | ~31 | ~42 [avatar contradiction clause 482]
526. k1_arytm_0(k6_numbers,X1) = X1 | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(k6_numbers,k1_numbers) <- (10) [superposition 216,158] {a:1,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
537. k1_arytm_0(k6_numbers,X1) = X1 | ~m1_subset_1(X1,k1_numbers) | ~m1_subset_1(k6_numbers,k1_numbers) <- (10) [duplicate literal removal 526] {a:1,w:11,wCS:4400,sine:2,thAx:0,allAx:2,thDist:-2}
538. k1_arytm_0(k6_numbers,X1) = X1 | ~m1_subset_1(X1,k1_numbers) <- (9, 10) [subsumption resolution 537,212] {a:1,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
654. k6_numbers = k3_arytm_0(k6_numbers) | ~m1_subset_1(k6_numbers,k1_numbers) | ~m1_subset_1(k3_arytm_0(k6_numbers),k1_numbers) <- (9, 10) [superposition 463,538] {a:2,w:11,wCS:4400,sine:1,thAx:0,allAx:3,thDist:-3}
664. 60 <=> m1_subset_1(k3_arytm_0(k6_numbers),k1_numbers) [avatar definition]
666. ~m1_subset_1(k3_arytm_0(k6_numbers),k1_numbers) <- (~60) [avatar component clause 664] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
669. ~m1_subset_1(k6_numbers,k1_numbers) | ~m1_subset_1(k3_arytm_0(k6_numbers),k1_numbers) <- (~1, 9, 10) [subsumption resolution 654,173] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
670. ~m1_subset_1(k3_arytm_0(k6_numbers),k1_numbers) <- (~1, 9, 10) [subsumption resolution 669,212] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
671. ~60 | 1 | ~9 | ~10 [avatar split clause 670,215,211,171,664]
673. ~m1_subset_1(k6_numbers,k1_numbers) <- (~60) [resolution 666,137] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
677. $false <- (9, ~60) [subsumption resolution 673,212] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
678. ~9 | 60 [avatar contradiction clause 677]
679. $false [avatar sat refutation 174,179,184,189,194,209,217,268,336,425,487,671,678] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:315,thDist:-315}
% SZS output end Proof for arithm__l3_arithm
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 50
% >>> Preprocessing
% Introduced skolems: 7
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 49
% Generated clauses: 318
% Active clauses: 105
% Passive clauses: 144
% Final active clauses: 101
% Final passive clauses: 29
% Main loop iterations started: 112
% 
% >>> Simplifying Inferences
% Duplicate literals: 27
% Fw subsumption resolutions: 32
% Fw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 45
% 
% >>> Generating Inferences
% Binary resolution: 101
% Unit resulting resolution: 7
% Forward superposition: 21
% Backward superposition: 19
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 18
% Sat splitting refutations: 11
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 72
% SAT solver unit clauses: 18
% SAT solver binary clauses: 35
% 
% Memory used [KB]: 5245
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit clause index maintenance: 0.001 s
% backward demodulation index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.003 s ( own 0.0000 s ) 

