i: [1,40,0,2,1,1,0,0] t16_relat_1
a: [3,68,0,2,1,40]
s: 68
i: [1,39,0,2,1,1,0,0] t16_relat_1
a: [3,73,0,2,1,39]
s: 73
i: [1,41,0,3,1,1,0,0] t16_relat_1
a: [3,63,0,3,1,41]
s: 63
i: [1,42,0,5,1,1,0,0] t16_relat_1
a: [3,58,0,5,1,42]
s: 58
i: [1,51,0,3,1,0,0,1] reflexivity_r1_tarski
s: 51
i: [1,50,0,4,1,0,0,255] existence_m1_subset_1
s: 50
i: [1,52,0,5,1,0,0,2] idempotence_k2_xboole_0
s: 52
i: [1,47,0,7,2,0,0,255] t3_subset
s: 47
d: [2,75,1,5,1,0,40,47,50]
s: 75
i: [1,48,0,7,2,0,0,255] t3_subset
s: 48
i: [1,53,0,7,1,0,0,2] commutativity_k2_xboole_0
s: 53
i: [1,46,0,8,3,0,0,255] cc2_relat_1
s: 46
i: [1,54,0,8,3,0,0,1] fc3_relat_1
s: 54
d: [2,77,1,6,2,0,40,46,50]
s: 77
i: [1,43,0,10,2,0,0,1] d6_relat_1
s: 43
i: [1,44,0,12,4,0,0,2] t11_relat_1
s: 44
d: [2,78,1,7,3,0,40,46,48]
s: 78
i: [1,45,0,12,4,0,0,2] t11_relat_1
s: 45
d: [2,89,0,10,3,0,21,45,78]
s: 89
i: [1,49,0,13,3,0,0,1] t13_xboole_1
s: 49
d: [2,83,1,11,4,0,44,54,43]
s: 83
d: [2,92,1,11,3,0,44,49,52]
s: 92
d: [2,96,1,11,3,0,44,49,52]
s: 96
d: [2,85,2,13,5,0,40,78,44]
d: [2,88,2,11,4,0,21,85,78]
s: 88
d: [2,93,1,13,3,0,44,49,53]
s: 93
d: [2,94,1,13,3,0,44,49,53]
s: 94
d: [2,118,3,8,3,1,40,88,63]
d: [2,121,3,6,2,2,21,118,68]
a: [3,125,3,3,1,121]
s: 125
d: [2,97,1,13,3,0,44,49,53]
s: 97
d: [2,119,3,10,3,0,40,88,75]
s: 119
d: [2,98,1,13,3,0,44,49,53]
s: 98
d: [2,95,1,16,4,0,44,49,43]
s: 95
d: [2,90,2,11,4,0,40,89,78]
s: 90
d: [2,99,1,16,4,0,44,49,43]
s: 99
d: [2,155,3,8,3,1,40,90,63]
d: [2,158,3,6,2,2,21,155,68]
a: [3,166,3,3,1,158]
s: 166
d: [2,100,2,12,4,0,40,92,78]
s: 100
d: [2,106,2,12,4,0,40,96,78]
s: 106
d: [2,156,3,10,3,0,40,90,75]
s: 156
d: [2,91,2,14,4,0,40,49,78]
s: 91
d: [2,104,2,14,4,0,44,92,43]
s: 104
d: [2,110,2,14,4,0,44,96,43]
s: 110
d: [2,132,2,14,4,0,40,93,78]
s: 132
d: [2,195,3,12,3,0,40,104,51]
s: 195
d: [2,143,2,14,4,0,44,95,52]
s: 143
d: [2,201,3,12,3,0,40,110,51]
s: 201
d: [2,171,2,14,4,0,44,99,52]
s: 171
d: [2,136,2,16,4,0,44,93,43]
s: 136
d: [2,190,3,14,5,0,40,91,54]
s: 190
d: [2,140,2,16,4,0,44,93,43]
s: 140
d: [2,192,3,14,4,0,44,91,53]
s: 192
d: [2,144,2,16,4,0,44,95,53]
s: 144
d: [2,145,2,16,4,0,44,95,53]
s: 145
d: [2,193,3,14,4,0,44,91,53]
s: 193
d: [2,172,2,16,4,0,44,99,53]
s: 172
d: [2,205,3,14,5,0,40,132,54]
s: 205
d: [2,173,2,16,4,0,44,99,53]
s: 173
d: [2,142,2,17,5,0,40,95,78]
s: 142
d: [2,232,3,14,4,0,40,142,51]
s: 232
d: [2,170,2,17,5,0,40,99,78]
s: 170
d: [2,238,4,13,4,0,44,232,52]
s: 238
d: [2,146,2,19,5,0,44,95,43]
s: 146
d: [2,174,2,19,5,0,44,99,43]
s: 174
d: [2,243,3,14,4,0,40,170,51]
s: 243
d: [2,179,3,15,5,0,44,100,43]
s: 179
d: [2,259,4,13,4,0,44,243,52]
s: 259
d: [2,185,3,15,5,0,44,106,43]
s: 185
d: [2,197,3,17,5,0,40,104,89]
d: [2,199,3,15,4,0,18,197]
s: 199
d: [2,263,4,13,4,0,40,179,51]
s: 263
d: [2,200,3,17,5,0,40,110,89]
d: [2,203,3,15,4,0,21,200,78]
s: 203
d: [2,237,4,14,5,0,40,232,54]
s: 237
d: [2,249,3,19,6,0,40,146,89]
d: [2,252,3,17,5,0,18,249]
d: [2,253,3,15,4,0,21,252,78]
s: 253
d: [2,115,3,16,5,0,40,88,96]
s: 115
d: [2,277,4,17,6,0,40,253,44]
d: [2,279,4,12,4,0,18,277]
d: [2,280,4,10,3,0,21,279,78]
s: 280
r: 253
d: [2,286,5,0,0,3,52,68,63,58,280]
e: 286
d: [2,287,5,5,2,1,40,280,58]
d: [2,292,5,3,1,2,21,287,68]
d: [2,293,5,0,0,3,21,292,63]
e: 293
f: 293,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t16_relat_1
% SZS output start Proof for relat_1__t16_relat_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => r1_tarski(k1_relat_1(X0),k1_relat_1(X1))))) [input t16_relat_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => r1_tarski(k1_relat_1(X0),k1_relat_1(X1))))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
5. ! [X0] : (v1_relat_1(X0) => k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0))) [input d6_relat_1]
16. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => (r1_tarski(X0,X1) => (r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)))))) [input t11_relat_1]
17. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),k2_xboole_0(X1,X3))) [input t13_xboole_1]
18. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
19. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 18]
20. ! [X0,X1,X2,X3] : ((r1_tarski(X2,X3) & r1_tarski(X0,X1)) => r1_tarski(k2_xboole_0(X0,X2),k2_xboole_0(X1,X3))) [flattening 17]
23. ? [X0] : (? [X1] : ((~r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) & r1_tarski(X0,X1)) & v1_relat_1(X1)) & v1_relat_1(X0)) [ennf transformation 2]
24. ? [X0] : (? [X1] : (~r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) & r1_tarski(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) [flattening 23]
25. ! [X0] : (k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0)) [ennf transformation 5]
26. ! [X0] : (! [X1] : (((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 16]
27. ! [X0] : (! [X1] : ((r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) & r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [flattening 26]
28. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 3]
29. ! [X0,X1,X2,X3] : (r1_tarski(k2_xboole_0(X0,X2),k2_xboole_0(X1,X3)) | (~r1_tarski(X2,X3) | ~r1_tarski(X0,X1))) [ennf transformation 20]
30. ! [X0,X1,X2,X3] : (r1_tarski(k2_xboole_0(X0,X2),k2_xboole_0(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1)) [flattening 29]
33. ? [X0] : (? [X1] : (~r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) & r1_tarski(X0,X1) & v1_relat_1(X1)) & v1_relat_1(X0)) => (? [X1] : (~r1_tarski(k1_relat_1(sK0),k1_relat_1(X1)) & r1_tarski(sK0,X1) & v1_relat_1(X1)) & v1_relat_1(sK0)) [choice axiom]
34. ? [X1] : (~r1_tarski(k1_relat_1(sK0),k1_relat_1(X1)) & r1_tarski(sK0,X1) & v1_relat_1(X1)) => (~r1_tarski(k1_relat_1(sK0),k1_relat_1(sK1)) & r1_tarski(sK0,sK1) & v1_relat_1(sK1)) [choice axiom]
35. (~r1_tarski(k1_relat_1(sK0),k1_relat_1(sK1)) & r1_tarski(sK0,sK1) & v1_relat_1(sK1)) & v1_relat_1(sK0) [skolemisation 24,34,33]
36. ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1)))) [nnf transformation 19]
40. v1_relat_1(sK1) [cnf transformation 35] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. r1_tarski(sK0,sK1) [cnf transformation 35] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~r1_tarski(k1_relat_1(sK0),k1_relat_1(sK1)) [cnf transformation 35] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
43. k1_relat_1(X0) = k2_xboole_0(k9_xtuple_0(X0),k10_xtuple_0(X0)) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
45. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:1,thDist:-1}
46. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
48. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
49. r1_tarski(k2_xboole_0(X0,X2),k2_xboole_0(X1,X3)) | ~r1_tarski(X2,X3) | ~r1_tarski(X0,X1) [cnf transformation 30] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. 1 <=> r1_tarski(k1_relat_1(sK0),k1_relat_1(sK1)) [avatar definition]
58. ~r1_tarski(k1_relat_1(sK0),k1_relat_1(sK1)) <- (~1) [avatar component clause 56] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
59. ~1 [avatar split clause 42,56]
61. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
63. r1_tarski(sK0,sK1) <- (2) [avatar component clause 61] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
64. 2 [avatar split clause 41,61]
66. 3 <=> v1_relat_1(sK1) [avatar definition]
68. v1_relat_1(sK1) <- (3) [avatar component clause 66] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
69. 3 [avatar split clause 40,66]
78. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 46,48] {a:1,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. r1_tarski(k10_xtuple_0(X0),k10_xtuple_0(X1)) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) [subsumption resolution 45,78] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
95. r1_tarski(k1_relat_1(X11),k2_xboole_0(X12,X13)) | ~r1_tarski(k10_xtuple_0(X11),X13) | ~r1_tarski(k9_xtuple_0(X11),X12) | ~v1_relat_1(X11) [superposition 49,43] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
146. ~r1_tarski(k10_xtuple_0(X9),k10_xtuple_0(X8)) | r1_tarski(k1_relat_1(X9),k1_relat_1(X8)) | ~r1_tarski(k9_xtuple_0(X9),k9_xtuple_0(X8)) | ~v1_relat_1(X9) | ~v1_relat_1(X8) [superposition 95,43] {a:2,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
249. r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) [resolution 146,89] {a:3,w:19,wCS:7600,sine:1,thAx:0,allAx:4,thDist:-4}
252. r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~v1_relat_1(X0) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) [duplicate literal removal 249] {a:3,w:17,wCS:6800,sine:1,thAx:0,allAx:4,thDist:-4}
253. ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) [subsumption resolution 252,78] {a:3,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
277. r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) | ~r1_tarski(X0,X1) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [resolution 253,44] {a:4,w:17,wCS:6800,sine:1,thAx:0,allAx:5,thDist:-5}
279. r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) | ~v1_relat_1(X0) [duplicate literal removal 277] {a:4,w:12,wCS:4800,sine:1,thAx:0,allAx:5,thDist:-5}
280. r1_tarski(k1_relat_1(X0),k1_relat_1(X1)) | ~v1_relat_1(X1) | ~r1_tarski(X0,X1) [subsumption resolution 279,78] {a:4,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
287. ~v1_relat_1(sK1) | ~r1_tarski(sK0,sK1) <- (~1) [resolution 280,58] {a:5,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
292. ~r1_tarski(sK0,sK1) <- (~1, 3) [subsumption resolution 287,68] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
293. $false <- (~1, 2, 3) [subsumption resolution 292,63] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
294. 1 | ~2 | ~3 [avatar contradiction clause 293]
295. $false [avatar sat refutation 59,64,69,294] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:58,thDist:-58}
% SZS output end Proof for relat_1__t16_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 16
% Generated clauses: 222
% Active clauses: 58
% Passive clauses: 131
% Final active clauses: 57
% Final passive clauses: 59
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 16
% Fw subsumption resolutions: 16
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 54
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 92
% Unit resulting resolution: 1
% Forward superposition: 75
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.002 s
% other: 0.003 s ( own 0.001 s ) 

