i: [1,160,0,2,1,1,0,0] t9_arytm_3
a: [3,247,0,2,1,160]
s: 247
i: [1,159,0,2,1,1,0,0] t9_arytm_3
a: [3,252,0,2,1,159]
s: 252
i: [1,161,0,6,2,1,0,0] t9_arytm_3
a: [3,238,0,3,1,161]
s: 238
i: [1,162,0,2,1,0,0,1] spc1_boole
a: [3,257,0,2,1,162]
s: 257
i: [1,163,0,2,1,0,0,1] fc1_xboole_0
a: [3,262,0,2,1,163]
s: 262
i: [1,168,0,2,1,0,0,255] fc6_ordinal1
a: [3,282,0,2,1,168]
s: 282
i: [1,167,0,2,1,0,0,255] fc6_ordinal1
a: [3,287,0,2,1,167]
s: 287
i: [1,222,0,7,3,0,0,1] reflexivity_r2_arytm_3
a: [3,306,0,2,1,222]
s: 306
d: [2,363,1,0,0,2,40,306,282]
e: 363
d: [2,362,1,0,0,2,40,306,252]
e: 362
d: [2,361,1,0,0,2,52,252,306]
e: 361
r: 306
i: [1,224,0,2,1,0,0,2] rc2_xboole_0
a: [3,314,0,2,1,224]
s: 314
i: [1,225,0,2,1,0,0,1] rc4_ordinal1
a: [3,319,0,2,1,225]
s: 319
i: [1,226,0,2,1,0,0,1] rc1_ordinal1
a: [3,324,0,2,1,226]
s: 324
i: [1,227,0,2,1,0,0,2] rc1_xboole_0
a: [3,329,0,2,1,227]
s: 329
i: [1,231,0,2,1,0,0,255] rc2_ordinal1
a: [3,334,0,2,1,231]
s: 334
i: [1,230,0,2,1,0,0,255] rc2_ordinal1
a: [3,339,0,2,1,230]
s: 339
i: [1,229,0,2,1,0,0,255] rc2_ordinal1
a: [3,344,0,2,1,229]
s: 344
i: [1,228,0,2,1,0,0,255] rc2_ordinal1
a: [3,349,0,2,1,228]
s: 349
i: [1,233,0,2,1,0,0,1] rc5_ordinal1
a: [3,354,0,2,1,233]
s: 354
i: [1,232,0,2,1,0,0,1] rc5_ordinal1
a: [3,359,0,2,1,232]
s: 359
i: [1,172,0,3,1,0,0,255] fc1_subset_1
s: 172
i: [1,208,0,3,1,0,0,255] rc2_subset_1
s: 208
i: [1,211,0,3,1,0,0,255] reflexivity_r1_tarski
s: 211
i: [1,164,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,267,0,3,1,164]
s: 267
i: [1,165,0,3,1,0,0,255] t1_numerals
a: [3,272,0,3,1,165]
s: 272
i: [1,171,0,3,1,0,0,255] spc1_numerals
a: [3,292,0,3,1,171]
s: 292
i: [1,170,0,3,1,0,0,255] spc1_numerals
a: [3,297,0,3,1,170]
s: 297
d: [2,367,0,3,1,2,22,297,267]
a: [3,371,0,3,1,367]
s: 371
i: [1,180,0,4,2,0,0,1] cc6_ordinal1
s: 180
i: [1,182,0,4,2,0,0,255] cc1_ordinal1
s: 182
d: [2,387,1,2,1,1,40,182,252]
a: [3,398,1,2,1,387]
s: 398
i: [1,183,0,4,2,0,0,255] cc1_ordinal1
s: 183
i: [1,191,0,4,2,0,0,255] cc9_ordinal1
s: 191
d: [2,406,1,2,1,1,40,183,252]
a: [3,417,1,2,1,406]
s: 417
i: [1,192,0,4,2,0,0,1] cc7_ordinal1
s: 192
d: [2,374,1,2,1,1,40,180,319]
a: [3,379,1,2,1,374]
s: 379
i: [1,193,0,4,2,0,0,1] cc3_ordinal1
s: 193
i: [1,206,0,4,1,0,0,2] existence_m1_subset_1
s: 206
d: [2,375,1,2,1,1,40,180,354]
a: [3,384,1,2,1,375]
s: 384
i: [1,210,0,4,1,0,0,255] rc3_subset_1
s: 210
d: [2,386,1,2,1,1,40,182,282]
a: [3,393,1,2,1,386]
s: 393
i: [1,166,0,4,1,0,0,255] dt_k5_numbers
a: [3,277,0,4,1,166]
s: 277
i: [1,169,0,4,1,0,0,255] spc1_numerals
a: [3,302,0,4,1,169]
s: 302
d: [2,388,1,2,1,1,40,182,324]
a: [3,403,1,2,1,388]
s: 403
d: [2,487,0,4,1,2,22,277,267]
a: [3,491,0,4,1,487]
s: 491
d: [2,405,1,2,1,1,40,183,282]
a: [3,412,1,2,1,405]
s: 412
d: [2,493,0,4,1,2,22,302,267]
a: [3,497,0,4,1,493]
s: 497
i: [1,176,0,5,2,0,0,255] rc1_subset_1
s: 176
d: [2,407,1,2,1,1,40,183,324]
a: [3,422,1,2,1,407]
s: 422
i: [1,194,0,5,2,0,0,1] t6_boole
s: 194
d: [2,502,1,5,2,1,44,238,194]
s: 502
i: [1,197,0,5,2,0,0,255] cc8_ordinal1
s: 197
i: [1,203,0,5,2,0,0,255] t7_boole
s: 203
d: [2,424,1,2,1,1,40,191,262]
a: [3,430,1,2,1,424]
s: 430
i: [1,207,0,5,1,0,0,255] rc2_subset_1
s: 207
d: [2,425,1,2,1,1,40,191,329]
a: [3,435,1,2,1,425]
s: 435
i: [1,209,0,5,1,0,0,255] rc3_subset_1
s: 209
a: [3,309,0,5,2,222]
s: 309
d: [2,518,2,4,2,2,40,309,502]
d: [2,519,2,2,1,3,21,518,252]
a: [3,523,2,2,1,519]
s: 523
i: [1,178,0,6,2,0,0,255] rc4_subset_1
s: 178
d: [2,437,1,2,1,1,40,192,262]
a: [3,443,1,2,1,437]
s: 443
i: [1,198,0,6,3,0,0,255] cc2_ordinal1
s: 198
i: [1,199,0,6,2,0,0,255] t1_subset
s: 199
d: [2,438,1,2,1,1,40,192,329]
a: [3,448,1,2,1,438]
s: 448
i: [1,212,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 212
d: [2,462,1,2,1,1,40,193,262]
a: [3,468,1,2,1,462]
s: 468
i: [1,175,0,7,2,0,0,255] rc1_subset_1
s: 175
i: [1,177,0,7,2,0,0,255] rc4_subset_1
s: 177
d: [2,463,1,2,1,1,40,193,329]
a: [3,473,1,2,1,463]
s: 473
i: [1,181,0,7,2,0,0,1] t35_ordinal2
s: 181
d: [2,504,1,2,1,1,40,197,371]
a: [3,510,1,2,1,504]
s: 510
i: [1,184,0,7,2,0,0,1] t39_ordinal2
s: 184
i: [1,185,0,7,2,0,0,1] t39_ordinal2
s: 185
d: [2,450,2,2,1,1,40,379,183]
a: [3,455,2,2,1,450]
s: 455
i: [1,189,0,7,3,0,0,1] cc2_ordinal2
s: 189
d: [2,451,2,2,1,1,40,379,182]
a: [3,460,2,2,1,451]
s: 460
i: [1,190,0,7,3,0,0,1] cc5_ordinal1
s: 190
i: [1,201,0,7,2,0,0,255] t3_subset
s: 201
d: [2,475,2,2,1,1,40,384,183]
a: [3,480,2,2,1,475]
s: 480
i: [1,202,0,7,3,0,0,2] t8_boole
s: 202
d: [2,476,2,2,1,1,40,384,182]
a: [3,485,2,2,1,476]
s: 485
i: [1,179,0,8,3,0,0,255] cc10_ordinal1
s: 179
i: [1,195,0,8,3,0,0,255] cc1_subset_1
s: 195
d: [2,586,1,4,2,1,40,179,491]
a: [3,595,1,2,1,586]
s: 595
i: [1,200,0,8,3,0,0,255] t2_subset
s: 200
d: [2,602,1,4,2,1,40,195,491]
d: [2,608,1,2,1,2,21,602,287]
a: [3,612,1,2,1,608]
s: 612
i: [1,223,0,8,3,0,0,2] dt_k11_ordinal2
s: 223
i: [1,196,0,9,3,0,0,255] cc4_subset_1
s: 196
d: [2,534,2,2,1,1,40,468,183]
a: [3,540,2,2,1,534]
s: 540
i: [1,205,0,9,3,0,0,255] t5_subset
s: 205
d: [2,535,2,2,1,1,40,468,182]
a: [3,545,2,2,1,535]
s: 545
i: [1,204,0,10,3,0,0,255] t4_subset
s: 204
i: [1,173,0,11,4,0,0,255] cc3_subset_1
s: 173
d: [2,547,2,2,1,1,40,473,183]
a: [3,552,2,2,1,547]
s: 552
i: [1,174,0,11,4,0,0,255] cc2_subset_1
s: 174
d: [2,548,2,2,1,1,40,473,182]
a: [3,557,2,2,1,548]
s: 557
i: [1,186,0,11,4,0,0,1] d3_arytm_3
s: 186
i: [1,217,0,12,5,0,0,255] dt_k9_ordinal3
s: 217
d: [2,560,2,2,1,1,40,510,180]
a: [3,564,2,2,1,560]
s: 564
d: [2,698,0,10,4,0,21,217,180]
d: [2,699,0,8,3,0,21,698,180]
s: 699
d: [2,700,3,2,1,1,40,564,183]
a: [3,705,3,2,1,700]
s: 705
i: [1,220,0,12,5,0,0,2] fc4_ordinal3
s: 220
i: [1,221,0,12,5,0,0,2] fc4_ordinal3
s: 221
d: [2,701,3,2,1,1,40,564,182]
a: [3,710,3,2,1,701]
s: 710
d: [2,715,0,10,4,0,21,221,180]
d: [2,716,0,8,3,0,21,715,180]
s: 716
d: [2,426,1,3,1,0,40,191,208]
s: 426
i: [1,234,0,13,4,0,0,1] d3_arytm_3
s: 234
d: [2,721,0,9,3,0,21,234,223]
s: 721
d: [2,722,2,8,3,1,40,721,502]
d: [2,729,2,6,2,2,21,722,252]
s: 729
i: [1,187,0,14,4,0,0,1] d3_arytm_3
s: 187
d: [2,439,1,3,1,0,40,192,208]
s: 439
i: [1,216,0,14,4,0,0,255] existence_m2_subset_1
s: 216
d: [2,748,0,12,3,0,21,216,195]
s: 748
d: [2,464,1,3,1,0,40,193,208]
s: 464
i: [1,213,0,15,5,0,0,255] dt_m2_subset_1
s: 213
d: [2,505,1,3,1,0,40,197,206]
a: [3,515,1,3,1,505]
s: 515
d: [2,751,0,13,4,0,21,213,195]
s: 751
i: [1,214,0,15,5,0,0,255] redefinition_m2_subset_1
s: 214
d: [2,615,1,5,2,1,40,200,272]
d: [2,625,1,3,1,2,21,615,287]
a: [3,629,1,3,1,625]
s: 629
d: [2,761,0,13,4,0,21,214,195]
s: 761
d: [2,616,1,5,2,1,40,200,292]
a: [3,634,1,3,1,616]
s: 634
i: [1,215,0,15,5,0,0,255] redefinition_m2_subset_1
s: 215
d: [2,782,0,13,4,0,21,215,195]
s: 782
d: [2,617,1,5,2,1,40,200,371]
d: [2,636,1,3,1,2,21,617,287]
a: [3,640,1,3,1,636]
s: 640
i: [1,218,0,15,5,0,0,255] commutativity_k9_ordinal3
s: 218
d: [2,747,2,3,1,0,40,439,180]
s: 747
d: [2,800,0,13,4,0,21,218,180]
d: [2,801,0,11,3,0,21,800,180]
s: 801
i: [1,219,0,15,5,0,0,255] redefinition_k9_ordinal3
s: 219
d: [2,749,2,3,1,0,40,464,183]
s: 749
d: [2,806,0,13,4,0,21,219,180]
d: [2,807,0,11,3,0,21,806,180]
s: 807
d: [2,815,3,10,4,2,44,729,807]
d: [2,833,3,8,3,3,21,815,247]
d: [2,834,3,6,2,3,21,833,180]
s: 834
d: [2,618,1,7,2,1,40,200,491]
d: [2,642,1,4,1,1,21,618,172]
a: [3,646,1,4,1,642]
s: 646
d: [2,500,1,5,2,1,44,272,194]
s: 500
d: [2,840,4,10,4,3,44,834,801]
d: [2,843,4,8,3,3,18,840]
d: [2,844,4,6,2,3,21,843,247]
s: 844
d: [2,573,1,5,2,0,40,189,206]
s: 573
d: [2,841,4,10,4,3,44,834,801]
d: [2,842,4,8,3,3,18,841]
d: [2,845,4,6,2,3,21,842,247]
s: 845
d: [2,590,1,5,2,0,40,179,207]
s: 590
d: [2,591,1,5,2,0,40,179,209]
s: 591
d: [2,750,2,3,1,0,40,464,182]
s: 750
d: [2,607,1,5,2,0,40,195,209]
s: 607
d: [2,752,2,3,1,1,40,515,180]
a: [3,756,2,3,1,752]
s: 756
d: [2,622,1,8,2,0,40,200,207]
d: [2,650,1,5,1,0,21,622,172]
s: 650
d: [2,623,1,8,2,0,40,200,209]
d: [2,651,1,5,1,0,21,623,172]
s: 651
d: [2,762,2,3,1,1,40,629,212]
a: [3,769,2,3,1,762]
s: 769
d: [2,690,1,9,3,0,40,173,209]
d: [2,694,1,5,2,0,21,690,210]
s: 694
d: [2,774,2,3,1,1,40,634,212]
a: [3,780,2,3,1,774]
s: 780
d: [2,724,1,9,4,0,44,721,184]
d: [2,727,1,7,3,0,18,724]
d: [2,730,1,5,2,1,21,727,564]
s: 730
d: [2,574,1,6,2,0,40,189,207]
s: 574
d: [2,792,2,3,1,1,40,640,212]
a: [3,798,2,3,1,792]
s: 798
d: [2,575,1,6,2,0,40,189,209]
s: 575
d: [2,867,3,3,1,1,40,756,183]
a: [3,872,3,3,1,867]
s: 872
d: [2,589,1,6,2,0,40,179,206]
s: 589
d: [2,605,1,6,2,0,40,195,206]
s: 605
d: [2,868,3,3,1,1,40,756,182]
a: [3,877,3,3,1,868]
s: 877
d: [2,621,1,6,2,0,40,200,206]
s: 621
d: [2,846,2,4,1,1,40,646,212]
a: [3,852,2,4,1,846]
s: 852
d: [2,665,1,6,2,0,40,196,207]
s: 665
d: [2,674,1,6,2,0,40,205,207]
s: 674
d: [2,668,3,4,2,1,44,540,194]
s: 668
d: [2,675,1,6,2,0,40,205,209]
s: 675
d: [2,676,3,4,2,1,44,545,194]
s: 676
d: [2,678,1,6,2,1,40,204,491]
s: 678
d: [2,689,1,9,3,0,40,173,207]
d: [2,693,1,6,2,0,21,689,208]
s: 693
d: [2,906,2,3,1,2,40,678,629]
a: [3,912,2,3,1,906]
s: 912
d: [2,499,1,7,3,0,44,194,194]
s: 499
d: [2,921,3,5,2,1,40,912,200]
d: [2,924,3,3,1,2,21,921,612]
a: [3,928,3,3,1,924]
s: 928
d: [2,585,1,7,3,0,40,179,201]
s: 585
d: [2,587,1,7,3,0,40,179,175]
s: 587
d: [2,930,4,3,1,1,40,928,212]
a: [3,937,4,3,1,930]
s: 937
d: [2,588,1,7,3,0,40,179,177]
s: 588
d: [2,908,2,6,2,1,40,678,621]
d: [2,914,2,4,1,2,21,908,287]
a: [3,918,2,4,1,914]
s: 918
d: [2,601,1,7,3,0,40,195,201]
s: 601
d: [2,614,1,10,3,0,40,200,201]
d: [2,624,1,7,2,0,21,614,172]
s: 624
d: [2,941,3,6,2,1,40,918,200]
d: [2,943,3,4,1,2,21,941,612]
a: [3,947,3,4,1,943]
s: 947
d: [2,619,1,10,3,0,40,200,175]
d: [2,648,1,7,2,0,21,619,172]
s: 648
d: [2,953,4,4,1,1,40,947,212]
a: [3,959,4,4,1,953]
s: 959
d: [2,620,1,10,3,0,40,200,177]
d: [2,649,1,7,2,0,21,620,172]
s: 649
d: [2,664,1,7,2,0,40,196,206]
s: 664
d: [2,765,2,5,2,1,44,629,194]
s: 765
d: [2,673,1,7,2,0,40,205,206]
s: 673
d: [2,855,2,7,3,1,40,500,200]
d: [2,857,2,5,2,2,21,855,287]
s: 857
d: [2,682,1,7,2,0,40,204,207]
s: 682
d: [2,683,1,7,2,0,40,204,209]
s: 683
d: [2,862,2,5,2,0,40,573,183]
s: 862
d: [2,784,1,9,3,1,40,782,491]
d: [2,790,1,7,2,2,21,784,287]
s: 790
d: [2,863,2,5,2,0,40,573,182]
s: 863
d: [2,809,1,11,4,0,44,807,181]
d: [2,822,1,9,3,1,21,809,443]
d: [2,823,1,7,2,1,21,822,180]
s: 823
d: [2,986,5,6,3,4,44,844,823]
d: [2,996,5,4,2,5,21,986,262]
d: [2,997,5,2,1,5,21,996,443]
d: [2,998,5,0,0,5,21,997,247]
e: 998
r: 844
r: 834
r: 729
r: 502
r: 238
a: [3,237,0,3,1,161]
s: 237
a: [3,242,0,3,1,161]
s: 242
d: [2,1001,2,0,0,3,52,252,242,730]
e: 1001
d: [2,1002,2,2,1,2,40,242,730]
d: [2,1004,2,0,0,3,21,1002,252]
e: 1004
d: [2,503,1,2,1,1,40,197,272]
d: [2,569,1,4,2,1,40,189,371]
d: [2,581,1,2,1,2,21,569,282]
f: 1004,998,581,503,367,170,168,165,164,163,159,160,161
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_3__t9_arytm_3
% SZS output start Proof for arytm_3__t9_arytm_3
1. ! [X0] : ((v7_ordinal1(X0) & v3_ordinal1(X0)) => (r2_arytm_3(1,X0) & r2_arytm_3(X0,k1_xboole_0))) [input t9_arytm_3]
2. ~! [X0] : ((v7_ordinal1(X0) & v3_ordinal1(X0)) => (r2_arytm_3(1,X0) & r2_arytm_3(X0,k1_xboole_0))) [negated conjecture 1]
8. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (m1_subset_1(X1,X0) => v3_ordinal1(X1))) [input cc2_ordinal2]
15. ! [X0] : (v7_ordinal1(X0) => v3_ordinal1(X0)) [input cc6_ordinal1]
17. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
19. ! [X0,X1] : ((v7_ordinal1(X1) & v3_ordinal1(X1) & v7_ordinal1(X0) & v3_ordinal1(X0)) => k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0)) [input commutativity_k9_ordinal3]
20. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_arytm_3(X0,X1) <=> ? [X2] : (k11_ordinal2(X0,X2) = X1 & v3_ordinal1(X2))))) [input d3_arytm_3]
21. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k11_ordinal2(X0,X1))) [input dt_k11_ordinal2]
33. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
35. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
46. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
47. ! [X0,X1] : ((v7_ordinal1(X1) & v3_ordinal1(X1) & v7_ordinal1(X0) & v3_ordinal1(X0)) => k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1)) [input redefinition_k9_ordinal3]
52. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
53. m1_subset_1(k1_xboole_0,k4_ordinal1) [input t1_numerals]
56. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0)) [input t35_ordinal2]
57. ! [X0] : (v3_ordinal1(X0) => (k11_ordinal2(X0,1) = X0 & k11_ordinal2(1,X0) = X0)) [input t39_ordinal2]
61. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
73. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 52]
75. ? [X0] : ((~r2_arytm_3(1,X0) | ~r2_arytm_3(X0,k1_xboole_0)) & (v7_ordinal1(X0) & v3_ordinal1(X0))) [ennf transformation 2]
76. ? [X0] : ((~r2_arytm_3(1,X0) | ~r2_arytm_3(X0,k1_xboole_0)) & v7_ordinal1(X0) & v3_ordinal1(X0)) [flattening 75]
84. ! [X0] : (v3_ordinal1(X0) | ~v7_ordinal1(X0)) [ennf transformation 15]
85. ! [X0] : (k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0) | ~v3_ordinal1(X0)) [ennf transformation 56]
87. ! [X0] : ((k11_ordinal2(X0,1) = X0 & k11_ordinal2(1,X0) = X0) | ~v3_ordinal1(X0)) [ennf transformation 57]
88. ! [X0] : (! [X1] : ((r2_arytm_3(X0,X1) <=> ? [X2] : (k11_ordinal2(X0,X2) = X1 & v3_ordinal1(X2))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 20]
89. ! [X0] : (! [X1] : (v3_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v3_ordinal1(X0)) [ennf transformation 8]
94. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 61]
97. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 17]
118. ! [X0,X1] : (k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0) | (~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0))) [ennf transformation 19]
119. ! [X0,X1] : (k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0) | ~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0)) [flattening 118]
120. ! [X0,X1] : (k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1) | (~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0))) [ennf transformation 47]
121. ! [X0,X1] : (k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1) | ~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0)) [flattening 120]
126. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 21]
127. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 126]
128. ? [X0] : ((~r2_arytm_3(1,X0) | ~r2_arytm_3(X0,k1_xboole_0)) & v7_ordinal1(X0) & v3_ordinal1(X0)) => ((~r2_arytm_3(1,sK0) | ~r2_arytm_3(sK0,k1_xboole_0)) & v7_ordinal1(sK0) & v3_ordinal1(sK0)) [choice axiom]
129. (~r2_arytm_3(1,sK0) | ~r2_arytm_3(sK0,k1_xboole_0)) & v7_ordinal1(sK0) & v3_ordinal1(sK0) [skolemisation 76,128]
134. ! [X0] : (! [X1] : (((r2_arytm_3(X0,X1) | ! [X2] : (k11_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2))) & (? [X2] : (k11_ordinal2(X0,X2) = X1 & v3_ordinal1(X2)) | ~r2_arytm_3(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 88]
135. ! [X0] : (! [X1] : (((r2_arytm_3(X0,X1) | ! [X2] : (k11_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2))) & (? [X3] : (k11_ordinal2(X0,X3) = X1 & v3_ordinal1(X3)) | ~r2_arytm_3(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [rectify 134]
136. ! [X1,X0] : (? [X3] : (k11_ordinal2(X0,X3) = X1 & v3_ordinal1(X3)) => (k11_ordinal2(X0,sK3(X0,X1)) = X1 & v3_ordinal1(sK3(X0,X1)))) [choice axiom]
137. ! [X0] : (! [X1] : (((r2_arytm_3(X0,X1) | ! [X2] : (k11_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2))) & ((k11_ordinal2(X0,sK3(X0,X1)) = X1 & v3_ordinal1(sK3(X0,X1))) | ~r2_arytm_3(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [skolemisation 135,136]
159. v3_ordinal1(sK0) [cnf transformation 129] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
160. v7_ordinal1(sK0) [cnf transformation 129] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
161. ~r2_arytm_3(1,sK0) | ~r2_arytm_3(sK0,k1_xboole_0) [cnf transformation 129] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
163. v1_xboole_0(k1_xboole_0) [cnf transformation 33] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
164. k4_ordinal1 = k5_numbers [cnf transformation 46] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
165. m1_subset_1(k1_xboole_0,k4_ordinal1) [cnf transformation 53] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
168. v3_ordinal1(k4_ordinal1) [cnf transformation 35] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
170. m1_subset_1(1,k5_numbers) [cnf transformation 73] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
180. ~v7_ordinal1(X0) | v3_ordinal1(X0) [cnf transformation 84] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
181. k1_xboole_0 = k11_ordinal2(k1_xboole_0,X0) | ~v3_ordinal1(X0) [cnf transformation 85] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
184. k11_ordinal2(1,X0) = X0 | ~v3_ordinal1(X0) [cnf transformation 87] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
188. r2_arytm_3(X0,X1) | k11_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 137] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
189. ~m1_subset_1(X1,X0) | v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 89] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
194. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 94] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
197. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 97] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
218. k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0) | ~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0) [cnf transformation 119] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
219. k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1) | ~v7_ordinal1(X1) | ~v3_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0) [cnf transformation 121] {a:0,w:15,wCS:6000,sine:255,thAx:0,allAx:1,thDist:-1}
223. v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 127] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
234. r2_arytm_3(X0,k11_ordinal2(X0,X2)) | ~v3_ordinal1(X2) | ~v3_ordinal1(k11_ordinal2(X0,X2)) | ~v3_ordinal1(X0) [equality resolution 188] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:2,thDist:-2}
236. 1 <=> r2_arytm_3(sK0,k1_xboole_0) [avatar definition]
238. ~r2_arytm_3(sK0,k1_xboole_0) <- (~1) [avatar component clause 236] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
240. 2 <=> r2_arytm_3(1,sK0) [avatar definition]
242. ~r2_arytm_3(1,sK0) <- (~2) [avatar component clause 240] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
243. ~1 | ~2 [avatar split clause 161,240,236]
245. 3 <=> v7_ordinal1(sK0) [avatar definition]
247. v7_ordinal1(sK0) <- (3) [avatar component clause 245] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
248. 3 [avatar split clause 160,245]
250. 4 <=> v3_ordinal1(sK0) [avatar definition]
252. v3_ordinal1(sK0) <- (4) [avatar component clause 250] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
253. 4 [avatar split clause 159,250]
260. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
262. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 260] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
263. 6 [avatar split clause 163,260]
265. 7 <=> k4_ordinal1 = k5_numbers [avatar definition]
267. k4_ordinal1 = k5_numbers <- (7) [avatar component clause 265] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
268. 7 [avatar split clause 164,265]
270. 8 <=> m1_subset_1(k1_xboole_0,k4_ordinal1) [avatar definition]
272. m1_subset_1(k1_xboole_0,k4_ordinal1) <- (8) [avatar component clause 270] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
273. 8 [avatar split clause 165,270]
280. 10 <=> v3_ordinal1(k4_ordinal1) [avatar definition]
282. v3_ordinal1(k4_ordinal1) <- (10) [avatar component clause 280] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
283. 10 [avatar split clause 168,280]
295. 13 <=> m1_subset_1(1,k5_numbers) [avatar definition]
297. m1_subset_1(1,k5_numbers) <- (13) [avatar component clause 295] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
298. 13 [avatar split clause 170,295]
367. m1_subset_1(1,k4_ordinal1) <- (7, 13) [forward demodulation 297,267] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
369. 27 <=> m1_subset_1(1,k4_ordinal1) [avatar definition]
371. m1_subset_1(1,k4_ordinal1) <- (27) [avatar component clause 369] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
372. 27 | ~7 | ~13 [avatar split clause 367,295,265,369]
441. 38 <=> v7_ordinal1(k1_xboole_0) [avatar definition]
443. v7_ordinal1(k1_xboole_0) <- (38) [avatar component clause 441] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
502. ~r2_arytm_3(sK0,X0) | ~v1_xboole_0(X0) <- (~1) [superposition 238,194] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
503. v7_ordinal1(k1_xboole_0) <- (8) [resolution 197,272] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
506. 38 | ~8 [avatar split clause 503,270,441]
562. 55 <=> v3_ordinal1(1) [avatar definition]
564. v3_ordinal1(1) <- (55) [avatar component clause 562] {a:2,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
569. v3_ordinal1(1) | ~v3_ordinal1(k4_ordinal1) <- (27) [resolution 189,371] {a:1,w:4,wCS:1600,sine:1,thAx:0,allAx:2,thDist:-2}
581. v3_ordinal1(1) <- (10, 27) [subsumption resolution 569,282] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
582. 55 | ~10 | ~27 [avatar split clause 581,369,280,562]
721. r2_arytm_3(X0,k11_ordinal2(X0,X2)) | ~v3_ordinal1(X2) | ~v3_ordinal1(X0) [subsumption resolution 234,223] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
722. ~v3_ordinal1(X0) | ~v3_ordinal1(sK0) | ~v1_xboole_0(k11_ordinal2(sK0,X0)) <- (~1) [resolution 721,502] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
724. r2_arytm_3(1,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(1) | ~v3_ordinal1(X1) [superposition 721,184] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:3,thDist:-3}
727. r2_arytm_3(1,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(1) [duplicate literal removal 724] {a:1,w:7,wCS:2800,sine:1,thAx:0,allAx:3,thDist:-3}
729. ~v1_xboole_0(k11_ordinal2(sK0,X0)) | ~v3_ordinal1(X0) <- (~1, 4) [subsumption resolution 722,252] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
730. r2_arytm_3(1,X1) | ~v3_ordinal1(X1) <- (55) [subsumption resolution 727,564] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
800. k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0) [subsumption resolution 218,180] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
801. k9_ordinal3(X0,X1) = k9_ordinal3(X1,X0) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [subsumption resolution 800,180] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
806. k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0) [subsumption resolution 219,180] {a:0,w:13,wCS:5200,sine:1,thAx:0,allAx:1,thDist:-1}
807. k9_ordinal3(X0,X1) = k11_ordinal2(X0,X1) | ~v7_ordinal1(X1) | ~v7_ordinal1(X0) [subsumption resolution 806,180] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
809. k1_xboole_0 = k9_ordinal3(k1_xboole_0,X0) | ~v7_ordinal1(X0) | ~v7_ordinal1(k1_xboole_0) | ~v3_ordinal1(X0) [superposition 807,181] {a:1,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
815. ~v1_xboole_0(k9_ordinal3(sK0,X2)) | ~v3_ordinal1(X2) | ~v7_ordinal1(X2) | ~v7_ordinal1(sK0) <- (~1, 4) [superposition 729,807] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
822. k1_xboole_0 = k9_ordinal3(k1_xboole_0,X0) | ~v7_ordinal1(X0) | ~v3_ordinal1(X0) <- (38) [subsumption resolution 809,443] {a:1,w:9,wCS:3600,sine:1,thAx:0,allAx:2,thDist:-2}
823. k1_xboole_0 = k9_ordinal3(k1_xboole_0,X0) | ~v7_ordinal1(X0) <- (38) [subsumption resolution 822,180] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
833. ~v1_xboole_0(k9_ordinal3(sK0,X2)) | ~v3_ordinal1(X2) | ~v7_ordinal1(X2) <- (~1, 3, 4) [subsumption resolution 815,247] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
834. ~v1_xboole_0(k9_ordinal3(sK0,X2)) | ~v7_ordinal1(X2) <- (~1, 3, 4) [subsumption resolution 833,180] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
840. ~v1_xboole_0(k9_ordinal3(X0,sK0)) | ~v7_ordinal1(X0) | ~v7_ordinal1(sK0) | ~v7_ordinal1(X0) <- (~1, 3, 4) [superposition 834,801] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
843. ~v1_xboole_0(k9_ordinal3(X0,sK0)) | ~v7_ordinal1(X0) | ~v7_ordinal1(sK0) <- (~1, 3, 4) [duplicate literal removal 840] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
844. ~v1_xboole_0(k9_ordinal3(X0,sK0)) | ~v7_ordinal1(X0) <- (~1, 3, 4) [subsumption resolution 843,247] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
986. ~v1_xboole_0(k1_xboole_0) | ~v7_ordinal1(k1_xboole_0) | ~v7_ordinal1(sK0) <- (~1, 3, 4, 38) [superposition 844,823] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
996. ~v7_ordinal1(k1_xboole_0) | ~v7_ordinal1(sK0) <- (~1, 3, 4, 6, 38) [subsumption resolution 986,262] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
997. ~v7_ordinal1(sK0) <- (~1, 3, 4, 6, 38) [subsumption resolution 996,443] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
998. $false <- (~1, 3, 4, 6, 38) [subsumption resolution 997,247] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
999. 1 | ~3 | ~4 | ~6 | ~38 [avatar contradiction clause 998]
1002. ~v3_ordinal1(sK0) <- (~2, 55) [resolution 242,730] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1004. $false <- (~2, 4, 55) [subsumption resolution 1002,252] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1005. 2 | ~4 | ~55 [avatar contradiction clause 1004]
1006. $false [avatar sat refutation 243,248,253,263,268,273,283,298,372,506,582,999,1005] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:391,thDist:-391}
% SZS output end Proof for arytm_3__t9_arytm_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 62
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 75
% Generated clauses: 528
% Active clauses: 166
% Passive clauses: 254
% Final active clauses: 160
% Final passive clauses: 67
% Main loop iterations started: 187
% 
% >>> Simplifying Inferences
% Duplicate literals: 29
% Fw subsumption resolutions: 69
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 21
% Forward subsumptions: 93
% 
% >>> Generating Inferences
% Binary resolution: 219
% Unit resulting resolution: 2
% Forward superposition: 35
% Backward superposition: 23
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 87
% SAT solver unit clauses: 23
% SAT solver binary clauses: 44
% 
% Memory used [KB]: 5373
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward demodulation: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.002 s ) 

