i: [1,136,0,3,1,1,0,0] l38_arytm_3
a: [3,202,0,3,1,136]
s: 202
i: [1,137,0,2,1,0,0,1] spc1_boole
a: [3,207,0,2,1,137]
s: 207
i: [1,138,0,2,1,0,0,255] fc1_xboole_0
a: [3,212,0,2,1,138]
s: 212
i: [1,143,0,2,1,0,0,255] fc6_ordinal1
a: [3,232,0,2,1,143]
s: 232
i: [1,142,0,2,1,0,0,255] fc6_ordinal1
a: [3,237,0,2,1,142]
s: 237
i: [1,189,0,2,1,0,0,2] rc2_xboole_0
a: [3,257,0,2,1,189]
s: 257
i: [1,190,0,2,1,0,0,255] rc4_ordinal1
a: [3,262,0,2,1,190]
s: 262
i: [1,191,0,2,1,0,0,2] rc1_ordinal1
a: [3,267,0,2,1,191]
s: 267
i: [1,192,0,2,1,0,0,2] rc1_xboole_0
a: [3,272,0,2,1,192]
s: 272
i: [1,196,0,2,1,0,0,255] rc2_ordinal1
a: [3,277,0,2,1,196]
s: 277
i: [1,195,0,2,1,0,0,255] rc2_ordinal1
a: [3,282,0,2,1,195]
s: 282
i: [1,194,0,2,1,0,0,255] rc2_ordinal1
a: [3,287,0,2,1,194]
s: 287
i: [1,193,0,2,1,0,0,255] rc2_ordinal1
a: [3,292,0,2,1,193]
s: 292
i: [1,198,0,2,1,0,0,255] rc5_ordinal1
a: [3,297,0,2,1,198]
s: 297
i: [1,197,0,2,1,0,0,255] rc5_ordinal1
a: [3,302,0,2,1,197]
s: 302
i: [1,147,0,3,1,0,0,255] fc1_subset_1
s: 147
i: [1,178,0,3,1,0,0,255] rc2_subset_1
s: 178
i: [1,181,0,3,1,0,0,255] reflexivity_r1_tarski
s: 181
i: [1,139,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,217,0,3,1,139]
s: 217
i: [1,140,0,3,1,0,0,255] t1_numerals
a: [3,222,0,3,1,140]
s: 222
i: [1,146,0,3,1,0,0,255] spc1_numerals
a: [3,242,0,3,1,146]
s: 242
i: [1,145,0,3,1,0,0,255] spc1_numerals
a: [3,247,0,3,1,145]
s: 247
d: [2,304,0,3,1,2,22,247,217]
a: [3,308,0,3,1,304]
s: 308
i: [1,155,0,4,2,0,0,255] cc6_ordinal1
s: 155
i: [1,157,0,4,2,0,0,255] cc1_ordinal1
s: 157
i: [1,158,0,4,2,0,0,255] cc1_ordinal1
s: 158
d: [2,310,1,2,1,1,40,155,262]
a: [3,315,1,2,1,310]
s: 315
i: [1,161,0,4,2,0,0,255] cc9_ordinal1
s: 161
d: [2,311,1,2,1,1,40,155,297]
a: [3,320,1,2,1,311]
s: 320
i: [1,162,0,4,2,0,0,255] cc7_ordinal1
s: 162
i: [1,163,0,4,2,0,0,2] cc3_ordinal1
s: 163
d: [2,322,1,2,1,1,40,157,232]
a: [3,328,1,2,1,322]
s: 328
i: [1,176,0,4,1,0,0,3] existence_m1_subset_1
s: 176
d: [2,323,1,2,1,1,40,157,267]
a: [3,333,1,2,1,323]
s: 333
i: [1,180,0,4,1,0,0,255] rc3_subset_1
s: 180
i: [1,141,0,4,1,0,0,255] dt_k5_numbers
a: [3,227,0,4,1,141]
s: 227
d: [2,335,1,2,1,1,40,158,232]
a: [3,341,1,2,1,335]
s: 341
i: [1,144,0,4,1,0,0,255] spc1_numerals
a: [3,252,0,4,1,144]
s: 252
d: [2,336,1,2,1,1,40,158,267]
a: [3,346,1,2,1,336]
s: 346
d: [2,411,0,4,1,2,22,227,217]
a: [3,415,0,4,1,411]
s: 415
d: [2,417,0,4,1,2,22,252,217]
a: [3,421,0,4,1,417]
s: 421
d: [2,360,1,2,1,1,40,161,212]
a: [3,366,1,2,1,360]
s: 366
i: [1,151,0,5,2,0,0,255] rc1_subset_1
s: 151
d: [2,361,1,2,1,1,40,161,272]
a: [3,371,1,2,1,361]
s: 371
i: [1,156,0,5,2,0,0,1] t2_arytm_3
s: 156
i: [1,164,0,5,2,0,0,255] t6_boole
s: 164
d: [2,423,1,2,1,1,40,156,202]
a: [3,427,1,2,1,423]
s: 427
i: [1,167,0,5,2,0,0,255] cc8_ordinal1
s: 167
d: [2,385,1,2,1,1,40,162,212]
a: [3,391,1,2,1,385]
s: 391
i: [1,173,0,5,2,0,0,255] t7_boole
s: 173
i: [1,177,0,5,1,0,0,255] rc2_subset_1
s: 177
d: [2,386,1,2,1,1,40,162,272]
a: [3,396,1,2,1,386]
s: 396
i: [1,179,0,5,1,0,0,255] rc3_subset_1
s: 179
d: [2,398,1,2,1,1,40,163,212]
a: [3,404,1,2,1,398]
s: 404
i: [1,153,0,6,2,0,0,255] rc4_subset_1
s: 153
i: [1,168,0,6,3,0,0,255] cc2_ordinal1
s: 168
d: [2,399,1,2,1,1,40,163,272]
a: [3,409,1,2,1,399]
s: 409
i: [1,169,0,6,2,0,0,255] t1_subset
s: 169
d: [2,433,1,2,1,1,40,167,308]
a: [3,439,1,2,1,433]
s: 439
d: [2,480,2,0,0,2,52,427,439,155]
e: 480
d: [2,481,2,2,1,1,40,439,155]
d: [2,483,2,0,0,2,21,481,427]
e: 483
f: 483,433,423,304,145,139,136
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_3__l38_arytm_3
% SZS output start Proof for arytm_3__l38_arytm_3
1. r1_arytm_3(1,1) [input l38_arytm_3]
2. ~r1_arytm_3(1,1) [negated conjecture 1]
15. ! [X0] : (v7_ordinal1(X0) => v3_ordinal1(X0)) [input cc6_ordinal1]
16. ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0)) [input cc8_ordinal1]
32. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
56. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
57. ! [X0] : (v3_ordinal1(X0) => r1_arytm_3(1,X0)) [input t2_arytm_3]
58. ~r1_arytm_3(1,1) [flattening 2]
68. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 56]
77. ! [X0] : (v3_ordinal1(X0) | ~v7_ordinal1(X0)) [ennf transformation 15]
78. ! [X0] : (r1_arytm_3(1,X0) | ~v3_ordinal1(X0)) [ennf transformation 57]
88. ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1)) [ennf transformation 16]
136. ~r1_arytm_3(1,1) [cnf transformation 58] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
139. k4_ordinal1 = k5_numbers [cnf transformation 32] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
145. m1_subset_1(1,k5_numbers) [cnf transformation 68] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
155. ~v7_ordinal1(X0) | v3_ordinal1(X0) [cnf transformation 77] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
156. r1_arytm_3(1,X0) | ~v3_ordinal1(X0) [cnf transformation 78] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
167. ~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0) [cnf transformation 88] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
200. 1 <=> r1_arytm_3(1,1) [avatar definition]
202. ~r1_arytm_3(1,1) <- (~1) [avatar component clause 200] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
203. ~1 [avatar split clause 136,200]
215. 4 <=> k4_ordinal1 = k5_numbers [avatar definition]
217. k4_ordinal1 = k5_numbers <- (4) [avatar component clause 215] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
218. 4 [avatar split clause 139,215]
245. 10 <=> m1_subset_1(1,k5_numbers) [avatar definition]
247. m1_subset_1(1,k5_numbers) <- (10) [avatar component clause 245] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
248. 10 [avatar split clause 145,245]
304. m1_subset_1(1,k4_ordinal1) <- (4, 10) [forward demodulation 247,217] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
306. 22 <=> m1_subset_1(1,k4_ordinal1) [avatar definition]
308. m1_subset_1(1,k4_ordinal1) <- (22) [avatar component clause 306] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
309. 22 | ~4 | ~10 [avatar split clause 304,245,215,306]
423. ~v3_ordinal1(1) <- (~1) [resolution 156,202] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
425. 41 <=> v3_ordinal1(1) [avatar definition]
427. ~v3_ordinal1(1) <- (~41) [avatar component clause 425] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
428. ~41 | 1 [avatar split clause 423,200,425]
433. v7_ordinal1(1) <- (22) [resolution 167,308] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
437. 42 <=> v7_ordinal1(1) [avatar definition]
439. v7_ordinal1(1) <- (42) [avatar component clause 437] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
440. 42 | ~22 [avatar split clause 433,306,437]
481. v3_ordinal1(1) <- (42) [resolution 439,155] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
483. $false <- (~41, 42) [subsumption resolution 481,427] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
484. 41 | ~42 [avatar contradiction clause 483]
485. $false [avatar sat refutation 203,218,248,309,428,440,484] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:200,thDist:-200}
% SZS output end Proof for arytm_3__l38_arytm_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 56
% >>> Preprocessing
% Introduced skolems: 12
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 63
% Generated clauses: 156
% Active clauses: 56
% Passive clauses: 94
% Final active clauses: 56
% Final passive clauses: 35
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 1
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 35
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 54
% SAT solver unit clauses: 22
% SAT solver binary clauses: 27
% 
% Memory used [KB]: 9594
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

