i: [1,51,0,11,2,1,0,0] t14_zfmisc_1
a: [3,82,0,3,1,51]
s: 82
a: [3,79,0,8,1,51]
s: 79
i: [1,52,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,89,0,2,1,52]
s: 89
i: [1,53,0,2,1,0,0,2] fc1_xboole_0
a: [3,94,0,2,1,53]
s: 94
i: [1,71,0,2,1,0,0,2] rc2_xboole_0
a: [3,104,0,2,1,71]
s: 104
i: [1,72,0,2,1,0,0,2] rc1_xboole_0
a: [3,109,0,2,1,72]
s: 109
i: [1,55,0,3,1,0,0,1] fc2_xboole_0
s: 55
i: [1,54,0,3,1,0,0,255] d2_xboole_0
a: [3,99,0,3,1,54]
s: 99
r: 89
i: [1,74,0,4,1,0,0,1] d1_tarski
s: 74
i: [1,56,0,5,1,0,0,1] t4_boole
s: 56
i: [1,57,0,5,1,0,0,1] t3_boole
s: 57
i: [1,58,0,5,2,0,0,2] t6_boole
s: 58
i: [1,67,0,5,2,0,0,2] t7_boole
s: 67
i: [1,69,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 69
i: [1,70,0,6,2,0,0,2] symmetry_r1_xboole_0
s: 70
i: [1,59,0,7,2,0,0,2] l20_zfmisc_1
s: 59
d: [2,117,1,4,1,0,40,69,74]
s: 117
i: [1,66,0,7,3,0,0,1] t8_boole
s: 66
i: [1,68,0,7,2,0,0,2] l18_zfmisc_1
s: 68
i: [1,75,0,7,2,0,0,1] d1_tarski
s: 75
i: [1,60,0,8,2,0,0,1] t83_xboole_1
s: 60
d: [2,121,1,10,2,1,44,79,60]
d: [2,122,1,5,1,1,19,121]
a: [3,126,1,5,1,122]
s: 126
i: [1,61,0,8,2,0,0,1] t83_xboole_1
s: 61
d: [2,128,2,4,1,1,40,126,59]
a: [3,132,2,4,1,128]
s: 132
d: [2,139,3,0,0,2,52,82,132,75]
e: 139
d: [2,140,3,3,1,1,40,132,75]
d: [2,144,3,0,0,2,21,140,82]
e: 144
r: 79
r: 82
a: [3,83,0,3,1,51]
s: 83
r: 132
r: 126
a: [3,78,0,8,1,51]
s: 78
d: [2,141,3,4,1,1,40,132,69]
a: [3,150,3,4,1,141]
s: 150
d: [2,159,0,8,1,2,22,78,83]
a: [3,163,0,8,1,159]
s: 163
d: [2,153,1,5,1,2,23,126,83]
a: [3,156,1,5,1,153]
s: 156
d: [2,171,2,0,0,1,52,74,156,68]
e: 171
d: [2,172,2,4,1,1,40,156,68]
d: [2,175,2,0,0,1,21,172,74]
e: 175
i: [1,50,0,11,2,1,0,0] t14_zfmisc_1
d: [2,166,1,10,2,1,44,61,163]
d: [2,167,1,5,1,1,19,166]
a: [3,125,1,5,1,122]
d: [2,169,1,5,1,2,22,125,83]
f: 175,169,167,159,140,128,122,50,51
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t14_zfmisc_1
% SZS output start Proof for zfmisc_1__t14_zfmisc_1
1. ! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)) <=> X0 != X1) [input t14_zfmisc_1]
2. ~! [X0] : ! [X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)) <=> X0 != X1) [negated conjecture 1]
4. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
12. ! [X0] : ! [X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [input l18_zfmisc_1]
13. ! [X0] : ! [X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [input l20_zfmisc_1]
21. ! [X0] : ! [X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [input t83_xboole_1]
23. ~! [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)) <=> X0 != X1) [flattening 2]
24. ! [X0,X1] : (~r2_hidden(X0,X1) => r1_xboole_0(k1_tarski(X0),X1)) [flattening 13]
25. ! [X0,X1] : (r1_xboole_0(X0,X1) <=> k4_xboole_0(X0,X1) = X0) [flattening 21]
26. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 4]
29. ! [X0,X1] : ~(r2_hidden(X0,X1) & r1_xboole_0(k1_tarski(X0),X1)) [flattening 12]
30. ? [X0,X1] : (k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)) <~> X0 != X1) [ennf transformation 23]
32. ! [X0,X1] : (r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1)) [ennf transformation 24]
35. ! [X0,X1] : (~r2_hidden(X0,X1) | ~r1_xboole_0(k1_tarski(X0),X1)) [ennf transformation 29]
38. ? [X0,X1] : ((X0 = X1 | k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),k1_tarski(X1))) & (X0 != X1 | k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)))) [nnf transformation 30]
39. ? [X0,X1] : ((X0 = X1 | k1_tarski(X0) != k4_xboole_0(k1_tarski(X0),k1_tarski(X1))) & (X0 != X1 | k1_tarski(X0) = k4_xboole_0(k1_tarski(X0),k1_tarski(X1)))) => ((sK0 = sK1 | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1))) & (sK0 != sK1 | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)))) [choice axiom]
40. (sK0 = sK1 | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1))) & (sK0 != sK1 | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1))) [skolemisation 38,39]
41. ! [X0,X1] : ((r1_xboole_0(X0,X1) | k4_xboole_0(X0,X1) != X0) & (k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1))) [nnf transformation 25]
42. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 26]
43. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 42]
44. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) [choice axiom]
45. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK2(X0,X1) != X0 | ~r2_hidden(sK2(X0,X1),X1)) & (sK2(X0,X1) = X0 | r2_hidden(sK2(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 43,44]
50. sK0 != sK1 | k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [cnf transformation 40] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. sK0 = sK1 | k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [cnf transformation 40] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. r1_xboole_0(k1_tarski(X0),X1) | r2_hidden(X0,X1) [cnf transformation 32] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
60. k4_xboole_0(X0,X1) = X0 | ~r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. k4_xboole_0(X0,X1) != X0 | r1_xboole_0(X0,X1) [cnf transformation 41] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 45] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
63. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 45] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
68. ~r1_xboole_0(k1_tarski(X0),X1) | ~r2_hidden(X0,X1) [cnf transformation 35] {a:0,w:7,wCS:2800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
73. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 63] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
74. r2_hidden(X3,k1_tarski(X3)) [equality resolution 73] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 62] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. 1 <=> k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [avatar definition]
78. k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (1) [avatar component clause 77] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
79. k1_tarski(sK0) != k4_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (~1) [avatar component clause 77] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
81. 2 <=> sK0 = sK1 [avatar definition]
83. sK0 = sK1 <- (2) [avatar component clause 81] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
84. ~1 | 2 [avatar split clause 51,81,77]
85. 1 | ~2 [avatar split clause 50,81,77]
121. k1_tarski(sK0) != k1_tarski(sK0) | ~r1_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (~1) [superposition 79,60] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
122. ~r1_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (~1) [trivial inequality removal 121] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
124. 8 <=> r1_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) [avatar definition]
125. r1_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (8) [avatar component clause 124] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
126. ~r1_xboole_0(k1_tarski(sK0),k1_tarski(sK1)) <- (~8) [avatar component clause 124] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
127. ~8 | 1 [avatar split clause 122,77,124]
128. r2_hidden(sK0,k1_tarski(sK1)) <- (~8) [resolution 126,59] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
130. 9 <=> r2_hidden(sK0,k1_tarski(sK1)) [avatar definition]
132. r2_hidden(sK0,k1_tarski(sK1)) <- (9) [avatar component clause 130] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
133. 9 | 8 [avatar split clause 128,124,130]
140. sK0 = sK1 <- (9) [resolution 132,75] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
146. 2 | ~9 [avatar split clause 140,130,81]
155. 11 <=> r1_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) [avatar definition]
156. r1_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (11) [avatar component clause 155] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
159. k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (1, 2) [forward demodulation 78,83] {a:0,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
161. 12 <=> k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) [avatar definition]
163. k1_tarski(sK0) = k4_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (12) [avatar component clause 161] {a:0,w:8,wCS:800,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
164. 12 | ~1 | ~2 [avatar split clause 159,81,77,161]
166. k1_tarski(sK0) != k1_tarski(sK0) | r1_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (12) [superposition 61,163] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
167. r1_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (12) [trivial inequality removal 166] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
168. 11 | ~12 [avatar split clause 167,161,155]
169. r1_xboole_0(k1_tarski(sK0),k1_tarski(sK0)) <- (2, 8) [forward demodulation 125,83] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 11 | ~2 | ~8 [avatar split clause 169,124,81,155]
172. ~r2_hidden(sK0,k1_tarski(sK0)) <- (11) [resolution 156,68] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
175. $false <- (11) [subsumption resolution 172,74] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
176. ~11 [avatar contradiction clause 175]
177. $false [avatar sat refutation 84,85,127,133,146,164,168,170,176] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:92,thDist:-92}
% SZS output end Proof for zfmisc_1__t14_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 23
% Generated clauses: 75
% Active clauses: 27
% Passive clauses: 38
% Final active clauses: 22
% Final passive clauses: 8
% Main loop iterations started: 29
% 
% >>> Simplifying Inferences
% Trivial inequalities: 4
% Fw subsumption resolutions: 2
% Fw demodulations: 3
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 1
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 2
% Forward superposition: 3
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 7
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

