i: [1,27,0,6,2,1,0,0] t12_gate_1
a: [3,37,0,2,1,27]
s: 37
i: [1,26,0,6,2,1,0,0] t12_gate_1
a: [3,46,0,2,1,26]
s: 46
a: [3,40,0,4,1,27]
s: 40
i: [1,32,0,2,1,0,0,1] rc2_xboole_0
a: [3,52,0,2,1,32]
s: 52
i: [1,33,0,2,1,0,0,1] rc1_xboole_0
a: [3,57,0,2,1,33]
s: 57
i: [1,30,0,6,2,0,0,1] fc13_gate_1
s: 30
i: [1,28,0,7,1,0,0,255] commutativity_k6_gate_1
s: 28
i: [1,29,0,7,1,0,0,255] commutativity_k6_gate_1
s: 29
d: [2,59,1,6,2,0,44,30,28]
s: 59
i: [1,31,0,8,3,0,0,1] fc14_gate_1
s: 31
d: [2,63,1,0,0,3,52,46,37,40,31]
e: 63
d: [2,64,1,4,2,1,40,31,40]
d: [2,70,1,2,1,2,21,64,37]
d: [2,71,1,0,0,3,21,70,46]
e: 71
r: 40
r: 46
a: [3,45,0,2,1,26]
s: 45
a: [3,41,0,4,1,27]
s: 41
d: [2,73,1,0,0,2,52,45,41,30]
e: 73
d: [2,75,1,2,1,1,40,41,30]
d: [2,77,1,0,0,2,21,75,45]
e: 77
r: 37
r: 45
i: [1,22,0,8,3,1,0,0] t12_gate_1
d: [2,74,2,2,1,1,40,41,59]
f: 75,74,71,22,26,27
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_1__t12_gate_1
% SZS output start Proof for gate_1__t12_gate_1
1. ! [X0] : ! [X1] : (~(v1_xboole_0(k6_gate_1(X0,X1)) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k6_gate_1(X0,X1)))) [input t12_gate_1]
2. ~! [X0] : ! [X1] : (~(v1_xboole_0(k6_gate_1(X0,X1)) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k6_gate_1(X0,X1)))) [negated conjecture 1]
3. ! [X0,X1] : k6_gate_1(X0,X1) = k6_gate_1(X1,X0) [input commutativity_k6_gate_1]
7. ! [X0,X1] : (v1_xboole_0(X0) => ~v1_xboole_0(k6_gate_1(X0,X1))) [input fc13_gate_1]
8. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => v1_xboole_0(k6_gate_1(X0,X1))) [input fc14_gate_1]
11. ~! [X0,X1] : (~(v1_xboole_0(k6_gate_1(X0,X1)) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0))) & ~(~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k6_gate_1(X0,X1)))) [flattening 2]
12. ? [X0,X1] : ((v1_xboole_0(k6_gate_1(X0,X1)) & (v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k6_gate_1(X0,X1)))) [ennf transformation 11]
13. ! [X0,X1] : (~v1_xboole_0(k6_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 7]
14. ! [X0,X1] : (v1_xboole_0(k6_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 8]
15. ! [X0,X1] : (v1_xboole_0(k6_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 14]
16. ? [X0,X1] : ((v1_xboole_0(k6_gate_1(X0,X1)) & (v1_xboole_0(X1) | v1_xboole_0(X0))) | (~v1_xboole_0(X1) & ~v1_xboole_0(X0) & ~v1_xboole_0(k6_gate_1(X0,X1)))) => ((v1_xboole_0(k6_gate_1(sK0,sK1)) & (v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k6_gate_1(sK0,sK1)))) [choice axiom]
17. (v1_xboole_0(k6_gate_1(sK0,sK1)) & (v1_xboole_0(sK1) | v1_xboole_0(sK0))) | (~v1_xboole_0(sK1) & ~v1_xboole_0(sK0) & ~v1_xboole_0(k6_gate_1(sK0,sK1))) [skolemisation 12,16]
22. v1_xboole_0(sK1) | v1_xboole_0(sK0) | ~v1_xboole_0(k6_gate_1(sK0,sK1)) [cnf transformation 17] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
26. v1_xboole_0(k6_gate_1(sK0,sK1)) | ~v1_xboole_0(sK0) [cnf transformation 17] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
27. v1_xboole_0(k6_gate_1(sK0,sK1)) | ~v1_xboole_0(sK1) [cnf transformation 17] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
28. k6_gate_1(X0,X1) = k6_gate_1(X1,X0) [cnf transformation 3] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
30. ~v1_xboole_0(k6_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 13] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
31. v1_xboole_0(k6_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 15] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
35. 1 <=> v1_xboole_0(sK1) [avatar definition]
37. ~v1_xboole_0(sK1) <- (~1) [avatar component clause 35] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
39. 2 <=> v1_xboole_0(k6_gate_1(sK0,sK1)) [avatar definition]
40. ~v1_xboole_0(k6_gate_1(sK0,sK1)) <- (~2) [avatar component clause 39] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. v1_xboole_0(k6_gate_1(sK0,sK1)) <- (2) [avatar component clause 39] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
42. ~1 | 2 [avatar split clause 27,39,35]
44. 3 <=> v1_xboole_0(sK0) [avatar definition]
46. ~v1_xboole_0(sK0) <- (~3) [avatar component clause 44] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
47. ~3 | 2 [avatar split clause 26,39,44]
48. ~2 | 3 | 1 [avatar split clause 22,35,44,39]
59. ~v1_xboole_0(k6_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 30,28] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
64. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (~2) [resolution 31,40] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
70. v1_xboole_0(sK0) <- (~1, ~2) [subsumption resolution 64,37] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
71. $false <- (~1, ~2, ~3) [subsumption resolution 70,46] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
72. 1 | 2 | 3 [avatar contradiction clause 71]
74. ~v1_xboole_0(sK1) <- (2) [resolution 41,59] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
75. ~v1_xboole_0(sK0) <- (2) [resolution 41,30] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
79. ~1 | ~2 [avatar split clause 74,39,35]
80. ~3 | ~2 [avatar split clause 75,39,44]
81. $false [avatar sat refutation 42,47,48,72,79,80] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:48,thDist:-48}
% SZS output end Proof for gate_1__t12_gate_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 39
% Active clauses: 11
% Passive clauses: 15
% Final active clauses: 7
% Final passive clauses: 3
% Main loop iterations started: 12
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 5
% Unit resulting resolution: 2
% Forward superposition: 4
% Backward superposition: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 2
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

