i: [1,19,0,2,1,1,0,0] t40_orders_1
a: [3,42,0,2,1,19]
s: 42
i: [1,20,0,3,1,1,0,0] t40_orders_1
a: [3,37,0,3,1,20]
s: 37
i: [1,21,0,4,1,1,0,0] t40_orders_1
a: [3,32,0,4,1,21]
s: 32
i: [1,22,0,5,2,0,0,1] dt_k2_relat_1
s: 22
i: [1,23,0,7,2,0,0,255] involutiveness_k2_relat_1
s: 23
i: [1,26,0,8,3,0,0,1] d6_orders_1
s: 26
d: [2,45,1,5,2,1,40,26,37]
d: [2,46,1,3,1,2,21,45,42]
a: [3,50,1,3,1,46]
s: 50
i: [1,27,0,8,3,0,0,1] d6_orders_1
s: 27
d: [2,52,1,5,2,1,40,27,37]
d: [2,53,1,3,1,2,21,52,42]
a: [3,57,1,3,1,53]
s: 57
i: [1,24,0,9,3,0,0,2] l58_orders_1
s: 24
i: [1,25,0,9,3,0,0,2] l59_orders_1
s: 25
d: [2,59,1,12,4,0,44,24,23]
d: [2,60,1,9,3,0,21,59,22]
s: 60
i: [1,28,0,11,4,0,0,1] d6_orders_1
s: 28
d: [2,61,1,12,4,0,44,25,23]
d: [2,62,1,9,3,0,21,61,22]
s: 62
d: [2,67,1,16,5,0,40,28,25]
d: [2,68,1,13,4,0,21,67,22]
s: 68
d: [2,72,2,14,5,0,40,68,24]
d: [2,74,2,12,4,0,18,72]
s: 74
d: [2,76,3,0,0,4,52,42,50,57,32,74]
e: 76
d: [2,77,3,8,3,1,40,74,32]
d: [2,82,3,5,2,2,21,77,57]
d: [2,83,3,3,1,3,21,82,42]
d: [2,84,3,0,0,4,21,83,50]
e: 84
f: 84,53,46,19,20,21
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t40_orders_1
% SZS output start Proof for orders_1__t40_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) => r1_orders_1(k2_relat_1(X0),X1))) [input t40_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) => r1_orders_1(k2_relat_1(X0),X1))) [negated conjecture 1]
3. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d6_orders_1]
4. ! [X0] : (v1_relat_1(X0) => v1_relat_1(k2_relat_1(X0))) [input dt_k2_relat_1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_relat_2(X0,X1) => r1_relat_2(k2_relat_1(X0),X1))) [input l58_orders_1]
7. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r8_relat_2(X0,X1) => r8_relat_2(k2_relat_1(X0),X1))) [input l59_orders_1]
8. ? [X0] : (? [X1] : (~r1_orders_1(k2_relat_1(X0),X1) & r1_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
9. ! [X0] : (v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0)) [ennf transformation 4]
11. ! [X0] : (! [X1] : (r1_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 6]
12. ! [X0] : (! [X1] : (r8_relat_2(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
13. ! [X0] : (! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 3]
14. ? [X0] : (? [X1] : (~r1_orders_1(k2_relat_1(X0),X1) & r1_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (~r1_orders_1(k2_relat_1(sK0),X1) & r1_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
15. ? [X1] : (~r1_orders_1(k2_relat_1(sK0),X1) & r1_orders_1(sK0,X1)) => (~r1_orders_1(k2_relat_1(sK0),sK1) & r1_orders_1(sK0,sK1)) [choice axiom]
16. (~r1_orders_1(k2_relat_1(sK0),sK1) & r1_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 8,15,14]
17. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | (~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 13]
18. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 17]
19. v1_relat_1(sK0) [cnf transformation 16] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
20. r1_orders_1(sK0,sK1) [cnf transformation 16] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
21. ~r1_orders_1(k2_relat_1(sK0),sK1) [cnf transformation 16] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
22. v1_relat_1(k2_relat_1(X0)) | ~v1_relat_1(X0) [cnf transformation 9] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
24. r1_relat_2(k2_relat_1(X0),X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 11] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
25. r8_relat_2(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 12] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
26. ~r1_orders_1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
27. ~r1_orders_1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
28. ~r8_relat_2(X0,X1) | r1_orders_1(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. 1 <=> r1_orders_1(k2_relat_1(sK0),sK1) [avatar definition]
32. ~r1_orders_1(k2_relat_1(sK0),sK1) <- (~1) [avatar component clause 30] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
33. ~1 [avatar split clause 21,30]
35. 2 <=> r1_orders_1(sK0,sK1) [avatar definition]
37. r1_orders_1(sK0,sK1) <- (2) [avatar component clause 35] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
38. 2 [avatar split clause 20,35]
40. 3 <=> v1_relat_1(sK0) [avatar definition]
42. v1_relat_1(sK0) <- (3) [avatar component clause 40] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
43. 3 [avatar split clause 19,40]
45. r1_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 26,37] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
46. r1_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 45,42] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
48. 4 <=> r1_relat_2(sK0,sK1) [avatar definition]
50. r1_relat_2(sK0,sK1) <- (4) [avatar component clause 48] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
51. 4 | ~2 | ~3 [avatar split clause 46,40,35,48]
52. r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 27,37] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
53. r8_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 52,42] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
55. 5 <=> r8_relat_2(sK0,sK1) [avatar definition]
57. r8_relat_2(sK0,sK1) <- (5) [avatar component clause 55] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
58. 5 | ~2 | ~3 [avatar split clause 53,40,35,55]
67. r1_orders_1(k2_relat_1(X0),X1) | ~r1_relat_2(k2_relat_1(X0),X1) | ~v1_relat_1(k2_relat_1(X0)) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 28,25] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
68. ~r1_relat_2(k2_relat_1(X0),X1) | r1_orders_1(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [subsumption resolution 67,22] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
72. r1_orders_1(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 68,24] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
74. r1_orders_1(k2_relat_1(X0),X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r1_relat_2(X0,X1) [duplicate literal removal 72] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
77. ~r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) | ~r1_relat_2(sK0,sK1) <- (~1) [resolution 74,32] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
82. ~v1_relat_1(sK0) | ~r1_relat_2(sK0,sK1) <- (~1, 5) [subsumption resolution 77,57] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
83. ~r1_relat_2(sK0,sK1) <- (~1, 3, 5) [subsumption resolution 82,42] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
84. $false <- (~1, 3, 4, 5) [subsumption resolution 83,50] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
85. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 84]
87. $false [avatar sat refutation 33,38,43,51,58,85] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:49,thDist:-49}
% SZS output end Proof for orders_1__t40_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 10
% Generated clauses: 46
% Active clauses: 16
% Passive clauses: 17
% Final active clauses: 16
% Final passive clauses: 1
% Main loop iterations started: 16
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 5
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 1
% Forward superposition: 6
% Backward superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

