i: [1,70,0,4,1,1,0,0] l48_subset_1
a: [3,108,0,4,1,70]
s: 108
i: [1,69,0,6,2,1,0,0] l48_subset_1
s: 69
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,113,0,2,1,71]
s: 113
i: [1,72,0,2,1,0,0,255] fc1_xboole_0
a: [3,118,0,2,1,72]
s: 118
i: [1,98,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,128,0,2,1,98]
s: 128
i: [1,99,0,2,1,0,0,2] rc2_xboole_0
a: [3,133,0,2,1,99]
s: 133
i: [1,100,0,2,1,0,0,2] rc1_xboole_0
a: [3,138,0,2,1,100]
s: 138
i: [1,102,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,143,0,2,1,102]
s: 143
i: [1,101,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,148,0,2,1,101]
s: 148
i: [1,74,0,3,1,0,0,1] fc1_subset_1
s: 74
i: [1,95,0,3,1,0,0,1] rc2_subset_1
s: 95
i: [1,96,0,3,1,0,0,255] reflexivity_r1_tarski
s: 96
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,123,0,3,1,73]
s: 123
r: 113
i: [1,75,0,4,2,0,0,255] cc2_zfmisc_1
s: 75
i: [1,78,0,4,2,0,0,255] cc1_zfmisc_1
s: 78
i: [1,93,0,4,1,0,0,1] existence_m1_subset_1
s: 93
d: [2,151,1,2,1,1,40,75,118]
a: [3,157,1,2,1,151]
s: 157
i: [1,77,0,5,2,0,0,1] rc1_subset_1
s: 77
d: [2,153,1,2,1,1,40,75,138]
a: [3,162,1,2,1,153]
s: 162
i: [1,79,0,5,2,0,0,255] t6_boole
s: 79
i: [1,92,0,5,2,0,0,1] t7_boole
s: 92
d: [2,152,1,3,1,0,40,75,95]
s: 152
i: [1,94,0,5,1,0,0,1] rc2_subset_1
s: 94
i: [1,97,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 97
i: [1,76,0,7,2,0,0,1] rc1_subset_1
s: 76
i: [1,82,0,7,3,0,0,1] d1_subset_1
s: 82
d: [2,167,1,5,2,0,40,82,93]
s: 167
i: [1,83,0,7,3,0,0,1] d1_subset_1
s: 83
d: [2,170,2,5,2,0,40,167,75]
s: 170
i: [1,91,0,7,3,0,0,255] t8_boole
s: 91
i: [1,103,0,7,2,0,0,255] d1_zfmisc_1
s: 103
i: [1,104,0,7,2,0,0,255] d1_zfmisc_1
s: 104
i: [1,80,0,8,3,0,0,1] d1_subset_1
s: 80
d: [2,180,1,8,2,0,40,80,94]
d: [2,182,1,5,1,0,21,180,74]
s: 182
i: [1,81,0,8,3,0,0,1] d1_subset_1
s: 81
d: [2,186,0,6,2,0,21,81,92]
s: 186
d: [2,183,2,4,1,0,40,182,104]
s: 183
i: [1,85,0,8,2,0,0,255] d3_tarski
s: 85
d: [2,189,1,8,2,0,40,85,69]
s: 189
i: [1,86,0,8,2,0,0,255] d3_tarski
s: 86
i: [1,84,0,9,3,0,0,255] d3_tarski
s: 84
d: [2,196,2,5,2,0,40,189,92]
a: [3,200,2,2,1,196]
s: 200
i: [1,89,0,14,3,0,0,255] d1_zfmisc_1
s: 89
d: [2,207,2,6,2,0,40,86,189]
d: [2,208,2,3,1,0,18,207]
a: [3,213,2,3,1,208]
s: 213
i: [1,90,0,14,3,0,0,255] d1_zfmisc_1
s: 90
d: [2,192,1,5,2,0,40,85,92]
s: 192
d: [2,194,2,8,2,0,40,189,186]
s: 194
d: [2,179,1,6,2,0,40,80,93]
s: 179
d: [2,225,2,6,2,0,40,179,69]
a: [3,233,2,4,1,225]
s: 233
d: [2,164,1,7,3,0,44,79,79]
s: 164
d: [2,174,1,7,2,0,40,103,97]
s: 174
d: [2,240,3,4,1,1,40,233,186]
a: [3,246,3,4,1,240]
s: 246
d: [2,178,1,10,3,0,40,80,76]
d: [2,181,1,7,2,0,21,178,74]
s: 181
d: [2,241,3,4,1,1,40,233,97]
a: [3,251,3,4,1,241]
s: 251
d: [2,187,1,7,2,0,40,186,103]
s: 187
d: [2,260,3,0,0,2,52,213,108,187]
e: 260
d: [2,261,2,3,1,1,40,187,108]
d: [2,265,2,0,0,2,21,261,213]
e: 265
f: 265,208,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for subset_1__l48_subset_1
% SZS output start Proof for subset_1__l48_subset_1
1. ! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1)) => m1_subset_1(X0,k1_zfmisc_1(X1))) [input l48_subset_1]
2. ~! [X0] : ! [X1] : (! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1)) => m1_subset_1(X0,k1_zfmisc_1(X1))) [negated conjecture 1]
6. ! [X0] : ! [X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [input d1_subset_1]
7. ! [X0] : ! [X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [input d1_zfmisc_1]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
25. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
27. ~! [X0,X1] : (! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1)) => m1_subset_1(X0,k1_zfmisc_1(X1))) [flattening 2]
28. ! [X0,X1] : ((v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> v1_xboole_0(X1))) & (~v1_xboole_0(X0) => (m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)))) [flattening 6]
29. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 9]
30. ! [X0,X1] : (k1_zfmisc_1(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> r1_tarski(X2,X0))) [flattening 7]
32. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 25]
34. ? [X0,X1] : (~m1_subset_1(X0,k1_zfmisc_1(X1)) & ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 27]
39. ! [X0,X1] : (((m1_subset_1(X1,X0) <=> v1_xboole_0(X1)) | ~v1_xboole_0(X0)) & ((m1_subset_1(X1,X0) <=> r2_hidden(X1,X0)) | v1_xboole_0(X0))) [ennf transformation 28]
40. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 29]
42. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 32]
44. ? [X0,X1] : (~m1_subset_1(X0,k1_zfmisc_1(X1)) & ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) => (~m1_subset_1(sK0,k1_zfmisc_1(sK1)) & ! [X2] : (r2_hidden(X2,sK1) | ~r2_hidden(X2,sK0))) [choice axiom]
45. ~m1_subset_1(sK0,k1_zfmisc_1(sK1)) & ! [X2] : (r2_hidden(X2,sK1) | ~r2_hidden(X2,sK0)) [skolemisation 34,44]
48. ! [X0,X1] : ((((m1_subset_1(X1,X0) | ~v1_xboole_0(X1)) & (v1_xboole_0(X1) | ~m1_subset_1(X1,X0))) | ~v1_xboole_0(X0)) & (((m1_subset_1(X1,X0) | ~r2_hidden(X1,X0)) & (r2_hidden(X1,X0) | ~m1_subset_1(X1,X0))) | v1_xboole_0(X0))) [nnf transformation 39]
49. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 40]
50. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 49]
51. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) [choice axiom]
52. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 50,51]
53. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | ~r1_tarski(X2,X0)) & (r1_tarski(X2,X0) | ~r2_hidden(X2,X1))) | k1_zfmisc_1(X0) != X1)) [nnf transformation 30]
54. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [rectify 53]
55. ! [X1,X0] : (? [X2] : ((~r1_tarski(X2,X0) | ~r2_hidden(X2,X1)) & (r1_tarski(X2,X0) | r2_hidden(X2,X1))) => ((~r1_tarski(sK4(X0,X1),X0) | ~r2_hidden(sK4(X0,X1),X1)) & (r1_tarski(sK4(X0,X1),X0) | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
56. ! [X0,X1] : ((k1_zfmisc_1(X0) = X1 | ((~r1_tarski(sK4(X0,X1),X0) | ~r2_hidden(sK4(X0,X1),X1)) & (r1_tarski(sK4(X0,X1),X0) | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | ~r1_tarski(X3,X0)) & (r1_tarski(X3,X0) | ~r2_hidden(X3,X1))) | k1_zfmisc_1(X0) != X1)) [skolemisation 54,55]
69. ~r2_hidden(X2,sK0) | r2_hidden(X2,sK1) [cnf transformation 45] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. ~m1_subset_1(sK0,k1_zfmisc_1(sK1)) [cnf transformation 45] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
81. m1_subset_1(X1,X0) | ~r2_hidden(X1,X0) | v1_xboole_0(X0) [cnf transformation 48] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:1,thDist:-1}
85. r2_hidden(sK3(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 52] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. ~r2_hidden(sK3(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 52] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
88. r2_hidden(X3,X1) | ~r1_tarski(X3,X0) | k1_zfmisc_1(X0) != X1 [cnf transformation 56] {a:0,w:10,wCS:4000,sine:255,thAx:0,allAx:2,thDist:-2}
92. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 42] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. r2_hidden(X3,k1_zfmisc_1(X0)) | ~r1_tarski(X3,X0) [equality resolution 88] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
106. 1 <=> m1_subset_1(sK0,k1_zfmisc_1(sK1)) [avatar definition]
108. ~m1_subset_1(sK0,k1_zfmisc_1(sK1)) <- (~1) [avatar component clause 106] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. ~1 [avatar split clause 70,106]
186. ~r2_hidden(X1,X0) | m1_subset_1(X1,X0) [subsumption resolution 81,92] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
187. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [resolution 186,103] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
189. r2_hidden(sK3(sK0,X0),sK1) | r1_tarski(sK0,X0) [resolution 85,69] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
207. r1_tarski(sK0,sK1) | r1_tarski(sK0,sK1) [resolution 86,189] {a:2,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
208. r1_tarski(sK0,sK1) [duplicate literal removal 207] {a:2,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
211. 14 <=> r1_tarski(sK0,sK1) [avatar definition]
213. r1_tarski(sK0,sK1) <- (14) [avatar component clause 211] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
214. 14 [avatar split clause 208,211]
261. ~r1_tarski(sK0,sK1) <- (~1) [resolution 187,108] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
265. $false <- (~1, 14) [subsumption resolution 261,213] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
266. 1 | ~14 [avatar contradiction clause 265]
267. $false [avatar sat refutation 109,214,266] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:86,thDist:-86}
% SZS output end Proof for subset_1__l48_subset_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 25
% Introduced skolems: 11
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 123
% Active clauses: 52
% Passive clauses: 69
% Final active clauses: 51
% Final passive clauses: 14
% Main loop iterations started: 55
% 
% >>> Simplifying Inferences
% Duplicate literals: 4
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 59
% Unit resulting resolution: 1
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 10
% SAT solver binary clauses: 8
% 
% Memory used [KB]: 9466
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

