i: [1,204,0,2,1,1,0,0] t15_euclmetr
a: [3,282,0,2,1,204]
s: 282
i: [1,203,0,2,1,1,0,0] t15_euclmetr
a: [3,287,0,2,1,203]
s: 287
i: [1,202,0,2,1,1,0,0] t15_euclmetr
a: [3,292,0,2,1,202]
s: 292
i: [1,201,0,2,1,1,0,0] t15_euclmetr
a: [3,297,0,2,1,201]
s: 297
i: [1,200,0,2,1,1,0,0] t15_euclmetr
a: [3,302,0,2,1,200]
s: 302
i: [1,199,0,2,1,1,0,0] t15_euclmetr
a: [3,307,0,2,1,199]
s: 307
i: [1,205,0,2,1,0,0,255] spc1_boole
a: [3,312,0,2,1,205]
s: 312
i: [1,210,0,2,1,0,0,255] fc6_ordinal1
a: [3,332,0,2,1,210]
s: 332
i: [1,209,0,2,1,0,0,255] fc6_ordinal1
a: [3,337,0,2,1,209]
s: 337
i: [1,264,0,2,1,0,0,255] existence_l1_analoaf
a: [3,357,0,2,1,264]
s: 357
i: [1,265,0,2,1,0,0,255] existence_l1_struct_0
a: [3,362,0,2,1,265]
s: 362
i: [1,266,0,2,1,0,0,1] existence_l1_analmetr
a: [3,367,0,2,1,266]
s: 367
i: [1,267,0,2,1,0,0,255] rc4_ordinal1
a: [3,372,0,2,1,267]
s: 372
i: [1,268,0,2,1,0,0,255] rc1_ordinal1
a: [3,377,0,2,1,268]
s: 377
i: [1,272,0,2,1,0,0,255] rc2_ordinal1
a: [3,382,0,2,1,272]
s: 382
i: [1,271,0,2,1,0,0,255] rc2_ordinal1
a: [3,387,0,2,1,271]
s: 387
i: [1,270,0,2,1,0,0,255] rc2_ordinal1
a: [3,392,0,2,1,270]
s: 392
i: [1,269,0,2,1,0,0,255] rc2_ordinal1
a: [3,397,0,2,1,269]
s: 397
i: [1,274,0,2,1,0,0,255] rc5_ordinal1
a: [3,402,0,2,1,274]
s: 402
i: [1,273,0,2,1,0,0,255] rc5_ordinal1
a: [3,407,0,2,1,273]
s: 407
i: [1,276,0,2,1,0,0,1] rc2_analmetr
a: [3,412,0,2,1,276]
s: 412
i: [1,275,0,2,1,0,0,1] rc2_analmetr
a: [3,417,0,2,1,275]
s: 417
i: [1,277,0,2,1,0,0,255] rc4_xreal_0
a: [3,422,0,2,1,277]
s: 422
i: [1,258,0,3,1,0,0,255] reflexivity_r1_tarski
s: 258
i: [1,206,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,317,0,3,1,206]
s: 317
i: [1,207,0,3,1,0,0,255] t1_numerals
a: [3,322,0,3,1,207]
s: 322
i: [1,213,0,3,1,0,0,255] spc1_numerals
a: [3,342,0,3,1,213]
s: 342
i: [1,212,0,3,1,0,0,255] spc1_numerals
a: [3,347,0,3,1,212]
s: 347
i: [1,278,0,3,1,0,0,255] rc1_real_1
a: [3,427,0,3,1,278]
s: 427
d: [2,429,0,3,1,2,22,347,317]
a: [3,433,0,3,1,429]
s: 433
i: [1,214,0,4,2,0,0,255] dt_l1_analoaf
s: 214
d: [2,435,1,2,1,1,40,214,357]
a: [3,439,1,2,1,435]
s: 439
i: [1,230,0,4,2,0,0,255] dt_l1_analmetr
s: 230
d: [2,441,1,2,1,1,40,230,297]
a: [3,447,1,2,1,441]
s: 447
i: [1,232,0,4,2,0,0,255] cc6_ordinal1
s: 232
i: [1,233,0,4,2,0,0,255] cc1_ordinal1
s: 233
d: [2,459,2,2,1,1,40,447,214]
a: [3,463,2,2,1,459]
s: 463
i: [1,234,0,4,2,0,0,255] cc1_ordinal1
s: 234
d: [2,442,1,2,1,1,40,230,367]
a: [3,452,1,2,1,442]
s: 452
i: [1,236,0,4,2,0,0,255] cc9_ordinal1
s: 236
i: [1,237,0,4,2,0,0,255] cc7_ordinal1
s: 237
d: [2,443,1,2,1,1,40,230,417]
a: [3,457,1,2,1,443]
s: 457
i: [1,238,0,4,2,0,0,255] cc3_ordinal1
s: 238
d: [2,465,1,2,1,1,40,232,372]
a: [3,470,1,2,1,465]
s: 470
i: [1,257,0,4,1,0,0,255] existence_m1_subset_1
s: 257
i: [1,208,0,4,1,0,0,255] dt_k5_numbers
a: [3,327,0,4,1,208]
s: 327
d: [2,466,1,2,1,1,40,232,402]
a: [3,475,1,2,1,466]
s: 475
i: [1,211,0,4,1,0,0,255] spc1_numerals
a: [3,352,0,4,1,211]
s: 352
d: [2,477,1,2,1,1,40,233,332]
a: [3,483,1,2,1,477]
s: 483
d: [2,545,0,4,1,2,22,327,317]
a: [3,549,0,4,1,545]
s: 549
d: [2,563,0,4,1,2,22,352,317]
a: [3,567,0,4,1,563]
s: 567
d: [2,478,1,2,1,1,40,233,377]
a: [3,488,1,2,1,478]
s: 488
i: [1,239,0,5,2,0,0,255] t6_boole
s: 239
d: [2,490,1,2,1,1,40,234,332]
a: [3,496,1,2,1,490]
s: 496
i: [1,240,0,5,2,0,0,255] cc8_ordinal1
s: 240
i: [1,254,0,5,2,0,0,255] t7_boole
s: 254
d: [2,491,1,2,1,1,40,234,377]
a: [3,501,1,2,1,491]
s: 501
i: [1,216,0,6,3,0,0,255] cc7_struct_0
s: 216
d: [2,509,1,2,1,1,40,236,422]
a: [3,513,1,2,1,509]
s: 513
i: [1,217,0,6,3,0,0,255] cc5_struct_0
s: 217
i: [1,219,0,6,3,0,0,255] cc17_struct_0
s: 219
d: [2,515,1,2,1,1,40,237,422]
a: [3,519,1,2,1,515]
s: 519
i: [1,220,0,6,3,0,0,255] cc2_struct_0
s: 220
d: [2,527,1,2,1,1,40,238,422]
a: [3,531,1,2,1,527]
s: 531
i: [1,221,0,6,3,0,0,255] cc6_struct_0
s: 221
i: [1,222,0,6,3,0,0,255] cc1_struct_0
s: 222
d: [2,571,1,2,1,1,40,240,433]
a: [3,577,1,2,1,571]
s: 577
i: [1,225,0,6,3,0,0,255] cc4_struct_0
s: 225
d: [2,572,1,2,1,1,40,240,322]
a: [3,582,1,2,1,572]
s: 582
i: [1,249,0,6,3,0,0,255] cc2_ordinal1
s: 249
i: [1,250,0,6,2,0,0,255] t1_subset
s: 250
d: [2,589,1,4,2,1,40,217,362]
a: [3,595,1,2,1,589]
s: 595
i: [1,259,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 259
d: [2,503,2,2,1,1,40,452,214]
a: [3,507,2,2,1,503]
s: 507
i: [1,223,0,7,3,0,0,255] cc8_struct_0
s: 223
i: [1,226,0,7,3,0,0,255] cc9_struct_0
s: 226
d: [2,521,2,2,1,1,40,457,214]
a: [3,525,2,2,1,521]
s: 525
i: [1,227,0,7,3,0,0,255] cc11_struct_0
s: 227
d: [2,533,2,2,1,1,40,470,234]
a: [3,538,2,2,1,533]
s: 538
i: [1,228,0,7,3,0,0,255] cc11_struct_0
s: 228
i: [1,235,0,7,3,0,0,255] cc5_ordinal1
s: 235
d: [2,676,1,4,2,1,40,235,342]
a: [3,685,1,2,1,676]
s: 685
i: [1,252,0,7,2,0,0,255] t3_subset
s: 252
d: [2,534,2,2,1,1,40,470,233]
a: [3,543,2,2,1,534]
s: 543
i: [1,253,0,7,3,0,0,255] t8_boole
s: 253
i: [1,215,0,8,3,0,0,255] cc10_ordinal1
s: 215
d: [2,698,1,4,2,1,40,215,549]
a: [3,703,1,2,1,698]
s: 703
i: [1,251,0,8,3,0,0,255] t2_subset
s: 251
d: [2,710,1,5,2,1,40,251,342]
a: [3,723,1,2,1,710]
s: 723
d: [2,751,2,0,0,2,52,703,723,236]
e: 751
d: [2,752,2,0,0,2,52,685,723,238]
e: 752
d: [2,753,2,2,1,1,40,723,238]
d: [2,758,2,0,0,2,21,753,685]
e: 758
d: [2,755,2,2,1,1,40,723,236]
d: [2,760,2,0,0,2,21,755,703]
e: 760
r: 685
r: 703
r: 723
i: [1,229,0,9,4,0,0,255] cc10_struct_0
s: 229
i: [1,256,0,9,3,0,0,255] t5_subset
s: 256
d: [2,681,1,4,2,1,40,235,427]
a: [3,694,1,2,1,681]
s: 694
r: 694
i: [1,241,0,10,5,0,0,2] t15_conmetr
s: 241
s: 685
i: [1,243,0,10,5,0,0,1] t10_euclmetr
s: 243
i: [1,244,0,10,5,0,0,1] t10_euclmetr
s: 244
a: [3,693,1,2,1,681]
s: 693
i: [1,245,0,10,5,0,0,1] t11_euclmetr
s: 245
d: [2,792,1,8,4,1,40,245,287]
d: [2,793,1,6,3,2,21,792,297]
d: [2,794,1,4,2,3,21,793,302]
d: [2,795,1,2,1,4,21,794,307]
a: [3,799,1,2,1,795]
s: 799
i: [1,246,0,10,5,0,0,1] t11_euclmetr
s: 246
i: [1,247,0,10,5,0,0,1] t8_euclmetr
s: 247
d: [2,802,1,8,4,1,40,247,292]
d: [2,803,1,6,3,2,21,802,297]
d: [2,804,1,4,2,3,21,803,302]
d: [2,805,1,2,1,4,21,804,307]
a: [3,809,1,2,1,805]
s: 809
i: [1,248,0,10,5,0,0,1] t8_euclmetr
s: 248
d: [2,811,2,8,4,1,40,809,241]
d: [2,812,2,6,3,2,21,811,297]
d: [2,813,2,4,2,3,21,812,302]
d: [2,814,2,2,1,4,21,813,307]
a: [3,818,2,2,1,814]
s: 818
i: [1,255,0,10,3,0,0,255] t4_subset
s: 255
i: [1,242,0,12,6,0,0,2] t13_conmetr
s: 242
d: [2,824,3,10,5,1,40,242,818]
d: [2,825,3,8,4,2,21,824,799]
d: [2,826,3,6,3,3,21,825,297]
d: [2,827,3,4,2,4,21,826,302]
d: [2,828,3,2,1,5,21,827,307]
a: [3,832,3,2,1,828]
s: 832
d: [2,834,4,0,0,5,52,307,302,297,282,832,244]
e: 834
d: [2,835,4,8,4,1,40,832,244]
d: [2,837,4,6,3,2,21,835,282]
d: [2,838,4,4,2,3,21,837,297]
d: [2,839,4,2,1,4,21,838,302]
d: [2,840,4,0,0,5,21,839,307]
e: 840
f: 840,828,814,805,795,199,200,201,202,203,204
% Refutation found. Thanks to Tanya!
% SZS status Theorem for euclmetr__t15_euclmetr
% SZS output start Proof for euclmetr__t15_euclmetr
1. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => ((v3_euclmetr(X0) & v1_euclmetr(X0)) => v2_euclmetr(X0))) [input t15_euclmetr]
2. ~! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => ((v3_euclmetr(X0) & v1_euclmetr(X0)) => v2_euclmetr(X0))) [negated conjecture 1]
61. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => (v2_euclmetr(X0) <=> v2_conmetr(X0))) [input t10_euclmetr]
62. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => (v3_euclmetr(X0) <=> v1_conaffm(X0))) [input t11_euclmetr]
63. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => ((v1_conaffm(X0) & v1_conmetr(X0)) => v2_conmetr(X0))) [input t13_conmetr]
64. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => (v3_conaffm(X0) => v1_conmetr(X0))) [input t15_conmetr]
74. ! [X0] : ((l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => (v1_euclmetr(X0) <=> v3_conaffm(X0))) [input t8_euclmetr]
104. ? [X0] : ((~v2_euclmetr(X0) & (v3_euclmetr(X0) & v1_euclmetr(X0))) & (l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0))) [ennf transformation 2]
105. ? [X0] : (~v2_euclmetr(X0) & v3_euclmetr(X0) & v1_euclmetr(X0) & l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) [flattening 104]
141. ! [X0] : ((v1_conmetr(X0) | ~v3_conaffm(X0)) | (~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0))) [ennf transformation 64]
142. ! [X0] : (v1_conmetr(X0) | ~v3_conaffm(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [flattening 141]
143. ! [X0] : ((v2_conmetr(X0) | (~v1_conaffm(X0) | ~v1_conmetr(X0))) | (~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0))) [ennf transformation 63]
144. ! [X0] : (v2_conmetr(X0) | ~v1_conaffm(X0) | ~v1_conmetr(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [flattening 143]
145. ! [X0] : ((v2_euclmetr(X0) <=> v2_conmetr(X0)) | (~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0))) [ennf transformation 61]
146. ! [X0] : ((v2_euclmetr(X0) <=> v2_conmetr(X0)) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [flattening 145]
147. ! [X0] : ((v3_euclmetr(X0) <=> v1_conaffm(X0)) | (~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0))) [ennf transformation 62]
148. ! [X0] : ((v3_euclmetr(X0) <=> v1_conaffm(X0)) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [flattening 147]
149. ! [X0] : ((v1_euclmetr(X0) <=> v3_conaffm(X0)) | (~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0))) [ennf transformation 74]
150. ! [X0] : ((v1_euclmetr(X0) <=> v3_conaffm(X0)) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [flattening 149]
169. ? [X0] : (~v2_euclmetr(X0) & v3_euclmetr(X0) & v1_euclmetr(X0) & l1_analmetr(X0) & v3_analmetr(X0) & ~v2_struct_0(X0)) => (~v2_euclmetr(sK0) & v3_euclmetr(sK0) & v1_euclmetr(sK0) & l1_analmetr(sK0) & v3_analmetr(sK0) & ~v2_struct_0(sK0)) [choice axiom]
170. ~v2_euclmetr(sK0) & v3_euclmetr(sK0) & v1_euclmetr(sK0) & l1_analmetr(sK0) & v3_analmetr(sK0) & ~v2_struct_0(sK0) [skolemisation 105,169]
171. ! [X0] : (((v2_euclmetr(X0) | ~v2_conmetr(X0)) & (v2_conmetr(X0) | ~v2_euclmetr(X0))) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [nnf transformation 146]
172. ! [X0] : (((v3_euclmetr(X0) | ~v1_conaffm(X0)) & (v1_conaffm(X0) | ~v3_euclmetr(X0))) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [nnf transformation 148]
173. ! [X0] : (((v1_euclmetr(X0) | ~v3_conaffm(X0)) & (v3_conaffm(X0) | ~v1_euclmetr(X0))) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0)) [nnf transformation 150]
199. ~v2_struct_0(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
200. v3_analmetr(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
201. l1_analmetr(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
202. v1_euclmetr(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
203. v3_euclmetr(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
204. ~v2_euclmetr(sK0) [cnf transformation 170] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
241. ~v3_conaffm(X0) | v1_conmetr(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0) [cnf transformation 142] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
242. ~v1_conmetr(X0) | ~v1_conaffm(X0) | v2_conmetr(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0) [cnf transformation 144] {a:0,w:12,wCS:4800,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
244. ~v2_conmetr(X0) | v2_euclmetr(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0) [cnf transformation 171] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
245. ~v3_euclmetr(X0) | v1_conaffm(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0) [cnf transformation 172] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
247. ~v1_euclmetr(X0) | v3_conaffm(X0) | ~l1_analmetr(X0) | ~v3_analmetr(X0) | v2_struct_0(X0) [cnf transformation 173] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
280. 1 <=> v2_euclmetr(sK0) [avatar definition]
282. ~v2_euclmetr(sK0) <- (~1) [avatar component clause 280] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
283. ~1 [avatar split clause 204,280]
285. 2 <=> v3_euclmetr(sK0) [avatar definition]
287. v3_euclmetr(sK0) <- (2) [avatar component clause 285] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
288. 2 [avatar split clause 203,285]
290. 3 <=> v1_euclmetr(sK0) [avatar definition]
292. v1_euclmetr(sK0) <- (3) [avatar component clause 290] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
293. 3 [avatar split clause 202,290]
295. 4 <=> l1_analmetr(sK0) [avatar definition]
297. l1_analmetr(sK0) <- (4) [avatar component clause 295] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
298. 4 [avatar split clause 201,295]
300. 5 <=> v3_analmetr(sK0) [avatar definition]
302. v3_analmetr(sK0) <- (5) [avatar component clause 300] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
303. 5 [avatar split clause 200,300]
305. 6 <=> v2_struct_0(sK0) [avatar definition]
307. ~v2_struct_0(sK0) <- (~6) [avatar component clause 305] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
308. ~6 [avatar split clause 199,305]
792. v1_conaffm(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (2) [resolution 245,287] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
793. v1_conaffm(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (2, 4) [subsumption resolution 792,297] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
794. v1_conaffm(sK0) | v2_struct_0(sK0) <- (2, 4, 5) [subsumption resolution 793,302] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
795. v1_conaffm(sK0) <- (2, 4, 5, ~6) [subsumption resolution 794,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
797. 85 <=> v1_conaffm(sK0) [avatar definition]
799. v1_conaffm(sK0) <- (85) [avatar component clause 797] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
800. 85 | ~2 | ~4 | ~5 | 6 [avatar split clause 795,305,300,295,285,797]
802. v3_conaffm(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (3) [resolution 247,292] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
803. v3_conaffm(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (3, 4) [subsumption resolution 802,297] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
804. v3_conaffm(sK0) | v2_struct_0(sK0) <- (3, 4, 5) [subsumption resolution 803,302] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
805. v3_conaffm(sK0) <- (3, 4, 5, ~6) [subsumption resolution 804,307] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
807. 86 <=> v3_conaffm(sK0) [avatar definition]
809. v3_conaffm(sK0) <- (86) [avatar component clause 807] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
810. 86 | ~3 | ~4 | ~5 | 6 [avatar split clause 805,305,300,295,290,807]
811. v1_conmetr(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (86) [resolution 809,241] {a:2,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
812. v1_conmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (4, 86) [subsumption resolution 811,297] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
813. v1_conmetr(sK0) | v2_struct_0(sK0) <- (4, 5, 86) [subsumption resolution 812,302] {a:2,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
814. v1_conmetr(sK0) <- (4, 5, ~6, 86) [subsumption resolution 813,307] {a:2,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
816. 87 <=> v1_conmetr(sK0) [avatar definition]
818. v1_conmetr(sK0) <- (87) [avatar component clause 816] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
819. 87 | ~4 | ~5 | 6 | ~86 [avatar split clause 814,807,305,300,295,816]
824. ~v1_conaffm(sK0) | v2_conmetr(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (87) [resolution 242,818] {a:3,w:10,wCS:1000,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
825. v2_conmetr(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (85, 87) [subsumption resolution 824,799] {a:3,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
826. v2_conmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (4, 85, 87) [subsumption resolution 825,297] {a:3,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
827. v2_conmetr(sK0) | v2_struct_0(sK0) <- (4, 5, 85, 87) [subsumption resolution 826,302] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
828. v2_conmetr(sK0) <- (4, 5, ~6, 85, 87) [subsumption resolution 827,307] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
830. 88 <=> v2_conmetr(sK0) [avatar definition]
832. v2_conmetr(sK0) <- (88) [avatar component clause 830] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
833. 88 | ~4 | ~5 | 6 | ~85 | ~87 [avatar split clause 828,816,797,305,300,295,830]
835. v2_euclmetr(sK0) | ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (88) [resolution 832,244] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
837. ~l1_analmetr(sK0) | ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (~1, 88) [subsumption resolution 835,282] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
838. ~v3_analmetr(sK0) | v2_struct_0(sK0) <- (~1, 4, 88) [subsumption resolution 837,297] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
839. v2_struct_0(sK0) <- (~1, 4, 5, 88) [subsumption resolution 838,302] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
840. $false <- (~1, 4, 5, ~6, 88) [subsumption resolution 839,307] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
841. 1 | ~4 | ~5 | 6 | ~88 [avatar contradiction clause 840]
842. $false [avatar sat refutation 283,288,293,298,303,308,800,810,819,833,841] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:436,thDist:-436}
% SZS output end Proof for euclmetr__t15_euclmetr
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 73
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 6
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 80
% Generated clauses: 282
% Active clauses: 100
% Passive clauses: 142
% Final active clauses: 96
% Final passive clauses: 35
% Main loop iterations started: 107
% 
% >>> Simplifying Inferences
% Duplicate literals: 3
% Fw subsumption resolutions: 23
% Fw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 6
% Forward subsumptions: 18
% 
% >>> Generating Inferences
% Binary resolution: 82
% Unit resulting resolution: 3
% Forward superposition: 3
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 13
% Split components: 26
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 95
% SAT solver unit clauses: 31
% SAT solver binary clauses: 36
% 
% Memory used [KB]: 9850
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% backward demodulation: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

