i: [1,138,0,2,1,1,0,0] t3_real
a: [3,199,0,2,1,138]
s: 199
i: [1,137,0,2,1,1,0,0] t3_real
a: [3,204,0,2,1,137]
s: 204
i: [1,135,0,2,1,1,0,0] t3_real
a: [3,214,0,2,1,135]
s: 214
i: [1,134,0,2,1,1,0,0] t3_real
a: [3,219,0,2,1,134]
s: 219
i: [1,136,0,3,1,1,0,0] t3_real
a: [3,209,0,3,1,136]
s: 209
i: [1,139,0,2,1,0,0,3] fc1_xboole_0
a: [3,224,0,2,1,139]
s: 224
i: [1,145,0,2,1,0,0,4] fc6_ordinal1
a: [3,249,0,2,1,145]
s: 249
i: [1,144,0,2,1,0,0,4] fc6_ordinal1
a: [3,254,0,2,1,144]
s: 254
i: [1,179,0,2,1,0,0,3] rc2_xboole_0
a: [3,266,0,2,1,179]
s: 266
i: [1,180,0,2,1,0,0,5] rc4_ordinal1
a: [3,271,0,2,1,180]
s: 271
i: [1,181,0,2,1,0,0,5] rc1_ordinal1
a: [3,276,0,2,1,181]
s: 276
i: [1,182,0,2,1,0,0,1] rc1_xreal_0
a: [3,281,0,2,1,182]
s: 281
i: [1,183,0,2,1,0,0,3] rc1_xboole_0
a: [3,286,0,2,1,183]
s: 286
i: [1,187,0,2,1,0,0,255] rc2_ordinal1
a: [3,291,0,2,1,187]
s: 291
i: [1,186,0,2,1,0,0,255] rc2_ordinal1
a: [3,296,0,2,1,186]
s: 296
i: [1,185,0,2,1,0,0,255] rc2_ordinal1
a: [3,301,0,2,1,185]
s: 301
i: [1,184,0,2,1,0,0,255] rc2_ordinal1
a: [3,306,0,2,1,184]
s: 306
i: [1,189,0,2,1,0,0,5] rc5_ordinal1
a: [3,311,0,2,1,189]
s: 311
i: [1,188,0,2,1,0,0,5] rc5_ordinal1
a: [3,316,0,2,1,188]
s: 316
i: [1,192,0,2,1,0,0,1] rc3_xreal_0
a: [3,321,0,2,1,192]
s: 321
i: [1,191,0,2,1,0,0,1] rc3_xreal_0
a: [3,326,0,2,1,191]
s: 326
i: [1,190,0,2,1,0,0,1] rc3_xreal_0
a: [3,331,0,2,1,190]
s: 331
i: [1,195,0,2,1,0,0,2] rc4_xreal_0
a: [3,336,0,2,1,195]
s: 336
i: [1,194,0,2,1,0,0,2] rc4_xreal_0
a: [3,341,0,2,1,194]
s: 341
i: [1,193,0,2,1,0,0,2] rc4_xreal_0
a: [3,346,0,2,1,193]
s: 346
i: [1,171,0,3,1,0,0,255] reflexivity_r1_tarski
s: 171
i: [1,140,0,3,1,0,0,2] redefinition_k6_numbers
a: [3,229,0,3,1,140]
s: 229
r: 224
d: [2,348,0,2,1,2,23,224,229]
a: [3,352,0,2,1,348]
s: 352
i: [1,141,0,3,1,0,0,3] redefinition_k5_numbers
a: [3,234,0,3,1,141]
s: 234
i: [1,149,0,4,2,0,0,5] cc6_ordinal1
s: 149
i: [1,150,0,4,2,0,0,5] cc2_xreal_0
s: 150
d: [2,354,1,2,1,1,40,149,271]
a: [3,359,1,2,1,354]
s: 359
i: [1,151,0,4,2,0,0,255] cc1_ordinal1
s: 151
d: [2,355,1,2,1,1,40,149,311]
a: [3,364,1,2,1,355]
s: 364
i: [1,152,0,4,2,0,0,255] cc1_ordinal1
s: 152
i: [1,154,0,4,2,0,0,2] cc4_xreal_0
s: 154
d: [2,428,1,2,1,1,40,154,219]
a: [3,436,1,2,1,428]
s: 436
i: [1,156,0,4,2,0,0,255] cc9_ordinal1
s: 156
d: [2,429,1,2,1,1,40,154,214]
a: [3,441,1,2,1,429]
s: 441
i: [1,157,0,4,2,0,0,5] cc7_ordinal1
s: 157
i: [1,158,0,4,2,0,0,5] cc3_ordinal1
s: 158
d: [2,366,1,2,1,1,40,150,271]
a: [3,371,1,2,1,366]
s: 371
i: [1,170,0,4,1,0,0,4] existence_m1_subset_1
s: 170
d: [2,367,1,2,1,1,40,150,311]
a: [3,376,1,2,1,367]
s: 376
i: [1,142,0,4,1,0,0,3] dt_k5_numbers
a: [3,239,0,4,1,142]
s: 239
i: [1,143,0,4,1,0,0,2] dt_k6_numbers
a: [3,244,0,4,1,143]
s: 244
d: [2,378,1,2,1,1,40,151,249]
a: [3,385,1,2,1,378]
s: 385
d: [2,514,0,4,1,2,22,239,234]
a: [3,518,0,4,1,514]
s: 518
d: [2,380,1,2,1,1,40,151,276]
a: [3,395,1,2,1,380]
s: 395
d: [2,520,0,4,1,2,22,244,234]
a: [3,524,0,4,1,520]
s: 524
i: [1,159,0,5,2,0,0,3] t6_boole
s: 159
d: [2,403,1,2,1,1,40,152,249]
a: [3,411,1,2,1,403]
s: 411
i: [1,160,0,5,2,0,0,3] cc1_xreal_0
s: 160
d: [2,405,1,2,1,1,40,152,276]
a: [3,421,1,2,1,405]
s: 421
i: [1,161,0,5,2,0,0,4] cc8_ordinal1
s: 161
i: [1,167,0,5,2,0,0,255] t7_boole
s: 167
d: [2,430,1,2,1,1,40,154,281]
a: [3,446,1,2,1,430]
s: 446
i: [1,177,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,261,0,5,2,177]
s: 261
d: [2,448,1,2,1,1,40,156,346]
a: [3,454,1,2,1,448]
s: 454
d: [2,526,0,5,2,1,22,159,229]
s: 526
i: [1,162,0,6,3,0,0,255] cc2_ordinal1
s: 162
d: [2,449,1,2,1,1,40,156,286]
a: [3,459,1,2,1,449]
s: 459
i: [1,163,0,6,2,0,0,255] t1_subset
s: 163
d: [2,450,1,2,1,1,40,156,352]
a: [3,464,1,2,1,450]
s: 464
i: [1,172,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 172
i: [1,147,0,7,3,0,0,1] d7_xxreal_0
s: 147
d: [2,466,1,2,1,1,40,157,346]
a: [3,472,1,2,1,466]
s: 472
i: [1,148,0,7,3,0,0,1] d7_xxreal_0
s: 148
d: [2,467,1,2,1,1,40,157,286]
a: [3,477,1,2,1,467]
s: 477
i: [1,153,0,7,3,0,0,5] cc5_ordinal1
s: 153
i: [1,165,0,7,2,0,0,255] t3_subset
s: 165
d: [2,468,1,2,1,1,40,157,352]
a: [3,482,1,2,1,468]
s: 482
i: [1,166,0,7,3,0,0,3] t8_boole
s: 166
d: [2,484,1,2,1,1,40,158,346]
a: [3,490,1,2,1,484]
s: 490
i: [1,146,0,8,3,0,0,255] cc10_ordinal1
s: 146
i: [1,164,0,8,3,0,0,255] t2_subset
s: 164
d: [2,485,1,2,1,1,40,158,286]
a: [3,495,1,2,1,485]
s: 495
i: [1,169,0,9,3,0,0,255] t5_subset
s: 169
d: [2,486,1,2,1,1,40,158,352]
a: [3,500,1,2,1,486]
s: 500
i: [1,168,0,10,3,0,0,255] t4_subset
s: 168
i: [1,178,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 178
d: [2,546,1,6,3,1,40,147,261]
d: [2,548,1,4,2,1,18,546]
a: [3,552,1,2,1,548]
s: 552
i: [1,176,0,14,4,0,0,3] existence_m2_subset_1
s: 176
d: [2,597,1,4,2,1,40,146,518]
a: [3,602,1,2,1,597]
s: 602
i: [1,155,0,15,6,0,0,1] l2_real
s: 155
i: [1,173,0,15,5,0,0,3] dt_m2_subset_1
s: 173
d: [2,668,1,12,5,1,40,155,209]
d: [2,673,1,10,4,2,21,668,214]
d: [2,674,1,8,3,3,21,673,219]
s: 674
i: [1,174,0,15,5,0,0,3] redefinition_m2_subset_1
s: 174
d: [2,694,2,9,4,3,40,674,148]
d: [2,697,2,7,3,4,21,694,204]
d: [2,698,2,5,2,5,21,697,436]
a: [3,702,2,3,1,698]
s: 702
d: [2,714,3,0,0,3,52,441,199,702,147]
e: 714
d: [2,715,3,4,2,1,40,702,147]
d: [2,719,3,2,1,2,21,715,199]
d: [2,720,3,0,0,3,21,719,441]
e: 720
d: [2,575,2,2,1,1,40,482,150]
f: 720,698,575,468,429,428,348,140,139,134,135,136,137,138
% Refutation found. Thanks to Tanya!
% SZS status Theorem for real__t3_real
% SZS output start Proof for real__t3_real
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(v3_xxreal_0(X1) & ~v3_xxreal_0(X0) & r1_xxreal_0(X0,X1)))) [input t3_real]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ~(v3_xxreal_0(X1) & ~v3_xxreal_0(X0) & r1_xxreal_0(X0,X1)))) [negated conjecture 1]
8. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
12. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
15. ! [X0] : (v1_xboole_0(X0) => v7_ordinal1(X0)) [input cc7_ordinal1]
19. ! [X0] : (v1_xxreal_0(X0) => (v3_xxreal_0(X0) <=> ~r1_xxreal_0(k6_numbers,X0))) [input d7_xxreal_0]
30. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
32. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input l2_real]
43. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
68. ? [X0] : (? [X1] : ((v3_xxreal_0(X1) & ~v3_xxreal_0(X0) & r1_xxreal_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
69. ? [X0] : (? [X1] : (v3_xxreal_0(X1) & ~v3_xxreal_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 68]
71. ! [X0] : ((v3_xxreal_0(X0) <=> ~r1_xxreal_0(k6_numbers,X0)) | ~v1_xxreal_0(X0)) [ennf transformation 19]
73. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 8]
76. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 12]
77. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 32]
78. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 77]
80. ! [X0] : (v7_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 15]
107. ? [X0] : (? [X1] : (v3_xxreal_0(X1) & ~v3_xxreal_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (v3_xxreal_0(X1) & ~v3_xxreal_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
108. ? [X1] : (v3_xxreal_0(X1) & ~v3_xxreal_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) => (v3_xxreal_0(sK1) & ~v3_xxreal_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) [choice axiom]
109. (v3_xxreal_0(sK1) & ~v3_xxreal_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 69,108,107]
110. ! [X0] : (((v3_xxreal_0(X0) | r1_xxreal_0(k6_numbers,X0)) & (~r1_xxreal_0(k6_numbers,X0) | ~v3_xxreal_0(X0))) | ~v1_xxreal_0(X0)) [nnf transformation 71]
134. v1_xreal_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
135. v1_xreal_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
136. r1_xxreal_0(sK0,sK1) [cnf transformation 109] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
137. ~v3_xxreal_0(sK0) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
138. v3_xxreal_0(sK1) [cnf transformation 109] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
139. v1_xboole_0(k1_xboole_0) [cnf transformation 30] {a:0,w:2,wCS:800,sine:3,thAx:0,allAx:1,thDist:-1}
140. k6_numbers = k1_xboole_0 [cnf transformation 43] {a:0,w:3,wCS:1200,sine:2,thAx:0,allAx:1,thDist:-1}
147. ~r1_xxreal_0(k6_numbers,X0) | ~v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
148. r1_xxreal_0(k6_numbers,X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 110] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
150. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 73] {a:0,w:4,wCS:1600,nSel:1,sine:5,thAx:0,allAx:1,thDist:-1}
154. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 76] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
155. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 78] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
157. ~v1_xboole_0(X0) | v7_ordinal1(X0) [cnf transformation 80] {a:0,w:4,wCS:1600,nSel:1,sine:5,thAx:0,allAx:1,thDist:-1}
197. 1 <=> v3_xxreal_0(sK1) [avatar definition]
199. v3_xxreal_0(sK1) <- (1) [avatar component clause 197] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
200. 1 [avatar split clause 138,197]
202. 2 <=> v3_xxreal_0(sK0) [avatar definition]
204. ~v3_xxreal_0(sK0) <- (~2) [avatar component clause 202] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
205. ~2 [avatar split clause 137,202]
207. 3 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
209. r1_xxreal_0(sK0,sK1) <- (3) [avatar component clause 207] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
210. 3 [avatar split clause 136,207]
212. 4 <=> v1_xreal_0(sK1) [avatar definition]
214. v1_xreal_0(sK1) <- (4) [avatar component clause 212] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
215. 4 [avatar split clause 135,212]
217. 5 <=> v1_xreal_0(sK0) [avatar definition]
219. v1_xreal_0(sK0) <- (5) [avatar component clause 217] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
220. 5 [avatar split clause 134,217]
222. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
224. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 222] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
225. 6 [avatar split clause 139,222]
227. 7 <=> k6_numbers = k1_xboole_0 [avatar definition]
229. k6_numbers = k1_xboole_0 <- (7) [avatar component clause 227] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
230. 7 [avatar split clause 140,227]
348. v1_xboole_0(k6_numbers) <- (6, 7) [backward demodulation 224,229] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
350. 32 <=> v1_xboole_0(k6_numbers) [avatar definition]
352. v1_xboole_0(k6_numbers) <- (32) [avatar component clause 350] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
353. 32 | ~6 | ~7 [avatar split clause 348,227,222,350]
428. v1_xxreal_0(sK0) <- (5) [resolution 154,219] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
429. v1_xxreal_0(sK1) <- (4) [resolution 154,214] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
434. 45 <=> v1_xxreal_0(sK0) [avatar definition]
436. v1_xxreal_0(sK0) <- (45) [avatar component clause 434] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
437. 45 | ~5 [avatar split clause 428,217,434]
439. 46 <=> v1_xxreal_0(sK1) [avatar definition]
441. v1_xxreal_0(sK1) <- (46) [avatar component clause 439] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
442. 46 | ~4 [avatar split clause 429,212,439]
468. v7_ordinal1(k6_numbers) <- (32) [resolution 157,352] {a:1,w:2,wCS:800,sine:5,thAx:0,allAx:2,thDist:-2}
480. 53 <=> v7_ordinal1(k6_numbers) [avatar definition]
482. v7_ordinal1(k6_numbers) <- (53) [avatar component clause 480] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
483. 53 | ~32 [avatar split clause 468,350,480]
575. v1_xreal_0(k6_numbers) <- (53) [resolution 482,150] {a:2,w:2,wCS:800,sine:5,thAx:0,allAx:3,thDist:-3}
579. 66 <=> v1_xreal_0(k6_numbers) [avatar definition]
582. 66 | ~53 [avatar split clause 575,480,579]
668. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X0) <- (3) [resolution 155,209] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
673. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X0) <- (3, 4) [subsumption resolution 668,214] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
674. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK1) | ~v1_xreal_0(X0) <- (3, 4, 5) [subsumption resolution 673,219] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
694. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (3, 4, 5) [resolution 674,148] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
697. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) | ~v1_xxreal_0(sK0) <- (~2, 3, 4, 5) [subsumption resolution 694,204] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
698. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) <- (~2, 3, 4, 5, 45) [subsumption resolution 697,436] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
700. 82 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
702. r1_xxreal_0(k6_numbers,sK1) <- (82) [avatar component clause 700] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
703. ~66 | 82 | 2 | ~3 | ~4 | ~5 | ~45 [avatar split clause 698,434,217,212,207,202,700,579]
715. ~v3_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (82) [resolution 702,147] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
719. ~v1_xxreal_0(sK1) <- (1, 82) [subsumption resolution 715,199] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
720. $false <- (1, 46, 82) [subsumption resolution 719,441] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
721. ~1 | ~46 | ~82 [avatar contradiction clause 720]
723. $false [avatar sat refutation 200,205,210,215,220,225,230,353,437,442,483,582,703,721] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:376,thDist:-376}
% SZS output end Proof for real__t3_real
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 53
% >>> Preprocessing
% Introduced skolems: 13
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 62
% Generated clauses: 256
% Active clauses: 86
% Passive clauses: 131
% Final active clauses: 85
% Final passive clauses: 41
% Main loop iterations started: 89
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 16
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 14
% 
% >>> Generating Inferences
% Binary resolution: 78
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 7
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 10
% Split components: 20
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 85
% SAT solver unit clauses: 31
% SAT solver binary clauses: 40
% 
% Memory used [KB]: 9722
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% forward subsumption: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

