i: [1,48,0,2,1,1,0,0] t27_orders_1
a: [3,76,0,2,1,48]
s: 76
i: [1,47,0,2,1,1,0,0] t27_orders_1
a: [3,81,0,2,1,47]
s: 81
i: [1,49,0,4,1,1,0,0] t27_orders_1
a: [3,71,0,4,1,49]
s: 71
i: [1,59,0,4,1,0,0,255] fc6_relat_1
s: 59
i: [1,60,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 60
i: [1,50,0,6,3,0,0,1] d5_orders_1
s: 50
d: [2,83,1,4,2,1,40,50,76]
d: [2,84,1,2,1,2,21,83,81]
a: [3,88,1,2,1,84]
s: 88
i: [1,51,0,6,3,0,0,1] d5_orders_1
s: 51
d: [2,90,1,4,2,1,40,51,76]
d: [2,91,1,2,1,2,21,90,81]
a: [3,95,1,2,1,91]
s: 95
i: [1,52,0,6,3,0,0,1] d5_orders_1
s: 52
i: [1,53,0,6,3,0,0,1] d5_orders_1
s: 53
d: [2,97,1,4,2,1,40,52,76]
d: [2,98,1,2,1,2,21,97,81]
a: [3,102,1,2,1,98]
s: 102
i: [1,62,0,6,2,0,0,1] dt_k2_wellord1
s: 62
d: [2,104,1,4,2,1,40,53,76]
d: [2,105,1,2,1,2,21,104,81]
a: [3,109,1,2,1,105]
s: 109
i: [1,63,0,6,2,0,0,3] fc1_relat_1
s: 63
i: [1,61,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 61
d: [2,112,1,6,2,0,44,63,61]
s: 112
i: [1,56,0,8,3,0,0,2] t16_wellord1
s: 56
d: [2,113,1,6,2,0,44,63,61]
s: 113
i: [1,57,0,8,3,0,0,2] t15_wellord1
s: 57
i: [1,58,0,8,3,0,0,2] t17_wellord1
s: 58
i: [1,64,0,10,4,0,0,2] fc18_relat_2
s: 64
i: [1,65,0,10,4,0,0,2] fc19_relat_2
s: 65
d: [2,118,1,10,4,0,44,64,61]
s: 118
i: [1,55,0,11,2,0,0,255] d6_wellord1
s: 55
i: [1,54,0,12,6,0,0,1] d5_orders_1
s: 54
d: [2,130,1,14,5,0,44,64,55]
d: [2,136,1,12,4,0,18,130]
d: [2,141,1,8,3,0,21,136,59]
s: 141
i: [1,66,0,12,5,0,0,2] fc5_relat_2
s: 66
d: [2,125,1,13,2,0,44,55,60]
d: [2,138,1,9,1,0,21,125,59]
s: 138
i: [1,67,0,12,5,0,0,2] fc22_relat_2
s: 67
d: [2,128,1,13,2,0,44,60,55]
d: [2,139,1,9,1,0,21,128,59]
s: 139
d: [2,119,1,10,4,0,44,64,61]
s: 119
d: [2,129,1,16,5,0,44,65,55]
d: [2,137,1,14,4,0,18,129]
d: [2,140,1,10,3,0,21,137,59]
s: 140
d: [2,126,1,11,2,0,44,55,61]
s: 126
d: [2,127,1,11,2,0,44,55,61]
s: 127
d: [2,132,1,11,2,0,44,61,55]
s: 132
d: [2,133,1,11,2,0,44,61,55]
s: 133
d: [2,143,1,24,7,0,40,54,56]
d: [2,144,1,20,6,0,21,143,62]
s: 144
d: [2,165,2,19,3,0,44,126,55]
d: [2,181,2,15,2,0,21,165,59]
d: [2,182,2,11,1,0,21,181,59]
s: 182
d: [2,175,2,19,3,0,44,55,126]
d: [2,183,2,15,2,0,21,175,59]
d: [2,184,2,11,1,0,21,183,59]
s: 184
d: [2,185,2,20,7,0,40,144,141]
d: [2,187,2,18,6,0,18,185]
s: 187
d: [2,193,3,10,2,0,44,62,182]
d: [2,205,3,6,1,0,21,193,59]
s: 205
d: [2,199,3,10,2,0,44,62,182]
d: [2,211,3,6,1,0,21,199,59]
s: 211
d: [2,189,3,14,3,0,44,141,182]
d: [2,201,3,10,2,0,21,189,59]
s: 201
d: [2,190,3,14,3,0,44,58,182]
d: [2,202,3,10,2,0,21,190,59]
s: 202
d: [2,191,3,14,3,0,44,57,182]
d: [2,203,3,10,2,0,21,191,59]
s: 203
d: [2,192,3,14,3,0,44,56,182]
d: [2,204,3,10,2,0,21,192,59]
s: 204
d: [2,195,3,14,3,0,44,141,182]
d: [2,207,3,10,2,0,21,195,59]
s: 207
d: [2,196,3,14,3,0,44,58,182]
d: [2,208,3,10,2,0,21,196,59]
s: 208
d: [2,197,3,14,3,0,44,57,182]
d: [2,209,3,10,2,0,21,197,59]
s: 209
d: [2,198,3,14,3,0,44,56,182]
d: [2,210,3,10,2,0,21,198,59]
s: 210
d: [2,224,4,10,2,0,44,201,182]
s: 224
d: [2,212,3,18,7,0,40,187,58]
d: [2,216,3,16,6,0,18,212]
s: 216
d: [2,225,4,10,2,0,44,201,182]
s: 225
d: [2,214,3,30,6,0,44,187,182]
d: [2,217,3,26,5,0,21,214,59]
s: 217
d: [2,215,3,30,6,0,44,187,182]
d: [2,218,3,26,5,0,21,215,59]
s: 218
d: [2,229,4,10,2,0,44,202,182]
s: 229
d: [2,188,3,32,6,0,44,144,182]
d: [2,200,3,28,5,0,21,188,59]
s: 200
d: [2,230,4,10,2,0,44,202,182]
s: 230
d: [2,194,3,32,6,0,44,144,182]
d: [2,206,3,28,5,0,21,194,59]
s: 206
d: [2,233,4,10,2,0,44,203,182]
s: 233
d: [2,234,4,10,2,0,44,203,182]
s: 234
d: [2,237,4,10,2,0,44,204,182]
s: 237
d: [2,238,4,10,2,0,44,204,182]
s: 238
d: [2,244,4,16,7,0,40,216,57]
d: [2,249,4,14,6,0,18,244]
s: 249
d: [2,281,5,0,0,6,52,88,95,102,81,109,71,249]
e: 281
d: [2,282,5,10,5,1,40,249,71]
d: [2,291,5,8,4,2,21,282,109]
d: [2,292,5,6,3,3,21,291,81]
d: [2,293,5,4,2,4,21,292,102]
d: [2,294,5,2,1,5,21,293,95]
d: [2,295,5,0,0,6,21,294,88]
e: 295
f: 295,105,98,91,84,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t27_orders_1
% SZS output start Proof for orders_1__t27_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v3_orders_1(X0) => v3_orders_1(k2_wellord1(X0,X1)))) [input t27_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (v3_orders_1(X0) => v3_orders_1(k2_wellord1(X0,X1)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d5_orders_1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : k2_wellord1(X0,X1) = k3_xboole_0(X0,k2_zfmisc_1(X1,X1))) [input d6_wellord1]
6. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k2_wellord1(X0,X1))) [input dt_k2_wellord1]
9. ! [X0,X1] : ((v1_relat_1(X1) & v4_relat_2(X0) & v1_relat_1(X0)) => v4_relat_2(k3_xboole_0(X0,X1))) [input fc18_relat_2]
14. ! [X0,X1] : v1_relat_1(k2_zfmisc_1(X0,X1)) [input fc6_relat_1]
16. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v1_relat_2(X1) => v1_relat_2(k2_wellord1(X1,X0)))) [input t15_wellord1]
17. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v6_relat_2(X1) => v6_relat_2(k2_wellord1(X1,X0)))) [input t16_wellord1]
18. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v8_relat_2(X1) => v8_relat_2(k2_wellord1(X1,X0)))) [input t17_wellord1]
19. ! [X0,X1] : (v1_relat_1(X1) => (v6_relat_2(X1) => v6_relat_2(k2_wellord1(X1,X0)))) [flattening 17]
20. ! [X0,X1] : (v1_relat_1(X1) => (v1_relat_2(X1) => v1_relat_2(k2_wellord1(X1,X0)))) [flattening 16]
21. ! [X0,X1] : (v1_relat_1(X1) => (v8_relat_2(X1) => v8_relat_2(k2_wellord1(X1,X0)))) [flattening 18]
23. ? [X0] : (? [X1] : (~v3_orders_1(k2_wellord1(X0,X1)) & v3_orders_1(X0)) & v1_relat_1(X0)) [ennf transformation 2]
24. ! [X0] : ((v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
25. ! [X0] : (! [X1] : k2_wellord1(X0,X1) = k3_xboole_0(X0,k2_zfmisc_1(X1,X1)) | ~v1_relat_1(X0)) [ennf transformation 5]
26. ! [X0,X1] : ((v6_relat_2(k2_wellord1(X1,X0)) | ~v6_relat_2(X1)) | ~v1_relat_1(X1)) [ennf transformation 19]
27. ! [X0,X1] : (v6_relat_2(k2_wellord1(X1,X0)) | ~v6_relat_2(X1) | ~v1_relat_1(X1)) [flattening 26]
28. ! [X0,X1] : ((v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1)) | ~v1_relat_1(X1)) [ennf transformation 20]
29. ! [X0,X1] : (v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1) | ~v1_relat_1(X1)) [flattening 28]
30. ! [X0,X1] : ((v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1)) | ~v1_relat_1(X1)) [ennf transformation 21]
31. ! [X0,X1] : (v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1) | ~v1_relat_1(X1)) [flattening 30]
32. ! [X0,X1] : (v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 6]
34. ! [X0,X1] : (v4_relat_2(k3_xboole_0(X0,X1)) | (~v1_relat_1(X1) | ~v4_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 9]
35. ! [X0,X1] : (v4_relat_2(k3_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v4_relat_2(X0) | ~v1_relat_1(X0)) [flattening 34]
42. ? [X0] : (? [X1] : (~v3_orders_1(k2_wellord1(X0,X1)) & v3_orders_1(X0)) & v1_relat_1(X0)) => (? [X1] : (~v3_orders_1(k2_wellord1(sK0,X1)) & v3_orders_1(sK0)) & v1_relat_1(sK0)) [choice axiom]
43. ? [X1] : (~v3_orders_1(k2_wellord1(sK0,X1)) & v3_orders_1(sK0)) => (~v3_orders_1(k2_wellord1(sK0,sK1)) & v3_orders_1(sK0)) [choice axiom]
44. (~v3_orders_1(k2_wellord1(sK0,sK1)) & v3_orders_1(sK0)) & v1_relat_1(sK0) [skolemisation 23,43,42]
45. ! [X0] : (((v3_orders_1(X0) | (~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 24]
46. ! [X0] : (((v3_orders_1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 45]
47. v1_relat_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. v3_orders_1(sK0) [cnf transformation 44] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. ~v3_orders_1(k2_wellord1(sK0,sK1)) [cnf transformation 44] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. ~v3_orders_1(X0) | v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. ~v3_orders_1(X0) | v8_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. ~v3_orders_1(X0) | v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
53. ~v3_orders_1(X0) | v6_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. ~v6_relat_2(X0) | v3_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. k2_wellord1(X0,X1) = k3_xboole_0(X0,k2_zfmisc_1(X1,X1)) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. v6_relat_2(k2_wellord1(X1,X0)) | ~v6_relat_2(X1) | ~v1_relat_1(X1) [cnf transformation 27] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
57. v1_relat_2(k2_wellord1(X1,X0)) | ~v1_relat_2(X1) | ~v1_relat_1(X1) [cnf transformation 29] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. v8_relat_2(k2_wellord1(X1,X0)) | ~v8_relat_2(X1) | ~v1_relat_1(X1) [cnf transformation 31] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
59. v1_relat_1(k2_zfmisc_1(X0,X1)) [cnf transformation 14] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
62. v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 32] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
64. v4_relat_2(k3_xboole_0(X0,X1)) | ~v1_relat_1(X1) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 35] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
69. 1 <=> v3_orders_1(k2_wellord1(sK0,sK1)) [avatar definition]
71. ~v3_orders_1(k2_wellord1(sK0,sK1)) <- (~1) [avatar component clause 69] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. ~1 [avatar split clause 49,69]
74. 2 <=> v3_orders_1(sK0) [avatar definition]
76. v3_orders_1(sK0) <- (2) [avatar component clause 74] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 2 [avatar split clause 48,74]
79. 3 <=> v1_relat_1(sK0) [avatar definition]
81. v1_relat_1(sK0) <- (3) [avatar component clause 79] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
82. 3 [avatar split clause 47,79]
83. v1_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 50,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
84. v1_relat_2(sK0) <- (2, 3) [subsumption resolution 83,81] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
86. 4 <=> v1_relat_2(sK0) [avatar definition]
88. v1_relat_2(sK0) <- (4) [avatar component clause 86] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
89. 4 | ~2 | ~3 [avatar split clause 84,79,74,86]
90. v8_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 51,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
91. v8_relat_2(sK0) <- (2, 3) [subsumption resolution 90,81] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
93. 5 <=> v8_relat_2(sK0) [avatar definition]
95. v8_relat_2(sK0) <- (5) [avatar component clause 93] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
96. 5 | ~2 | ~3 [avatar split clause 91,79,74,93]
97. v4_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 52,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
98. v4_relat_2(sK0) <- (2, 3) [subsumption resolution 97,81] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. 6 <=> v4_relat_2(sK0) [avatar definition]
102. v4_relat_2(sK0) <- (6) [avatar component clause 100] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
103. 6 | ~2 | ~3 [avatar split clause 98,79,74,100]
104. v6_relat_2(sK0) | ~v1_relat_1(sK0) <- (2) [resolution 53,76] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
105. v6_relat_2(sK0) <- (2, 3) [subsumption resolution 104,81] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
107. 7 <=> v6_relat_2(sK0) [avatar definition]
109. v6_relat_2(sK0) <- (7) [avatar component clause 107] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 7 | ~2 | ~3 [avatar split clause 105,79,74,107]
130. v4_relat_2(k2_wellord1(X3,X4)) | ~v1_relat_1(k2_zfmisc_1(X4,X4)) | ~v4_relat_2(X3) | ~v1_relat_1(X3) | ~v1_relat_1(X3) [superposition 64,55] {a:1,w:14,wCS:5600,sine:2,thAx:0,allAx:2,thDist:-2}
136. v4_relat_2(k2_wellord1(X3,X4)) | ~v1_relat_1(k2_zfmisc_1(X4,X4)) | ~v4_relat_2(X3) | ~v1_relat_1(X3) [duplicate literal removal 130] {a:1,w:12,wCS:4800,sine:2,thAx:0,allAx:2,thDist:-2}
141. v4_relat_2(k2_wellord1(X3,X4)) | ~v4_relat_2(X3) | ~v1_relat_1(X3) [subsumption resolution 136,59] {a:1,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
143. v3_orders_1(k2_wellord1(X0,X1)) | ~v4_relat_2(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [resolution 54,56] {a:1,w:24,wCS:9600,sine:1,thAx:0,allAx:2,thDist:-2}
144. ~v4_relat_2(k2_wellord1(X0,X1)) | v3_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) [subsumption resolution 143,62] {a:1,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
185. v3_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v1_relat_1(X0) [resolution 144,141] {a:2,w:20,wCS:8000,sine:1,thAx:0,allAx:4,thDist:-4}
187. ~v8_relat_2(k2_wellord1(X0,X1)) | v3_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) [duplicate literal removal 185] {a:2,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
212. v3_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_1(X0) [resolution 187,58] {a:3,w:18,wCS:7200,sine:1,thAx:0,allAx:5,thDist:-5}
216. ~v1_relat_2(k2_wellord1(X0,X1)) | v3_orders_1(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) [duplicate literal removal 212] {a:3,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
244. v3_orders_1(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [resolution 216,57] {a:4,w:16,wCS:6400,sine:1,thAx:0,allAx:6,thDist:-6}
249. v3_orders_1(k2_wellord1(X0,X1)) | ~v6_relat_2(X0) | ~v1_relat_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) [duplicate literal removal 244] {a:4,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
282. ~v6_relat_2(sK0) | ~v1_relat_1(sK0) | ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1) [resolution 249,71] {a:5,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
291. ~v1_relat_1(sK0) | ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 7) [subsumption resolution 282,109] {a:5,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
292. ~v4_relat_2(sK0) | ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3, 7) [subsumption resolution 291,81] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
293. ~v8_relat_2(sK0) | ~v1_relat_2(sK0) <- (~1, 3, 6, 7) [subsumption resolution 292,102] {a:5,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
294. ~v1_relat_2(sK0) <- (~1, 3, 5, 6, 7) [subsumption resolution 293,95] {a:5,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
295. $false <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 294,88] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
296. 1 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 295]
299. $false [avatar sat refutation 72,77,82,89,96,103,110,296] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:77,thDist:-77}
% SZS output end Proof for orders_1__t27_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 21
% Generated clauses: 222
% Active clauses: 46
% Passive clauses: 84
% Final active clauses: 46
% Final passive clauses: 19
% Main loop iterations started: 65
% 
% >>> Simplifying Inferences
% Duplicate literals: 19
% Fw subsumption resolutions: 44
% 
% >>> Deletion Inferences
% Simple tautologies: 19
% Forward subsumptions: 66
% 
% >>> Generating Inferences
% Binary resolution: 28
% Unit resulting resolution: 1
% Forward superposition: 68
% Backward superposition: 36
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 11
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

