i: [1,66,0,9,2,1,0,0] t136_zfmisc_1
a: [3,110,0,3,1,66]
s: 110
a: [3,107,0,6,1,66]
s: 107
d: [2,159,0,6,1,2,22,107,110]
a: [3,163,0,6,1,159]
s: 163
i: [1,67,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,122,0,2,1,67]
s: 122
i: [1,68,0,2,1,0,0,255] fc1_xboole_0
a: [3,127,0,2,1,68]
s: 127
i: [1,93,0,2,1,0,0,2] rc2_zfmisc_1
a: [3,137,0,2,1,93]
s: 137
i: [1,94,0,2,1,0,0,2] rc2_xboole_0
a: [3,142,0,2,1,94]
s: 142
i: [1,95,0,2,1,0,0,2] rc1_xboole_0
a: [3,147,0,2,1,95]
s: 147
i: [1,97,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,152,0,2,1,97]
s: 152
i: [1,96,0,2,1,0,0,2] rc1_zfmisc_1
a: [3,157,0,2,1,96]
s: 157
i: [1,70,0,3,1,0,0,1] fc2_xboole_0
s: 70
i: [1,71,0,3,1,0,0,1] fc2_zfmisc_1
s: 71
i: [1,69,0,3,1,0,0,255] d2_xboole_0
a: [3,132,0,3,1,69]
s: 132
r: 122
i: [1,73,0,4,2,0,0,2] cc2_zfmisc_1
s: 73
i: [1,74,0,4,2,0,0,2] cc1_zfmisc_1
s: 74
i: [1,99,0,4,1,0,0,1] d1_tarski
s: 99
d: [2,166,1,2,1,1,40,73,147]
a: [3,171,1,2,1,166]
s: 171
i: [1,72,0,5,1,0,0,255] t1_boole
s: 72
d: [2,167,1,2,1,1,40,73,127]
a: [3,176,1,2,1,167]
s: 176
i: [1,75,0,5,2,0,0,255] t6_boole
s: 75
i: [1,81,0,5,2,0,0,1] t7_boole
s: 81
i: [1,88,0,5,1,0,0,1] idempotence_k2_xboole_0
s: 88
i: [1,90,0,6,2,0,0,1] fc4_xboole_0
s: 90
i: [1,91,0,6,2,0,0,1] fc5_xboole_0
s: 91
i: [1,92,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 92
i: [1,80,0,7,3,0,0,1] t8_boole
s: 80
d: [2,187,1,4,1,0,40,92,99]
s: 187
i: [1,89,0,7,1,0,0,1] commutativity_k2_xboole_0
s: 89
d: [2,188,1,5,1,0,44,89,72]
s: 188
i: [1,100,0,7,2,0,0,1] d1_tarski
s: 100
i: [1,101,0,8,2,0,0,1] d3_xboole_0
s: 101
d: [2,204,1,0,0,1,52,99,163,101]
e: 204
d: [2,205,1,4,1,1,40,101,163]
d: [2,215,1,0,0,1,21,205,99]
e: 215
r: 110
r: 163
a: [3,111,0,3,1,66]
s: 111
i: [1,65,0,9,2,1,0,0] t136_zfmisc_1
a: [3,115,0,3,1,65]
s: 115
d: [2,218,1,2,1,1,40,115,81]
a: [3,227,1,2,1,218]
s: 227
s: 107
a: [3,162,0,6,1,159]
s: 162
d: [2,217,1,3,1,1,40,115,92]
a: [3,222,1,3,1,217]
s: 222
i: [1,102,0,8,2,0,0,1] d3_xboole_0
s: 102
d: [2,247,1,0,0,2,52,115,107,102]
e: 247
d: [2,251,1,3,1,1,40,102,107]
d: [2,258,1,0,0,2,21,251,115]
e: 258
r: 107
r: 115
a: [3,116,0,3,1,65]
s: 116
a: [3,106,0,6,1,66]
s: 106
i: [1,103,0,11,3,0,0,1] d3_xboole_0
s: 103
d: [2,229,1,4,1,1,40,107,101]
a: [3,232,1,4,1,229]
s: 232
d: [2,284,2,0,0,2,52,111,232,100]
e: 284
d: [2,285,2,3,1,1,40,232,100]
d: [2,289,2,0,0,2,21,285,111]
e: 289
i: [1,64,0,12,3,1,0,0] t136_zfmisc_1
d: [2,275,1,7,2,1,40,103,106]
d: [2,282,1,4,1,2,21,275,116]
f: 289,282,258,229,215,159,64,65,66
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t136_zfmisc_1
% SZS output start Proof for zfmisc_1__t136_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))) <=> (X0 = X1 | r2_hidden(X0,X2))) [input t136_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))) <=> (X0 = X1 | r2_hidden(X0,X2))) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
9. ! [X0] : ! [X1] : ! [X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [input d3_xboole_0]
28. ~! [X0,X1,X2] : (r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))) <=> (X0 = X1 | r2_hidden(X0,X2))) [flattening 2]
29. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
32. ! [X0,X1,X2] : (k2_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) | r2_hidden(X3,X0)))) [flattening 9]
34. ? [X0,X1,X2] : (r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))) <~> (X0 = X1 | r2_hidden(X0,X2))) [ennf transformation 28]
43. ? [X0,X1,X2] : (((X0 != X1 & ~r2_hidden(X0,X2)) | ~r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1)))) & ((X0 = X1 | r2_hidden(X0,X2)) | r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))))) [nnf transformation 34]
44. ? [X0,X1,X2] : (((X0 != X1 & ~r2_hidden(X0,X2)) | ~r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1)))) & (X0 = X1 | r2_hidden(X0,X2) | r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))))) [flattening 43]
45. ? [X0,X1,X2] : (((X0 != X1 & ~r2_hidden(X0,X2)) | ~r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1)))) & (X0 = X1 | r2_hidden(X0,X2) | r2_hidden(X0,k2_xboole_0(X2,k1_tarski(X1))))) => (((sK0 != sK1 & ~r2_hidden(sK0,sK2)) | ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1)))) & (sK0 = sK1 | r2_hidden(sK0,sK2) | r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))))) [choice axiom]
46. ((sK0 != sK1 & ~r2_hidden(sK0,sK2)) | ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1)))) & (sK0 = sK1 | r2_hidden(sK0,sK2) | r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1)))) [skolemisation 44,45]
47. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 29]
48. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 47]
49. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) [choice axiom]
50. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK3(X0,X1) != X0 | ~r2_hidden(sK3(X0,X1),X1)) & (sK3(X0,X1) = X0 | r2_hidden(sK3(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 48,49]
51. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) | r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [nnf transformation 32]
52. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) & ~r2_hidden(X3,X0))) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | ~r2_hidden(X3,X2))) | k2_xboole_0(X0,X1) != X2)) [flattening 51]
53. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [rectify 52]
54. ! [X2,X1,X0] : (? [X3] : (((~r2_hidden(X3,X1) & ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & (r2_hidden(X3,X1) | r2_hidden(X3,X0) | r2_hidden(X3,X2))) => (((~r2_hidden(sK4(X0,X1,X2),X1) & ~r2_hidden(sK4(X0,X1,X2),X0)) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(sK4(X0,X1,X2),X1) | r2_hidden(sK4(X0,X1,X2),X0) | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
55. ! [X0,X1,X2] : ((k2_xboole_0(X0,X1) = X2 | (((~r2_hidden(sK4(X0,X1,X2),X1) & ~r2_hidden(sK4(X0,X1,X2),X0)) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (r2_hidden(sK4(X0,X1,X2),X1) | r2_hidden(sK4(X0,X1,X2),X0) | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (~r2_hidden(X4,X1) & ~r2_hidden(X4,X0))) & (r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2))) | k2_xboole_0(X0,X1) != X2)) [skolemisation 53,54]
64. sK0 = sK1 | r2_hidden(sK0,sK2) | r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) [cnf transformation 46] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
65. ~r2_hidden(sK0,sK2) | ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) [cnf transformation 46] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
66. sK0 != sK1 | ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) [cnf transformation 46] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. X0 = X3 | ~r2_hidden(X3,X1) | k1_tarski(X0) != X1 [cnf transformation 50] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
77. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 50] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
82. r2_hidden(X4,X1) | r2_hidden(X4,X0) | ~r2_hidden(X4,X2) | k2_xboole_0(X0,X1) != X2 [cnf transformation 55] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
83. r2_hidden(X4,X2) | ~r2_hidden(X4,X0) | k2_xboole_0(X0,X1) != X2 [cnf transformation 55] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
84. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | k2_xboole_0(X0,X1) != X2 [cnf transformation 55] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
98. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 77] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
99. r2_hidden(X3,k1_tarski(X3)) [equality resolution 98] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
100. ~r2_hidden(X3,k1_tarski(X0)) | X0 = X3 [equality resolution 76] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
101. r2_hidden(X4,k2_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) [equality resolution 84] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
102. r2_hidden(X4,k2_xboole_0(X0,X1)) | ~r2_hidden(X4,X0) [equality resolution 83] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
103. ~r2_hidden(X4,k2_xboole_0(X0,X1)) | r2_hidden(X4,X0) | r2_hidden(X4,X1) [equality resolution 82] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
105. 1 <=> r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) [avatar definition]
106. r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) <- (1) [avatar component clause 105] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
107. ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK1))) <- (~1) [avatar component clause 105] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
109. 2 <=> sK0 = sK1 [avatar definition]
110. sK0 = sK1 <- (2) [avatar component clause 109] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
111. sK0 != sK1 <- (~2) [avatar component clause 109] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. ~1 | ~2 [avatar split clause 66,109,105]
114. 3 <=> r2_hidden(sK0,sK2) [avatar definition]
115. r2_hidden(sK0,sK2) <- (3) [avatar component clause 114] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
116. ~r2_hidden(sK0,sK2) <- (~3) [avatar component clause 114] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. ~1 | ~3 [avatar split clause 65,114,105]
118. 1 | 3 | 2 [avatar split clause 64,109,114,105]
159. ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK0))) <- (~1, 2) [forward demodulation 107,110] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
161. 12 <=> r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK0))) [avatar definition]
163. ~r2_hidden(sK0,k2_xboole_0(sK2,k1_tarski(sK0))) <- (~12) [avatar component clause 161] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
164. ~12 | 1 | ~2 [avatar split clause 159,109,105,161]
205. ~r2_hidden(sK0,k1_tarski(sK0)) <- (~12) [resolution 101,163] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
215. $false <- (~12) [subsumption resolution 205,99] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
216. 12 [avatar contradiction clause 215]
229. ~r2_hidden(sK0,k1_tarski(sK1)) <- (~1) [resolution 107,101] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
231. 17 <=> r2_hidden(sK0,k1_tarski(sK1)) [avatar definition]
232. r2_hidden(sK0,k1_tarski(sK1)) <- (17) [avatar component clause 231] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
234. ~17 | 1 [avatar split clause 229,105,231]
251. ~r2_hidden(sK0,sK2) <- (~1) [resolution 102,107] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
258. $false <- (~1, 3) [subsumption resolution 251,115] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
259. 1 | ~3 [avatar contradiction clause 258]
275. r2_hidden(sK0,sK2) | r2_hidden(sK0,k1_tarski(sK1)) <- (1) [resolution 103,106] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
282. r2_hidden(sK0,k1_tarski(sK1)) <- (1, ~3) [subsumption resolution 275,116] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
283. 17 | ~1 | 3 [avatar split clause 282,114,105,231]
285. sK0 = sK1 <- (17) [resolution 232,100] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
289. $false <- (~2, 17) [subsumption resolution 285,111] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
290. 2 | ~17 [avatar contradiction clause 289]
291. $false [avatar sat refutation 112,117,118,164,216,234,259,283,290] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:125,thDist:-125}
% SZS output end Proof for zfmisc_1__t136_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 34
% Generated clauses: 146
% Active clauses: 40
% Passive clauses: 68
% Final active clauses: 35
% Final passive clauses: 25
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 4
% Fw demodulations: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 14
% Equational tautologies: 1
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 3
% Forward superposition: 24
% Backward superposition: 13
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 30
% SAT solver unit clauses: 10
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

