i: [1,205,0,4,1,1,0,0] t6_borsuk_5
a: [3,335,0,4,1,205]
s: 335
i: [1,204,0,4,1,1,0,0] t6_borsuk_5
a: [3,340,0,4,1,204]
s: 340
i: [1,207,0,7,1,1,0,0] t6_borsuk_5
a: [3,325,0,7,1,207]
s: 325
i: [1,208,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,345,0,2,1,208]
s: 345
i: [1,209,0,2,1,0,0,255] fc1_xboole_0
a: [3,350,0,2,1,209]
s: 350
i: [1,278,0,2,1,0,0,2] rc2_xboole_0
a: [3,367,0,2,1,278]
s: 367
i: [1,279,0,2,1,0,0,255] rc2_rat_1
a: [3,372,0,2,1,279]
s: 372
i: [1,280,0,2,1,0,0,255] rc2_int_1
a: [3,377,0,2,1,280]
s: 377
i: [1,281,0,2,1,0,0,255] rc1_xreal_0
a: [3,382,0,2,1,281]
s: 382
i: [1,282,0,2,1,0,0,2] rc1_xboole_0
a: [3,387,0,2,1,282]
s: 387
i: [1,283,0,2,1,0,0,255] rc1_xxreal_0
a: [3,392,0,2,1,283]
s: 392
i: [1,285,0,2,1,0,0,255] rc1_membered
a: [3,397,0,2,1,285]
s: 397
i: [1,284,0,2,1,0,0,255] rc1_membered
a: [3,402,0,2,1,284]
s: 402
i: [1,287,0,2,1,0,0,255] rc2_membered
a: [3,407,0,2,1,287]
s: 407
i: [1,286,0,2,1,0,0,255] rc2_membered
a: [3,412,0,2,1,286]
s: 412
i: [1,289,0,2,1,0,0,255] rc3_membered
a: [3,417,0,2,1,289]
s: 417
i: [1,288,0,2,1,0,0,255] rc3_membered
a: [3,422,0,2,1,288]
s: 422
i: [1,290,0,2,1,0,0,255] rc1_finset_1
a: [3,427,0,2,1,290]
s: 427
i: [1,291,0,2,1,0,0,255] rc8_finset_1
a: [3,432,0,2,1,291]
s: 432
i: [1,295,0,2,1,0,0,255] rc2_xreal_0
a: [3,437,0,2,1,295]
s: 437
i: [1,294,0,2,1,0,0,255] rc2_xreal_0
a: [3,442,0,2,1,294]
s: 442
i: [1,293,0,2,1,0,0,255] rc2_xreal_0
a: [3,447,0,2,1,293]
s: 447
i: [1,292,0,2,1,0,0,255] rc2_xreal_0
a: [3,452,0,2,1,292]
s: 452
i: [1,299,0,2,1,0,0,255] rc3_xreal_0
a: [3,457,0,2,1,299]
s: 457
i: [1,298,0,2,1,0,0,255] rc3_xreal_0
a: [3,462,0,2,1,298]
s: 462
i: [1,297,0,2,1,0,0,255] rc3_xreal_0
a: [3,467,0,2,1,297]
s: 467
i: [1,296,0,2,1,0,0,255] rc3_xreal_0
a: [3,472,0,2,1,296]
s: 472
i: [1,303,0,2,1,0,0,255] rc4_xreal_0
a: [3,477,0,2,1,303]
s: 477
i: [1,302,0,2,1,0,0,255] rc4_xreal_0
a: [3,482,0,2,1,302]
s: 482
i: [1,301,0,2,1,0,0,255] rc4_xreal_0
a: [3,487,0,2,1,301]
s: 487
i: [1,300,0,2,1,0,0,255] rc4_xreal_0
a: [3,492,0,2,1,300]
s: 492
i: [1,305,0,2,1,0,0,255] rc4_xxreal_0
a: [3,497,0,2,1,305]
s: 497
i: [1,304,0,2,1,0,0,255] rc4_xxreal_0
a: [3,502,0,2,1,304]
s: 502
i: [1,307,0,2,1,0,0,255] rc2_xxreal_0
a: [3,507,0,2,1,307]
s: 507
i: [1,306,0,2,1,0,0,255] rc2_xxreal_0
a: [3,512,0,2,1,306]
s: 512
i: [1,309,0,2,1,0,0,255] rc3_xxreal_0
a: [3,517,0,2,1,309]
s: 517
i: [1,308,0,2,1,0,0,255] rc3_xxreal_0
a: [3,522,0,2,1,308]
s: 522
i: [1,206,0,9,1,1,0,0] t6_borsuk_5
a: [3,330,0,9,1,206]
s: 330
i: [1,210,0,3,1,0,0,255] d2_xboole_0
a: [3,355,0,3,1,210]
s: 355
r: 345
i: [1,211,0,4,2,0,0,255] cc2_int_1
s: 211
i: [1,212,0,4,2,0,0,255] cc2_xreal_0
s: 212
i: [1,213,0,4,2,0,0,255] cc2_xxreal_0
s: 213
i: [1,216,0,4,2,0,0,255] cc2_membered
s: 216
i: [1,217,0,4,2,0,0,255] cc3_membered
s: 217
i: [1,220,0,4,2,0,0,255] cc4_membered
s: 220
i: [1,221,0,4,2,0,0,255] cc5_membered
s: 221
i: [1,223,0,4,2,0,0,255] cc1_rat_1
s: 223
i: [1,224,0,4,2,0,0,255] cc2_rat_1
s: 224
d: [2,525,1,2,1,1,40,223,372]
a: [3,529,1,2,1,525]
s: 529
i: [1,225,0,4,2,0,0,255] cc3_int_1
s: 225
i: [1,226,0,4,2,0,0,255] cc1_membered
s: 226
d: [2,531,1,2,1,1,40,224,377]
a: [3,535,1,2,1,531]
s: 535
i: [1,228,0,4,2,0,0,255] cc3_xreal_0
s: 228
d: [2,537,1,2,1,1,40,225,377]
a: [3,541,1,2,1,537]
s: 541
i: [1,229,0,4,2,0,0,255] cc4_xreal_0
s: 229
i: [1,230,0,4,2,0,0,255] cc18_membered
s: 230
d: [2,543,1,2,1,1,40,226,397]
a: [3,549,1,2,1,543]
s: 549
i: [1,264,0,4,1,0,0,255] existence_m1_subset_1
s: 264
d: [2,544,1,2,1,1,40,226,407]
a: [3,554,1,2,1,544]
s: 554
i: [1,310,0,4,1,0,0,1] d5_zfmisc_1
s: 310
i: [1,311,0,4,1,0,0,1] d5_zfmisc_1
s: 311
d: [2,545,1,2,1,1,40,226,417]
a: [3,559,1,2,1,545]
s: 559
i: [1,312,0,4,1,0,0,1] d5_zfmisc_1
s: 312
d: [2,564,1,2,1,1,40,228,382]
a: [3,576,1,2,1,564]
s: 576
i: [1,231,0,5,2,0,0,255] t6_boole
s: 231
i: [1,250,0,5,2,0,0,255] t7_boole
s: 250
d: [2,586,1,2,1,1,40,229,382]
a: [3,603,1,2,1,586]
s: 603
i: [1,271,0,5,1,0,0,1] fc2_subset_1
s: 271
d: [2,605,1,2,1,1,40,230,350]
a: [3,612,1,2,1,605]
s: 612
i: [1,267,0,7,3,0,0,1] irreflexivity_r1_subset_1
a: [3,362,0,5,2,267]
s: 362
i: [1,237,0,6,3,0,0,255] cc7_xxreal_0
s: 237
d: [2,606,1,2,1,1,40,230,387]
a: [3,617,1,2,1,606]
s: 617
i: [1,238,0,6,3,0,0,255] cc7_xxreal_0
s: 238
d: [2,607,1,2,1,1,40,230,492]
a: [3,622,1,2,1,607]
s: 622
i: [1,241,0,6,3,0,0,255] cc5_xxreal_0
s: 241
i: [1,243,0,6,3,0,0,255] cc5_xxreal_0
s: 243
d: [2,608,1,2,1,1,40,230,502]
a: [3,627,1,2,1,608]
s: 627
i: [1,244,0,6,3,0,0,255] cc3_xxreal_0
s: 244
d: [2,656,1,4,2,1,40,237,442]
d: [2,658,1,2,1,2,21,656,447]
a: [3,662,1,2,1,658]
s: 662
i: [1,246,0,6,3,0,0,255] cc3_xxreal_0
s: 246
i: [1,247,0,6,2,0,0,255] t1_subset
s: 247
d: [2,657,1,4,2,1,40,237,507]
d: [2,664,1,2,1,2,21,657,512]
a: [3,668,1,2,1,664]
s: 668
i: [1,265,0,6,2,0,0,255] symmetry_r1_xboole_0
s: 265
d: [2,676,1,4,2,1,40,238,462]
d: [2,678,1,2,1,2,21,676,467]
a: [3,682,1,2,1,678]
s: 682
i: [1,266,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 266
i: [1,214,0,7,3,0,0,255] cc12_membered
s: 214
d: [2,677,1,4,2,1,40,238,517]
d: [2,684,1,2,1,2,21,677,522]
a: [3,688,1,2,1,684]
s: 688
i: [1,215,0,7,3,0,0,255] cc13_membered
s: 215
d: [2,696,1,4,2,1,40,243,462]
d: [2,698,1,2,1,2,21,696,467]
a: [3,702,1,2,1,698]
s: 702
i: [1,218,0,7,3,0,0,255] cc15_membered
s: 218
i: [1,219,0,7,3,0,0,255] cc16_membered
s: 219
d: [2,697,1,4,2,1,40,243,517]
d: [2,704,1,2,1,2,21,697,522]
a: [3,708,1,2,1,704]
s: 708
i: [1,222,0,7,3,0,0,255] cc14_membered
s: 222
d: [2,563,2,2,1,1,40,228,529]
a: [3,571,2,2,1,563]
s: 571
i: [1,227,0,7,3,0,0,255] cc17_membered
s: 227
i: [1,249,0,7,3,0,0,2] t8_boole
s: 249
d: [2,578,2,2,1,1,40,541,228]
a: [3,582,2,2,1,578]
s: 582
i: [1,233,0,8,4,0,0,255] cc6_xxreal_0
s: 233
d: [2,723,1,6,3,1,40,233,392]
a: [3,746,1,2,1,723]
s: 746
r: 746
i: [1,235,0,8,4,0,0,255] cc4_xxreal_0
s: 235
i: [1,239,0,8,4,0,0,255] cc8_xxreal_0
s: 239
a: [3,745,1,2,1,723]
s: 745
i: [1,248,0,8,3,0,0,255] t2_subset
s: 248
a: [3,750,1,2,1,723]
s: 750
i: [1,313,0,9,1,0,0,1] t5_borsuk_5
s: 313
i: [1,314,0,9,1,0,0,1] t5_borsuk_5
s: 314
a: [3,753,1,2,1,723]
s: 753
i: [1,315,0,9,1,0,0,1] t5_borsuk_5
s: 315
d: [2,584,2,2,1,1,40,229,529]
a: [3,593,2,2,1,584]
s: 593
i: [1,316,0,9,1,0,0,1] t5_borsuk_5
s: 316
i: [1,317,0,9,1,0,0,1] t5_borsuk_5
s: 317
d: [2,585,2,2,1,1,40,229,541]
a: [3,598,2,2,1,585]
s: 598
i: [1,318,0,9,1,0,0,1] t5_borsuk_5
s: 318
d: [2,629,2,2,1,1,40,549,216]
a: [3,633,2,2,1,629]
s: 633
i: [1,319,0,9,1,0,0,1] t5_borsuk_5
s: 319
i: [1,320,0,9,1,0,0,1] t5_borsuk_5
s: 320
d: [2,635,2,2,1,1,40,554,216]
a: [3,639,2,2,1,635]
s: 639
i: [1,321,0,9,1,0,0,1] t5_borsuk_5
s: 321
d: [2,641,2,2,1,1,40,559,216]
a: [3,645,2,2,1,641]
s: 645
i: [1,268,0,10,4,0,0,1] symmetry_r1_subset_1
s: 268
i: [1,269,0,10,4,0,0,255] redefinition_r1_subset_1
s: 269
d: [2,649,2,2,1,1,40,612,226]
a: [3,654,2,2,1,649]
s: 654
i: [1,270,0,10,4,0,0,255] redefinition_r1_subset_1
s: 270
d: [2,670,2,2,1,1,40,617,226]
a: [3,674,2,2,1,670]
s: 674
i: [1,272,0,11,4,0,0,255] fc24_membered
s: 272
i: [1,273,0,11,4,0,0,255] fc22_membered
s: 273
d: [2,690,2,2,1,1,40,622,226]
a: [3,694,2,2,1,690]
s: 694
i: [1,274,0,11,4,0,0,255] fc23_membered
s: 274
d: [2,710,2,2,1,1,40,627,226]
a: [3,714,2,2,1,710]
s: 714
i: [1,275,0,11,4,0,0,255] fc19_membered
s: 275
i: [1,276,0,11,4,0,0,255] fc21_membered
s: 276
d: [2,722,2,6,3,1,40,233,603]
a: [3,733,2,2,1,722]
s: 733
i: [1,277,0,11,4,0,0,255] fc20_membered
s: 277
d: [2,767,3,6,3,1,40,593,233]
a: [3,771,3,2,1,767]
s: 771
r: 771
i: [1,263,0,52,16,0,0,1] d8_zfmisc_1
s: 263
d: [2,871,1,45,15,1,40,263,325]
a: [3,875,1,3,1,871]
s: 875
r: 330
r: 325
r: 340
d: [2,935,0,4,1,2,23,340,875]
d: [2,946,0,0,0,2,21,935,312]
e: 946
r: 875
a: [3,874,1,3,1,871]
s: 874
s: 340
a: [3,879,1,3,1,871]
s: 879
r: 340
d: [2,948,0,4,1,2,23,340,879]
d: [2,949,0,0,0,2,21,948,311]
e: 949
r: 879
s: 340
s: 325
a: [3,878,1,3,1,871]
s: 878
s: 330
a: [3,883,1,3,1,871]
s: 883
r: 330
r: 325
r: 335
d: [2,962,0,9,1,2,23,330,883]
d: [2,965,0,0,0,2,21,962,321]
e: 965
r: 883
s: 335
s: 325
a: [3,882,1,3,1,871]
s: 882
s: 330
a: [3,887,1,3,1,871]
s: 887
r: 330
r: 325
r: 335
d: [2,979,0,9,1,2,23,330,887]
d: [2,982,0,0,0,2,21,979,320]
e: 982
r: 887
s: 335
s: 325
a: [3,886,1,3,1,871]
s: 886
s: 330
a: [3,891,1,3,1,871]
s: 891
r: 330
r: 325
r: 335
d: [2,997,0,9,1,2,23,330,891]
d: [2,1000,0,0,0,2,21,997,319]
e: 1000
r: 891
s: 335
s: 325
a: [3,890,1,3,1,871]
s: 890
s: 330
a: [3,895,1,3,1,871]
s: 895
r: 330
r: 325
r: 340
d: [2,1018,0,4,1,2,23,340,895]
d: [2,1030,0,0,0,2,21,1018,310]
e: 1030
r: 895
s: 340
s: 325
a: [3,894,1,3,1,871]
s: 894
s: 330
a: [3,899,1,3,1,871]
s: 899
r: 330
r: 325
r: 335
d: [2,1047,0,9,1,2,23,330,899]
d: [2,1051,0,0,0,2,21,1047,318]
e: 1051
r: 899
s: 335
s: 325
a: [3,898,1,3,1,871]
s: 898
s: 330
a: [3,903,1,3,1,871]
s: 903
r: 330
r: 325
r: 335
d: [2,1069,0,9,1,2,23,330,903]
d: [2,1073,0,0,0,2,21,1069,317]
e: 1073
r: 903
s: 335
s: 325
a: [3,902,1,3,1,871]
s: 902
s: 330
a: [3,907,1,3,1,871]
s: 907
r: 330
r: 325
r: 335
d: [2,1092,0,9,1,2,23,330,907]
d: [2,1096,0,0,0,2,21,1092,316]
e: 1096
r: 907
s: 335
s: 325
a: [3,906,1,3,1,871]
s: 906
s: 330
a: [3,911,1,3,1,871]
s: 911
r: 330
r: 325
r: 335
d: [2,1116,0,9,1,2,23,330,911]
d: [2,1121,0,0,0,2,21,1116,315]
e: 1121
r: 911
s: 335
s: 325
a: [3,910,1,3,1,871]
s: 910
s: 330
a: [3,915,1,3,1,871]
s: 915
r: 330
r: 325
r: 335
d: [2,1142,0,9,1,2,23,330,915]
d: [2,1147,0,0,0,2,21,1142,314]
e: 1147
r: 915
s: 335
s: 325
a: [3,914,1,3,1,871]
s: 914
s: 330
a: [3,919,1,3,1,871]
s: 919
r: 330
r: 325
r: 335
d: [2,1169,0,9,1,2,23,330,919]
d: [2,1174,0,0,0,2,21,1169,313]
e: 1174
r: 919
s: 335
s: 325
a: [3,918,1,3,1,871]
s: 918
s: 330
a: [3,923,1,3,1,871]
s: 923
r: 330
r: 325
r: 335
d: [2,1199,0,4,1,2,23,335,923]
d: [2,1213,0,0,0,2,21,1199,312]
e: 1213
r: 923
s: 335
s: 325
a: [3,922,1,3,1,871]
s: 922
s: 330
a: [3,927,1,3,1,871]
s: 927
r: 330
r: 325
r: 335
d: [2,1239,0,4,1,2,23,335,927]
d: [2,1253,0,0,0,2,21,1239,311]
e: 1253
r: 927
s: 335
s: 325
a: [3,926,1,3,1,871]
s: 926
s: 330
a: [3,931,1,3,1,871]
s: 931
r: 325
r: 330
r: 335
d: [2,1280,0,4,1,2,23,335,931]
d: [2,1295,0,0,0,2,21,1280,310]
e: 1295
d: [2,1255,1,45,15,1,40,325,263]
d: [2,1256,1,42,14,2,21,1255,922]
d: [2,1257,1,39,13,3,21,1256,918]
d: [2,1258,1,36,12,4,21,1257,914]
d: [2,1259,1,33,11,5,21,1258,910]
d: [2,1260,1,30,10,6,21,1259,906]
d: [2,1261,1,27,9,7,21,1260,902]
d: [2,1262,1,24,8,8,21,1261,898]
d: [2,1263,1,21,7,9,21,1262,894]
d: [2,1264,1,18,6,10,21,1263,890]
d: [2,1265,1,15,5,11,21,1264,886]
d: [2,1266,1,12,4,12,21,1265,882]
d: [2,1267,1,9,3,13,21,1266,878]
d: [2,1268,1,6,2,14,21,1267,874]
f: 1295,1268,1253,1213,1174,1147,1121,1096,1073,1051,1030,1000,982,965,949,946,204,205,206,207
% Refutation found. Thanks to Tanya!
% SZS status Theorem for borsuk_5__t6_borsuk_5
% SZS output start Proof for borsuk_5__t6_borsuk_5
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2)) => r4_zfmisc_1(X0,X1,X2,X3,X4,X5)) [input t6_borsuk_5]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ((r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2)) => r4_zfmisc_1(X0,X1,X2,X3,X4,X5)) [negated conjecture 1]
37. ! [X0] : ! [X1] : ! [X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [input d5_zfmisc_1]
38. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (r4_zfmisc_1(X0,X1,X2,X3,X4,X5) <=> (X4 != X5 & X3 != X5 & X3 != X4 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [input d8_zfmisc_1]
76. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) => (X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X0 != X5 & X0 != X4 & X0 != X3)) [input t5_borsuk_5]
80. ~! [X0,X1,X2,X3,X4,X5] : ((r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2)) => r4_zfmisc_1(X0,X1,X2,X3,X4,X5)) [flattening 2]
85. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) <=> (X1 != X2 & X0 != X2 & X0 != X1)) [flattening 37]
86. ! [X0,X1,X2,X3,X4,X5] : (r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) => (X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X0 != X5 & X0 != X4 & X0 != X3)) [flattening 76]
87. ! [X0,X1,X2,X3,X4,X5] : (r4_zfmisc_1(X0,X1,X2,X3,X4,X5) <=> (X4 != X5 & X3 != X5 & X3 != X4 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1)) [flattening 38]
88. ! [X0,X1,X2,X3,X4,X5] : ((X4 != X5 & X3 != X5 & X3 != X4 & X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X1 != X2 & X0 != X5 & X0 != X4 & X0 != X3 & X0 != X2 & X0 != X1) => r4_zfmisc_1(X0,X1,X2,X3,X4,X5)) [unused predicate definition removal 87]
89. ! [X0,X1,X2] : (r1_zfmisc_1(X0,X1,X2) => (X1 != X2 & X0 != X2 & X0 != X1)) [unused predicate definition removal 85]
101. ? [X0,X1,X2,X3,X4,X5] : (~r4_zfmisc_1(X0,X1,X2,X3,X4,X5) & (r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2))) [ennf transformation 80]
102. ? [X0,X1,X2,X3,X4,X5] : (~r4_zfmisc_1(X0,X1,X2,X3,X4,X5) & r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2)) [flattening 101]
141. ! [X0,X1,X2] : ((X1 != X2 & X0 != X2 & X0 != X1) | ~r1_zfmisc_1(X0,X1,X2)) [ennf transformation 89]
142. ! [X0,X1,X2,X3,X4,X5] : ((X2 != X5 & X2 != X4 & X2 != X3 & X1 != X5 & X1 != X4 & X1 != X3 & X0 != X5 & X0 != X4 & X0 != X3) | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5))) [ennf transformation 86]
143. ! [X0,X1,X2,X3,X4,X5] : (r4_zfmisc_1(X0,X1,X2,X3,X4,X5) | (X4 = X5 | X3 = X5 | X3 = X4 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1)) [ennf transformation 88]
144. ! [X0,X1,X2,X3,X4,X5] : (r4_zfmisc_1(X0,X1,X2,X3,X4,X5) | X4 = X5 | X3 = X5 | X3 = X4 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1) [flattening 143]
165. ? [X0,X1,X2,X3,X4,X5] : (~r4_zfmisc_1(X0,X1,X2,X3,X4,X5) & r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) & r1_zfmisc_1(X3,X4,X5) & r1_zfmisc_1(X0,X1,X2)) => (~r4_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5) & r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK5)) & r1_zfmisc_1(sK3,sK4,sK5) & r1_zfmisc_1(sK0,sK1,sK2)) [choice axiom]
166. ~r4_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5) & r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK5)) & r1_zfmisc_1(sK3,sK4,sK5) & r1_zfmisc_1(sK0,sK1,sK2) [skolemisation 102,165]
204. r1_zfmisc_1(sK0,sK1,sK2) [cnf transformation 166] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
205. r1_zfmisc_1(sK3,sK4,sK5) [cnf transformation 166] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
206. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK5)) [cnf transformation 166] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
207. ~r4_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5) [cnf transformation 166] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
251. X0 != X1 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 141] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
252. X0 != X2 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 141] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
253. X1 != X2 | ~r1_zfmisc_1(X0,X1,X2) [cnf transformation 141] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:1,thDist:-1}
254. X0 != X3 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
255. X0 != X4 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
256. X0 != X5 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
257. X1 != X3 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
258. X1 != X4 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
259. X1 != X5 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
260. X2 != X3 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
261. X2 != X4 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
262. X2 != X5 | ~r1_subset_1(k1_enumset1(X0,X1,X2),k1_enumset1(X3,X4,X5)) [cnf transformation 142] {a:0,w:12,wCS:4800,sine:1,thAx:0,allAx:1,thDist:-1}
263. r4_zfmisc_1(X0,X1,X2,X3,X4,X5) | X4 = X5 | X3 = X5 | X3 = X4 | X2 = X5 | X2 = X4 | X2 = X3 | X1 = X5 | X1 = X4 | X1 = X3 | X1 = X2 | X0 = X5 | X0 = X4 | X0 = X3 | X0 = X2 | X0 = X1 [cnf transformation 144] {a:0,w:52,wCS:20800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
310. ~r1_zfmisc_1(X0,X2,X2) [equality resolution 253] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
311. ~r1_zfmisc_1(X2,X1,X2) [equality resolution 252] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
312. ~r1_zfmisc_1(X1,X1,X2) [equality resolution 251] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
313. ~r1_subset_1(k1_enumset1(X0,X1,X5),k1_enumset1(X3,X4,X5)) [equality resolution 262] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
314. ~r1_subset_1(k1_enumset1(X0,X1,X4),k1_enumset1(X3,X4,X5)) [equality resolution 261] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
315. ~r1_subset_1(k1_enumset1(X0,X1,X3),k1_enumset1(X3,X4,X5)) [equality resolution 260] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
316. ~r1_subset_1(k1_enumset1(X0,X5,X2),k1_enumset1(X3,X4,X5)) [equality resolution 259] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
317. ~r1_subset_1(k1_enumset1(X0,X4,X2),k1_enumset1(X3,X4,X5)) [equality resolution 258] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
318. ~r1_subset_1(k1_enumset1(X0,X3,X2),k1_enumset1(X3,X4,X5)) [equality resolution 257] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
319. ~r1_subset_1(k1_enumset1(X5,X1,X2),k1_enumset1(X3,X4,X5)) [equality resolution 256] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
320. ~r1_subset_1(k1_enumset1(X4,X1,X2),k1_enumset1(X3,X4,X5)) [equality resolution 255] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
321. ~r1_subset_1(k1_enumset1(X3,X1,X2),k1_enumset1(X3,X4,X5)) [equality resolution 254] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
323. 1 <=> r4_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5) [avatar definition]
325. ~r4_zfmisc_1(sK0,sK1,sK2,sK3,sK4,sK5) <- (~1) [avatar component clause 323] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
326. ~1 [avatar split clause 207,323]
328. 2 <=> r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK5)) [avatar definition]
330. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK5)) <- (2) [avatar component clause 328] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
331. 2 [avatar split clause 206,328]
333. 3 <=> r1_zfmisc_1(sK3,sK4,sK5) [avatar definition]
335. r1_zfmisc_1(sK3,sK4,sK5) <- (3) [avatar component clause 333] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
336. 3 [avatar split clause 205,333]
338. 4 <=> r1_zfmisc_1(sK0,sK1,sK2) [avatar definition]
340. r1_zfmisc_1(sK0,sK1,sK2) <- (4) [avatar component clause 338] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
341. 4 [avatar split clause 204,338]
873. 95 <=> sK0 = sK1 [avatar definition]
874. sK0 != sK1 <- (~95) [avatar component clause 873] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
875. sK0 = sK1 <- (95) [avatar component clause 873] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
877. 96 <=> sK0 = sK2 [avatar definition]
878. sK0 != sK2 <- (~96) [avatar component clause 877] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
879. sK0 = sK2 <- (96) [avatar component clause 877] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
881. 97 <=> sK0 = sK3 [avatar definition]
882. sK0 != sK3 <- (~97) [avatar component clause 881] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
883. sK0 = sK3 <- (97) [avatar component clause 881] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
885. 98 <=> sK0 = sK4 [avatar definition]
886. sK0 != sK4 <- (~98) [avatar component clause 885] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
887. sK0 = sK4 <- (98) [avatar component clause 885] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
889. 99 <=> sK0 = sK5 [avatar definition]
890. sK0 != sK5 <- (~99) [avatar component clause 889] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
891. sK0 = sK5 <- (99) [avatar component clause 889] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
893. 100 <=> sK1 = sK2 [avatar definition]
894. sK1 != sK2 <- (~100) [avatar component clause 893] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
895. sK1 = sK2 <- (100) [avatar component clause 893] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
897. 101 <=> sK1 = sK3 [avatar definition]
898. sK1 != sK3 <- (~101) [avatar component clause 897] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
899. sK1 = sK3 <- (101) [avatar component clause 897] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
901. 102 <=> sK1 = sK4 [avatar definition]
902. sK1 != sK4 <- (~102) [avatar component clause 901] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
903. sK1 = sK4 <- (102) [avatar component clause 901] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
905. 103 <=> sK1 = sK5 [avatar definition]
906. sK1 != sK5 <- (~103) [avatar component clause 905] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
907. sK1 = sK5 <- (103) [avatar component clause 905] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
909. 104 <=> sK2 = sK3 [avatar definition]
910. sK2 != sK3 <- (~104) [avatar component clause 909] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
911. sK2 = sK3 <- (104) [avatar component clause 909] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
913. 105 <=> sK2 = sK4 [avatar definition]
914. sK2 != sK4 <- (~105) [avatar component clause 913] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
915. sK2 = sK4 <- (105) [avatar component clause 913] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
917. 106 <=> sK2 = sK5 [avatar definition]
918. sK2 != sK5 <- (~106) [avatar component clause 917] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
919. sK2 = sK5 <- (106) [avatar component clause 917] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
921. 107 <=> sK3 = sK4 [avatar definition]
922. sK3 != sK4 <- (~107) [avatar component clause 921] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
923. sK3 = sK4 <- (107) [avatar component clause 921] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
925. 108 <=> sK3 = sK5 [avatar definition]
927. sK3 = sK5 <- (108) [avatar component clause 925] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
929. 109 <=> sK4 = sK5 [avatar definition]
931. sK4 = sK5 <- (109) [avatar component clause 929] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
935. r1_zfmisc_1(sK0,sK0,sK2) <- (4, 95) [backward demodulation 340,875] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
946. $false <- (4, 95) [subsumption resolution 935,312] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
947. ~4 | ~95 [avatar contradiction clause 946]
948. r1_zfmisc_1(sK0,sK1,sK0) <- (4, 96) [backward demodulation 340,879] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
949. $false <- (4, 96) [subsumption resolution 948,311] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
950. ~4 | ~96 [avatar contradiction clause 949]
962. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK0,sK4,sK5)) <- (2, 97) [backward demodulation 330,883] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
965. $false <- (2, 97) [subsumption resolution 962,321] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
966. ~2 | ~97 [avatar contradiction clause 965]
979. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK0,sK5)) <- (2, 98) [backward demodulation 330,887] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
982. $false <- (2, 98) [subsumption resolution 979,320] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
983. ~2 | ~98 [avatar contradiction clause 982]
997. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK0)) <- (2, 99) [backward demodulation 330,891] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1000. $false <- (2, 99) [subsumption resolution 997,319] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1001. ~2 | ~99 [avatar contradiction clause 1000]
1018. r1_zfmisc_1(sK0,sK1,sK1) <- (4, 100) [backward demodulation 340,895] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1030. $false <- (4, 100) [subsumption resolution 1018,310] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1031. ~4 | ~100 [avatar contradiction clause 1030]
1047. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK1,sK4,sK5)) <- (2, 101) [backward demodulation 330,899] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1051. $false <- (2, 101) [subsumption resolution 1047,318] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1052. ~2 | ~101 [avatar contradiction clause 1051]
1069. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK1,sK5)) <- (2, 102) [backward demodulation 330,903] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1073. $false <- (2, 102) [subsumption resolution 1069,317] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1074. ~2 | ~102 [avatar contradiction clause 1073]
1092. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK1)) <- (2, 103) [backward demodulation 330,907] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1096. $false <- (2, 103) [subsumption resolution 1092,316] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1097. ~2 | ~103 [avatar contradiction clause 1096]
1116. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK2,sK4,sK5)) <- (2, 104) [backward demodulation 330,911] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1121. $false <- (2, 104) [subsumption resolution 1116,315] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1122. ~2 | ~104 [avatar contradiction clause 1121]
1142. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK2,sK5)) <- (2, 105) [backward demodulation 330,915] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1147. $false <- (2, 105) [subsumption resolution 1142,314] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1148. ~2 | ~105 [avatar contradiction clause 1147]
1169. r1_subset_1(k1_enumset1(sK0,sK1,sK2),k1_enumset1(sK3,sK4,sK2)) <- (2, 106) [backward demodulation 330,919] {a:0,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1174. $false <- (2, 106) [subsumption resolution 1169,313] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1175. ~2 | ~106 [avatar contradiction clause 1174]
1199. r1_zfmisc_1(sK3,sK3,sK5) <- (3, 107) [backward demodulation 335,923] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1213. $false <- (3, 107) [subsumption resolution 1199,312] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1214. ~3 | ~107 [avatar contradiction clause 1213]
1239. r1_zfmisc_1(sK3,sK4,sK3) <- (3, 108) [backward demodulation 335,927] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1253. $false <- (3, 108) [subsumption resolution 1239,311] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1254. ~3 | ~108 [avatar contradiction clause 1253]
1255. sK4 = sK5 | sK3 = sK5 | sK3 = sK4 | sK2 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1) [resolution 325,263] {a:1,w:45,wCS:4500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1256. sK4 = sK5 | sK3 = sK5 | sK2 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~107) [subsumption resolution 1255,922] {a:1,w:42,wCS:4200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1257. sK4 = sK5 | sK3 = sK5 | sK2 = sK4 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~106, ~107) [subsumption resolution 1256,918] {a:1,w:39,wCS:3900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1258. sK4 = sK5 | sK3 = sK5 | sK2 = sK3 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~105, ~106, ~107) [subsumption resolution 1257,914] {a:1,w:36,wCS:3600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1259. sK4 = sK5 | sK3 = sK5 | sK1 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~104, ~105, ~106, ~107) [subsumption resolution 1258,910] {a:1,w:33,wCS:3300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1260. sK4 = sK5 | sK3 = sK5 | sK1 = sK4 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1259,906] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1261. sK4 = sK5 | sK3 = sK5 | sK1 = sK3 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1260,902] {a:1,w:27,wCS:2700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1262. sK4 = sK5 | sK3 = sK5 | sK1 = sK2 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1261,898] {a:1,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1263. sK4 = sK5 | sK3 = sK5 | sK0 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1262,894] {a:1,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1264. sK4 = sK5 | sK3 = sK5 | sK0 = sK4 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~99, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1263,890] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1265. sK4 = sK5 | sK3 = sK5 | sK0 = sK3 | sK0 = sK2 | sK0 = sK1 <- (~1, ~98, ~99, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1264,886] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1266. sK4 = sK5 | sK3 = sK5 | sK0 = sK2 | sK0 = sK1 <- (~1, ~97, ~98, ~99, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1265,882] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1267. sK4 = sK5 | sK3 = sK5 | sK0 = sK1 <- (~1, ~96, ~97, ~98, ~99, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1266,878] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1268. sK4 = sK5 | sK3 = sK5 <- (~1, ~95, ~96, ~97, ~98, ~99, ~100, ~101, ~102, ~103, ~104, ~105, ~106, ~107) [subsumption resolution 1267,874] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1269. 108 | 109 | 1 | 95 | 96 | 97 | 98 | 99 | 100 | 101 | 102 | 103 | 104 | 105 | 106 | 107 [avatar split clause 1268,921,917,913,909,905,901,897,893,889,885,881,877,873,323,929,925]
1280. r1_zfmisc_1(sK3,sK4,sK4) <- (3, 109) [backward demodulation 335,931] {a:0,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1295. $false <- (3, 109) [subsumption resolution 1280,310] {a:0,w:0,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1296. ~3 | ~109 [avatar contradiction clause 1295]
1297. $false [avatar sat refutation 326,331,336,341,947,950,966,983,1001,1031,1052,1074,1097,1122,1148,1175,1214,1254,1269,1296] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:878,thDist:-878}
% SZS output end Proof for borsuk_5__t6_borsuk_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 78
% >>> Preprocessing
% Introduced skolems: 24
% Pure predicates: 4
% Unused predicate definitions: 2
% 
% >>> Saturation
% Initial clauses: 106
% Generated clauses: 612
% Active clauses: 199
% Passive clauses: 231
% Final active clauses: 139
% Final passive clauses: 27
% Main loop iterations started: 204
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 169
% Bw demodulations: 44
% 
% >>> Deletion Inferences
% Simple tautologies: 7
% Forward subsumptions: 40
% 
% >>> Generating Inferences
% Binary resolution: 112
% Forward superposition: 2
% Backward superposition: 21
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 19
% Split components: 133
% Sat splitting refutations: 15
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 154
% SAT solver unit clauses: 39
% SAT solver binary clauses: 78
% 
% Memory used [KB]: 5884
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% splitting model update: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.002 s
% passive container maintenance: 0.001 s
% other: 0.009 s ( own 0.004 s ) 

