i: [1,125,0,4,1,1,0,0] t56_classes1
a: [3,188,0,4,1,125]
s: 188
i: [1,126,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,193,0,2,1,126]
s: 193
i: [1,127,0,2,1,0,0,1] fc1_xboole_0
a: [3,198,0,2,1,127]
s: 198
i: [1,158,0,2,1,0,0,2] rc2_xboole_0
a: [3,208,0,2,1,158]
s: 208
i: [1,159,0,2,1,0,0,255] rc4_ordinal1
a: [3,213,0,2,1,159]
s: 213
i: [1,160,0,2,1,0,0,2] rc1_xboole_0
a: [3,218,0,2,1,160]
s: 218
i: [1,161,0,2,1,0,0,255] rc1_ordinal1
a: [3,223,0,2,1,161]
s: 223
i: [1,163,0,2,1,0,0,255] rc7_funct_1
a: [3,228,0,2,1,163]
s: 228
i: [1,162,0,2,1,0,0,255] rc7_funct_1
a: [3,233,0,2,1,162]
s: 233
i: [1,167,0,2,1,0,0,255] rc2_ordinal1
a: [3,238,0,2,1,167]
s: 238
i: [1,166,0,2,1,0,0,255] rc2_ordinal1
a: [3,243,0,2,1,166]
s: 243
i: [1,165,0,2,1,0,0,255] rc2_ordinal1
a: [3,248,0,2,1,165]
s: 248
i: [1,164,0,2,1,0,0,255] rc2_ordinal1
a: [3,253,0,2,1,164]
s: 253
i: [1,169,0,2,1,0,0,255] rc5_ordinal1
a: [3,258,0,2,1,169]
s: 258
i: [1,168,0,2,1,0,0,255] rc5_ordinal1
a: [3,263,0,2,1,168]
s: 263
i: [1,171,0,2,1,0,0,255] rc1_relat_1
a: [3,268,0,2,1,171]
s: 268
i: [1,170,0,2,1,0,0,255] rc1_relat_1
a: [3,273,0,2,1,170]
s: 273
i: [1,174,0,2,1,0,0,255] rc4_funct_1
a: [3,278,0,2,1,174]
s: 278
i: [1,173,0,2,1,0,0,255] rc4_funct_1
a: [3,283,0,2,1,173]
s: 283
i: [1,172,0,2,1,0,0,255] rc4_funct_1
a: [3,288,0,2,1,172]
s: 288
i: [1,175,0,2,1,0,0,255] rc2_relat_1
a: [3,293,0,2,1,175]
s: 293
i: [1,177,0,2,1,0,0,255] rc3_funct_1
a: [3,298,0,2,1,177]
s: 298
i: [1,176,0,2,1,0,0,255] rc3_funct_1
a: [3,303,0,2,1,176]
s: 303
i: [1,179,0,2,1,0,0,255] rc1_funct_1
a: [3,308,0,2,1,179]
s: 308
i: [1,178,0,2,1,0,0,255] rc1_funct_1
a: [3,313,0,2,1,178]
s: 313
i: [1,182,0,2,1,0,0,255] rc5_funct_1
a: [3,318,0,2,1,182]
s: 318
i: [1,181,0,2,1,0,0,255] rc5_funct_1
a: [3,323,0,2,1,181]
s: 323
i: [1,180,0,2,1,0,0,255] rc5_funct_1
a: [3,328,0,2,1,180]
s: 328
i: [1,184,0,2,1,0,0,255] rc2_funct_1
a: [3,333,0,2,1,184]
s: 333
i: [1,183,0,2,1,0,0,255] rc2_funct_1
a: [3,338,0,2,1,183]
s: 338
i: [1,128,0,3,1,0,0,255] d2_xboole_0
a: [3,203,0,3,1,128]
s: 203
r: 193
i: [1,132,0,4,2,0,0,255] cc6_ordinal1
s: 132
i: [1,133,0,4,2,0,0,255] cc7_funct_1
s: 133
d: [2,341,1,2,1,1,40,132,213]
a: [3,346,1,2,1,341]
s: 346
i: [1,134,0,4,2,0,0,255] cc7_ordinal1
s: 134
i: [1,135,0,4,2,0,0,255] cc3_ordinal1
s: 135
d: [2,342,1,2,1,1,40,132,258]
a: [3,351,1,2,1,342]
s: 351
i: [1,136,0,4,2,0,0,255] cc1_funct_1
s: 136
d: [2,353,1,2,1,1,40,133,218]
a: [3,358,1,2,1,353]
s: 358
i: [1,137,0,4,2,0,0,255] cc1_relat_1
s: 137
i: [1,139,0,4,2,0,0,255] cc1_ordinal1
s: 139
d: [2,354,1,2,1,1,40,133,198]
a: [3,363,1,2,1,354]
s: 363
i: [1,140,0,4,2,0,0,255] cc1_ordinal1
s: 140
d: [2,365,1,2,1,1,40,134,218]
a: [3,370,1,2,1,365]
s: 370
i: [1,156,0,4,1,0,0,255] existence_m1_subset_1
s: 156
i: [1,138,0,5,2,0,0,1] t6_boole
s: 138
d: [2,456,1,6,2,1,44,188,138]
s: 456
i: [1,155,0,5,2,0,0,255] t7_boole
s: 155
d: [2,366,1,2,1,1,40,134,198]
a: [3,375,1,2,1,366]
s: 375
i: [1,131,0,6,2,0,0,1] t55_classes1
s: 131
i: [1,142,0,6,3,0,0,255] cc2_ordinal1
s: 142
d: [2,460,1,5,2,1,44,188,131]
d: [2,463,1,2,1,1,19,460]
a: [3,467,1,2,1,463]
s: 467
i: [1,152,0,6,2,0,0,255] t1_subset
s: 152
d: [2,461,2,7,3,1,44,456,131]
d: [2,462,2,4,2,1,19,461]
s: 462
i: [1,157,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 157
i: [1,129,0,7,3,0,0,255] cc8_funct_1
s: 129
d: [2,471,3,2,1,2,40,462,218]
a: [3,476,3,2,1,471]
s: 476
i: [1,130,0,7,3,0,0,255] cc8_funct_1
s: 130
d: [2,470,2,4,2,1,44,467,138]
s: 470
i: [1,141,0,7,3,0,0,255] cc5_ordinal1
s: 141
i: [1,154,0,7,3,0,0,1] t8_boole
s: 154
d: [2,377,1,2,1,1,40,135,218]
a: [3,382,1,2,1,377]
s: 382
d: [2,481,4,0,0,2,52,476,382,139]
e: 481
d: [2,483,2,2,1,1,40,382,139]
d: [2,485,2,0,0,2,21,483,476]
e: 485
f: 485,471,377,160,125
% Refutation found. Thanks to Tanya!
% SZS status Theorem for classes1__t56_classes1
% SZS output start Proof for classes1__t56_classes1
1. k1_xboole_0 = k5_classes1(k1_xboole_0) [input t56_classes1]
2. ~k1_xboole_0 = k5_classes1(k1_xboole_0) [negated conjecture 1]
18. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
34. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
39. ? [X0] : v1_xboole_0(X0) [input rc1_xboole_0]
47. ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0) [input t6_boole]
49. ! [X0] : (v1_ordinal1(X0) => k5_classes1(X0) = X0) [input t55_classes1]
50. k1_xboole_0 != k5_classes1(k1_xboole_0) [flattening 2]
67. ! [X0] : (k5_classes1(X0) = X0 | ~v1_ordinal1(X0)) [ennf transformation 49]
71. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 34]
74. ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) [ennf transformation 47]
75. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 18]
101. ? [X0] : v1_xboole_0(X0) => v1_xboole_0(sK3) [choice axiom]
102. v1_xboole_0(sK3) [skolemisation 39,101]
125. k1_xboole_0 != k5_classes1(k1_xboole_0) [cnf transformation 50] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:1,thDist:-1}
131. k5_classes1(X0) = X0 | ~v1_ordinal1(X0) [cnf transformation 67] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
135. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 71] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
138. k1_xboole_0 = X0 | ~v1_xboole_0(X0) [cnf transformation 74] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
139. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 75] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
160. v1_xboole_0(sK3) [cnf transformation 102] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:2,thDist:-2}
186. 1 <=> k1_xboole_0 = k5_classes1(k1_xboole_0) [avatar definition]
188. k1_xboole_0 != k5_classes1(k1_xboole_0) <- (~1) [avatar component clause 186] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:1,thDist:-1}
189. ~1 [avatar split clause 125,186]
216. 7 <=> v1_xboole_0(sK3) [avatar definition]
218. v1_xboole_0(sK3) <- (7) [avatar component clause 216] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
219. 7 [avatar split clause 160,216]
377. v3_ordinal1(sK3) <- (7) [resolution 135,218] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
380. 38 <=> v3_ordinal1(sK3) [avatar definition]
382. v3_ordinal1(sK3) <- (38) [avatar component clause 380] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:3,thDist:-3}
383. 38 | ~7 [avatar split clause 377,216,380]
456. k5_classes1(X0) != X0 | ~v1_xboole_0(X0) <- (~1) [superposition 188,138] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
461. X0 != X0 | ~v1_xboole_0(X0) | ~v1_ordinal1(X0) <- (~1) [superposition 456,131] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
462. ~v1_xboole_0(X0) | ~v1_ordinal1(X0) <- (~1) [trivial inequality removal 461] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
471. ~v1_ordinal1(sK3) <- (~1, 7) [resolution 462,218] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
474. 51 <=> v1_ordinal1(sK3) [avatar definition]
476. ~v1_ordinal1(sK3) <- (~51) [avatar component clause 474] {a:3,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
477. ~51 | 1 | ~7 [avatar split clause 471,216,186,474]
483. v1_ordinal1(sK3) <- (38) [resolution 382,139] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:4,thDist:-4}
485. $false <- (38, ~51) [subsumption resolution 483,476] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
486. ~38 | 51 [avatar contradiction clause 485]
487. $false [avatar sat refutation 189,219,383,477,486] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:214,thDist:-214}
% SZS output end Proof for classes1__t56_classes1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 48
% >>> Preprocessing
% Introduced skolems: 15
% Pure predicates: 6
% 
% >>> Saturation
% Initial clauses: 60
% Generated clauses: 154
% Active clauses: 61
% Passive clauses: 79
% Final active clauses: 60
% Final passive clauses: 17
% Main loop iterations started: 62
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 1
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 30
% Unit resulting resolution: 1
% Forward superposition: 2
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 57
% SAT solver unit clauses: 31
% SAT solver binary clauses: 23
% 
% Memory used [KB]: 9594
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

