i: [1,79,0,8,2,1,0,0] t19_mcart_1
a: [3,118,0,4,1,79]
s: 118
i: [1,78,0,8,2,1,0,0] t19_mcart_1
a: [3,127,0,4,1,78]
s: 127
i: [1,80,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,143,0,2,1,80]
s: 143
i: [1,81,0,2,1,0,0,255] fc1_xboole_0
a: [3,148,0,2,1,81]
s: 148
i: [1,105,0,2,1,0,0,2] rc2_xboole_0
a: [3,158,0,2,1,105]
s: 158
i: [1,106,0,2,1,0,0,2] rc1_xboole_0
a: [3,163,0,2,1,106]
s: 163
i: [1,108,0,2,1,0,0,2] rc1_relat_1
a: [3,168,0,2,1,108]
s: 168
i: [1,107,0,2,1,0,0,2] rc1_relat_1
a: [3,173,0,2,1,107]
s: 173
i: [1,109,0,2,1,0,0,255] rc2_relat_1
a: [3,178,0,2,1,109]
s: 178
i: [1,75,0,9,1,1,0,0] t19_mcart_1
a: [3,138,0,9,1,75]
s: 138
i: [1,82,0,3,1,0,0,255] d2_xboole_0
a: [3,153,0,3,1,82]
s: 153
r: 143
i: [1,83,0,4,2,0,0,2] cc1_relat_1
s: 83
i: [1,99,0,4,1,0,0,255] existence_m1_subset_1
s: 99
d: [2,181,1,2,1,1,40,83,163]
a: [3,186,1,2,1,181]
s: 186
i: [1,100,0,4,1,0,0,1] fc3_xboole_0
s: 100
i: [1,101,0,4,1,0,0,1] fc6_relat_1
s: 101
d: [2,182,1,2,1,1,40,83,148]
a: [3,191,1,2,1,182]
s: 191
i: [1,84,0,5,2,0,0,255] t6_boole
s: 84
i: [1,90,0,5,2,0,0,1] t7_boole
s: 90
i: [1,111,0,5,1,0,0,1] d2_tarski
s: 111
i: [1,113,0,5,1,0,0,1] d2_tarski
s: 113
d: [2,196,1,8,1,1,40,90,138]
a: [3,200,1,8,1,196]
s: 200
i: [1,88,0,6,2,0,0,255] t1_subset
s: 88
d: [2,204,1,9,1,1,40,88,138]
a: [3,210,1,9,1,204]
s: 210
i: [1,103,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 103
i: [1,89,0,7,3,0,0,1] t8_boole
s: 89
d: [2,212,1,9,1,1,40,103,138]
a: [3,218,1,9,1,212]
s: 218
i: [1,102,0,7,1,0,0,1] commutativity_k2_tarski
s: 102
d: [2,205,1,5,1,0,40,88,111]
s: 205
i: [1,87,0,8,3,0,0,255] t2_subset
s: 87
i: [1,104,0,8,3,0,0,1] fc10_subset_1
s: 104
d: [2,206,1,5,1,0,40,88,113]
s: 206
i: [1,91,0,9,2,0,0,1] t10_mcart_1
s: 91
d: [2,234,1,6,1,1,40,91,138]
a: [3,238,1,6,1,234]
s: 238
i: [1,92,0,9,2,0,0,1] t10_mcart_1
s: 92
i: [1,114,0,11,3,0,0,1] d2_tarski
s: 114
d: [2,259,2,0,0,3,52,118,127,238,114]
e: 259
d: [2,262,2,8,2,1,40,114,238]
d: [2,266,2,4,1,2,21,262,127]
d: [2,267,2,0,0,3,21,266,118]
e: 267
r: 127
a: [3,122,0,4,1,79]
s: 122
a: [3,126,0,4,1,78]
s: 126
r: 238
r: 118
d: [2,272,0,3,1,2,23,118,126]
a: [3,276,0,3,1,272]
s: 276
i: [1,77,0,8,2,1,0,0] t19_mcart_1
a: [3,132,0,4,1,77]
s: 132
i: [1,97,0,17,3,0,0,1] d2_tarski
s: 97
d: [2,253,1,6,1,1,40,92,138]
a: [3,257,1,6,1,253]
s: 257
d: [2,278,2,0,0,3,52,122,132,257,114]
e: 278
d: [2,279,2,8,2,1,40,257,114]
d: [2,284,2,4,1,2,21,279,132]
d: [2,285,2,0,0,3,21,284,122]
e: 285
i: [1,76,0,8,2,1,0,0] t19_mcart_1
d: [2,269,2,4,1,2,21,262,118]
f: 285,269,253,234,75,76,77,78,79
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t19_mcart_1
% SZS output start Proof for mcart_1__t19_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4))) => ((k2_xtuple_0(X0) = X4 | k2_xtuple_0(X0) = X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [input t19_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4))) => ((k2_xtuple_0(X0) = X4 | k2_xtuple_0(X0) = X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [negated conjecture 1]
8. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
26. ! [X0] : ! [X1] : ! [X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [input t10_mcart_1]
32. ~! [X0,X1,X2,X3,X4] : (r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4))) => ((k2_xtuple_0(X0) = X4 | k2_xtuple_0(X0) = X3) & (k1_xtuple_0(X0) = X2 | k1_xtuple_0(X0) = X1))) [flattening 2]
37. ! [X0,X1,X2] : (r2_hidden(X0,k2_zfmisc_1(X1,X2)) => (r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1))) [flattening 26]
38. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 8]
42. ? [X0,X1,X2,X3,X4] : (((k2_xtuple_0(X0) != X4 & k2_xtuple_0(X0) != X3) | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4)))) [ennf transformation 32]
54. ! [X0,X1,X2] : ((r2_hidden(k2_xtuple_0(X0),X2) & r2_hidden(k1_xtuple_0(X0),X1)) | ~r2_hidden(X0,k2_zfmisc_1(X1,X2))) [ennf transformation 37]
58. ? [X0,X1,X2,X3,X4] : (((k2_xtuple_0(X0) != X4 & k2_xtuple_0(X0) != X3) | (k1_xtuple_0(X0) != X2 & k1_xtuple_0(X0) != X1)) & r2_hidden(X0,k2_zfmisc_1(k2_tarski(X1,X2),k2_tarski(X3,X4)))) => (((sK4 != k2_xtuple_0(sK0) & sK3 != k2_xtuple_0(sK0)) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k2_tarski(sK3,sK4)))) [choice axiom]
59. ((sK4 != k2_xtuple_0(sK0) & sK3 != k2_xtuple_0(sK0)) | (sK2 != k1_xtuple_0(sK0) & sK1 != k1_xtuple_0(sK0))) & r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k2_tarski(sK3,sK4))) [skolemisation 42,58]
60. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 38]
61. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 60]
62. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 61]
63. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK5(X0,X1,X2) != X1 & sK5(X0,X1,X2) != X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & (sK5(X0,X1,X2) = X1 | sK5(X0,X1,X2) = X0 | r2_hidden(sK5(X0,X1,X2),X2)))) [choice axiom]
64. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK5(X0,X1,X2) != X1 & sK5(X0,X1,X2) != X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & (sK5(X0,X1,X2) = X1 | sK5(X0,X1,X2) = X0 | r2_hidden(sK5(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 62,63]
75. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k2_tarski(sK3,sK4))) [cnf transformation 59] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
76. sK3 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
77. sK3 != k2_xtuple_0(sK0) | sK2 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
78. sK4 != k2_xtuple_0(sK0) | sK1 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
79. sK4 != k2_xtuple_0(sK0) | sK2 != k1_xtuple_0(sK0) [cnf transformation 59] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
91. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k1_xtuple_0(X0),X1) [cnf transformation 54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
92. ~r2_hidden(X0,k2_zfmisc_1(X1,X2)) | r2_hidden(k2_xtuple_0(X0),X2) [cnf transformation 54] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
93. X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2) | k2_tarski(X0,X1) != X2 [cnf transformation 64] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
114. ~r2_hidden(X4,k2_tarski(X0,X1)) | X0 = X4 | X1 = X4 [equality resolution 93] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
116. 1 <=> sK2 = k1_xtuple_0(sK0) [avatar definition]
118. sK2 != k1_xtuple_0(sK0) <- (~1) [avatar component clause 116] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
120. 2 <=> sK4 = k2_xtuple_0(sK0) [avatar definition]
122. sK4 != k2_xtuple_0(sK0) <- (~2) [avatar component clause 120] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
123. ~1 | ~2 [avatar split clause 79,120,116]
125. 3 <=> sK1 = k1_xtuple_0(sK0) [avatar definition]
128. ~3 | ~2 [avatar split clause 78,120,125]
130. 4 <=> sK3 = k2_xtuple_0(sK0) [avatar definition]
132. sK3 != k2_xtuple_0(sK0) <- (~4) [avatar component clause 130] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
133. ~1 | ~4 [avatar split clause 77,130,116]
134. ~3 | ~4 [avatar split clause 76,130,125]
136. 5 <=> r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k2_tarski(sK3,sK4))) [avatar definition]
138. r2_hidden(sK0,k2_zfmisc_1(k2_tarski(sK1,sK2),k2_tarski(sK3,sK4))) <- (5) [avatar component clause 136] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
139. 5 [avatar split clause 75,136]
234. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (5) [resolution 91,138] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
236. 19 <=> r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) [avatar definition]
238. r2_hidden(k1_xtuple_0(sK0),k2_tarski(sK1,sK2)) <- (19) [avatar component clause 236] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
239. 19 | ~5 [avatar split clause 234,136,236]
253. r2_hidden(k2_xtuple_0(sK0),k2_tarski(sK3,sK4)) <- (5) [resolution 92,138] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
255. 22 <=> r2_hidden(k2_xtuple_0(sK0),k2_tarski(sK3,sK4)) [avatar definition]
257. r2_hidden(k2_xtuple_0(sK0),k2_tarski(sK3,sK4)) <- (22) [avatar component clause 255] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
258. 22 | ~5 [avatar split clause 253,136,255]
262. sK1 = k1_xtuple_0(sK0) | sK2 = k1_xtuple_0(sK0) <- (19) [resolution 114,238] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
269. sK1 = k1_xtuple_0(sK0) <- (~1, 19) [subsumption resolution 262,118] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
270. 3 | 1 | ~19 [avatar split clause 269,236,116,125]
279. sK3 = k2_xtuple_0(sK0) | sK4 = k2_xtuple_0(sK0) <- (22) [resolution 257,114] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
284. sK4 = k2_xtuple_0(sK0) <- (~4, 22) [subsumption resolution 279,132] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
285. $false <- (~2, ~4, 22) [subsumption resolution 284,122] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
286. 2 | 4 | ~22 [avatar contradiction clause 285]
287. $false [avatar sat refutation 123,128,133,134,139,239,258,270,286] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:127,thDist:-127}
% SZS output end Proof for mcart_1__t19_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 30
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 2
% 
% >>> Saturation
% Initial clauses: 35
% Generated clauses: 112
% Active clauses: 42
% Passive clauses: 52
% Final active clauses: 38
% Final passive clauses: 10
% Main loop iterations started: 42
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 5
% Bw demodulations: 3
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 2
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 27
% Unit resulting resolution: 2
% Forward superposition: 6
% Backward superposition: 8
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 4
% Split components: 8
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 30
% SAT solver unit clauses: 9
% SAT solver binary clauses: 13
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

