i: [1,106,0,2,1,1,0,0] t26_ordinal3
a: [3,171,0,2,1,106]
s: 171
i: [1,105,0,2,1,1,0,0] t26_ordinal3
a: [3,176,0,2,1,105]
s: 176
i: [1,108,0,6,2,1,0,0] t26_ordinal3
a: [3,157,0,3,1,108]
s: 157
i: [1,107,0,5,1,1,0,0] t26_ordinal3
a: [3,166,0,5,1,107]
s: 166
i: [1,109,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,181,0,2,1,109]
s: 181
i: [1,110,0,2,1,0,0,1] fc1_xboole_0
a: [3,186,0,2,1,110]
s: 186
i: [1,139,0,7,3,0,0,2] reflexivity_r1_ordinal1
a: [3,195,0,2,1,139]
s: 195
d: [2,252,1,0,0,2,40,195,176]
e: 252
d: [2,251,1,0,0,2,40,195,171]
e: 251
d: [2,250,1,0,0,2,52,171,195]
e: 250
r: 195
i: [1,144,0,2,1,0,0,2] rc2_xboole_0
a: [3,203,0,2,1,144]
s: 203
i: [1,145,0,2,1,0,0,2] rc4_ordinal1
a: [3,208,0,2,1,145]
s: 208
i: [1,146,0,2,1,0,0,2] rc1_xboole_0
a: [3,213,0,2,1,146]
s: 213
i: [1,147,0,2,1,0,0,1] rc1_ordinal1
a: [3,218,0,2,1,147]
s: 218
i: [1,151,0,2,1,0,0,255] rc2_ordinal1
a: [3,223,0,2,1,151]
s: 223
i: [1,150,0,2,1,0,0,255] rc2_ordinal1
a: [3,228,0,2,1,150]
s: 228
i: [1,149,0,2,1,0,0,255] rc2_ordinal1
a: [3,233,0,2,1,149]
s: 233
i: [1,148,0,2,1,0,0,255] rc2_ordinal1
a: [3,238,0,2,1,148]
s: 238
i: [1,153,0,2,1,0,0,2] rc5_ordinal1
a: [3,243,0,2,1,153]
s: 243
i: [1,152,0,2,1,0,0,2] rc5_ordinal1
a: [3,248,0,2,1,152]
s: 248
i: [1,135,0,3,1,0,0,3] reflexivity_r1_tarski
s: 135
i: [1,111,0,3,1,0,0,255] d2_xboole_0
a: [3,191,0,3,1,111]
s: 191
r: 181
i: [1,113,0,4,2,0,0,2] cc6_ordinal1
s: 113
i: [1,114,0,4,2,0,0,255] cc9_ordinal1
s: 114
d: [2,257,1,2,1,1,40,113,208]
a: [3,262,1,2,1,257]
s: 262
i: [1,115,0,4,2,0,0,2] cc7_ordinal1
s: 115
d: [2,258,1,2,1,1,40,113,243]
a: [3,267,1,2,1,258]
s: 267
i: [1,116,0,4,2,0,0,1] cc3_ordinal1
s: 116
i: [1,118,0,4,2,0,0,255] cc1_ordinal1
s: 118
d: [2,305,1,2,1,1,40,118,176]
a: [3,314,1,2,1,305]
s: 314
i: [1,119,0,4,2,0,0,255] cc1_ordinal1
s: 119
d: [2,306,1,2,1,1,40,118,171]
a: [3,319,1,2,1,306]
s: 319
i: [1,134,0,4,1,0,0,4] existence_m1_subset_1
s: 134
i: [1,117,0,5,2,0,0,1] t6_boole
s: 117
d: [2,336,1,2,1,1,40,119,176]
a: [3,345,1,2,1,336]
s: 345
i: [1,131,0,5,2,0,0,5] t7_boole
s: 131
d: [2,337,1,2,1,1,40,119,171]
a: [3,350,1,2,1,337]
s: 350
a: [3,198,0,5,2,139]
s: 198
i: [1,124,0,6,2,0,0,3] t3_xboole_1
s: 124
d: [2,269,1,2,1,1,40,114,213]
a: [3,274,1,2,1,269]
s: 274
i: [1,125,0,6,3,0,0,255] cc2_ordinal1
s: 125
d: [2,270,1,2,1,1,40,114,186]
a: [3,279,1,2,1,270]
s: 279
i: [1,126,0,6,2,0,0,5] t1_subset
s: 126
i: [1,136,0,6,2,0,0,5] antisymmetry_r2_hidden
s: 136
d: [2,281,1,2,1,1,40,115,213]
a: [3,286,1,2,1,281]
s: 286
i: [1,122,0,7,3,0,0,4] cc2_ordinal2
s: 122
d: [2,282,1,2,1,1,40,115,186]
a: [3,291,1,2,1,282]
s: 291
i: [1,123,0,7,3,0,0,4] cc5_ordinal1
s: 123
i: [1,128,0,7,2,0,0,3] t3_subset
s: 128
d: [2,293,1,2,1,1,40,116,213]
a: [3,298,1,2,1,293]
s: 298
i: [1,129,0,7,2,0,0,3] t3_subset
s: 129
d: [2,294,1,2,1,1,40,116,186]
a: [3,303,1,2,1,294]
s: 303
i: [1,130,0,7,3,0,0,1] t8_boole
s: 130
i: [1,112,0,8,3,0,0,255] cc10_ordinal1
s: 112
d: [2,308,1,2,1,1,40,118,218]
a: [3,329,1,2,1,308]
s: 329
i: [1,127,0,8,3,0,0,5] t2_subset
s: 127
d: [2,339,1,2,1,1,40,119,218]
a: [3,360,1,2,1,339]
s: 360
i: [1,137,0,8,3,0,0,1] fc5_ordinal2
s: 137
i: [1,138,0,8,3,0,0,1] fc5_ordinal2
s: 138
d: [2,307,2,2,1,1,40,118,262]
a: [3,324,2,2,1,307]
s: 324
i: [1,140,0,8,3,0,0,1] dt_k10_ordinal2
s: 140
d: [2,310,2,2,1,1,40,118,267]
a: [3,334,2,2,1,310]
s: 334
i: [1,120,0,9,3,0,0,1] t24_ordinal3
s: 120
i: [1,121,0,9,3,0,0,1] t24_ordinal3
s: 121
d: [2,422,1,7,3,1,44,120,166]
d: [2,423,1,5,2,2,21,422,171]
d: [2,424,1,3,1,3,21,423,176]
a: [3,428,1,3,1,424]
s: 428
i: [1,133,0,9,3,0,0,4] t5_subset
s: 133
d: [2,430,1,7,3,1,44,121,166]
d: [2,431,1,5,2,2,21,430,171]
d: [2,432,1,3,1,3,21,431,176]
a: [3,436,1,3,1,432]
s: 436
i: [1,132,0,10,3,0,0,4] t4_subset
s: 132
i: [1,141,0,10,4,0,0,2] connectedness_r1_ordinal1
s: 141
d: [2,438,2,5,2,1,44,428,117]
s: 438
i: [1,142,0,10,4,0,0,2] redefinition_r1_ordinal1
s: 142
d: [2,450,2,7,3,1,40,142,428]
d: [2,459,2,5,2,2,21,450,303]
d: [2,460,2,3,1,3,21,459,176]
a: [3,464,2,3,1,460]
s: 464
d: [2,475,3,0,0,2,52,157,464,124]
e: 475
d: [2,476,3,3,1,1,40,464,124]
d: [2,479,3,0,0,2,21,476,157]
e: 479
r: 157
a: [3,156,0,3,1,108]
s: 156
r: 166
r: 176
r: 314
r: 345
r: 428
r: 438
r: 464
a: [3,161,0,3,1,108]
s: 161
d: [2,486,1,3,1,2,23,428,156]
a: [3,499,1,3,1,486]
s: 499
d: [2,482,0,5,1,2,23,166,156]
a: [3,492,0,5,1,482]
s: 492
d: [2,452,2,7,3,1,40,142,436]
d: [2,468,2,5,2,2,21,452,303]
d: [2,469,2,3,1,3,21,468,171]
a: [3,473,2,3,1,469]
s: 473
d: [2,510,3,0,0,2,52,161,473,124]
e: 510
d: [2,511,3,3,1,1,40,473,124]
d: [2,514,3,0,0,2,21,511,161]
e: 514
f: 514,476,469,460,432,424,294,110,105,106,107,108
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t26_ordinal3
% SZS output start Proof for ordinal3__t26_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k10_ordinal2(X0,X1) = k1_xboole_0 => (k1_xboole_0 = X1 & k1_xboole_0 = X0)))) [input t26_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k10_ordinal2(X0,X1) = k1_xboole_0 => (k1_xboole_0 = X1 & k1_xboole_0 = X0)))) [negated conjecture 1]
8. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
22. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
30. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
34. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r1_ordinal1(X1,k10_ordinal2(X0,X1)) & r1_ordinal1(X0,k10_ordinal2(X0,X1))))) [input t24_ordinal3]
37. ! [X0] : (r1_tarski(X0,k1_xboole_0) => k1_xboole_0 = X0) [input t3_xboole_1]
52. ? [X0] : (? [X1] : (((k1_xboole_0 != X1 | k1_xboole_0 != X0) & k10_ordinal2(X0,X1) = k1_xboole_0) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
53. ? [X0] : (? [X1] : ((k1_xboole_0 != X1 | k1_xboole_0 != X0) & k10_ordinal2(X0,X1) = k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 52]
58. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 8]
61. ! [X0] : (! [X1] : ((r1_ordinal1(X1,k10_ordinal2(X0,X1)) & r1_ordinal1(X0,k10_ordinal2(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 34]
64. ! [X0] : (k1_xboole_0 = X0 | ~r1_tarski(X0,k1_xboole_0)) [ennf transformation 37]
84. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 30]
85. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 84]
86. ? [X0] : (? [X1] : ((k1_xboole_0 != X1 | k1_xboole_0 != X0) & k10_ordinal2(X0,X1) = k1_xboole_0 & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : ((k1_xboole_0 != X1 | k1_xboole_0 != sK0) & k1_xboole_0 = k10_ordinal2(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
87. ? [X1] : ((k1_xboole_0 != X1 | k1_xboole_0 != sK0) & k1_xboole_0 = k10_ordinal2(sK0,X1) & v3_ordinal1(X1)) => ((k1_xboole_0 != sK1 | k1_xboole_0 != sK0) & k1_xboole_0 = k10_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
88. ((k1_xboole_0 != sK1 | k1_xboole_0 != sK0) & k1_xboole_0 = k10_ordinal2(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 53,87,86]
92. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 85]
105. v3_ordinal1(sK0) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
106. v3_ordinal1(sK1) [cnf transformation 88] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
107. k1_xboole_0 = k10_ordinal2(sK0,sK1) [cnf transformation 88] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
108. k1_xboole_0 != sK1 | k1_xboole_0 != sK0 [cnf transformation 88] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
110. v1_xboole_0(k1_xboole_0) [cnf transformation 22] {a:0,w:2,wCS:800,sine:1,thAx:0,allAx:1,thDist:-1}
116. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 58] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
120. r1_ordinal1(X0,k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 61] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
121. r1_ordinal1(X1,k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 61] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
124. ~r1_tarski(X0,k1_xboole_0) | k1_xboole_0 = X0 [cnf transformation 64] {a:0,w:6,wCS:2400,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
142. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 92] {a:0,w:10,wCS:4000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
155. 1 <=> k1_xboole_0 = sK0 [avatar definition]
159. 2 <=> k1_xboole_0 = sK1 [avatar definition]
161. k1_xboole_0 != sK1 <- (~2) [avatar component clause 159] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
162. ~1 | ~2 [avatar split clause 108,159,155]
164. 3 <=> k1_xboole_0 = k10_ordinal2(sK0,sK1) [avatar definition]
166. k1_xboole_0 = k10_ordinal2(sK0,sK1) <- (3) [avatar component clause 164] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
167. 3 [avatar split clause 107,164]
169. 4 <=> v3_ordinal1(sK1) [avatar definition]
171. v3_ordinal1(sK1) <- (4) [avatar component clause 169] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
172. 4 [avatar split clause 106,169]
174. 5 <=> v3_ordinal1(sK0) [avatar definition]
176. v3_ordinal1(sK0) <- (5) [avatar component clause 174] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
177. 5 [avatar split clause 105,174]
184. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
186. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 184] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
187. 7 [avatar split clause 110,184]
294. v3_ordinal1(k1_xboole_0) <- (7) [resolution 116,186] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
301. 28 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
303. v3_ordinal1(k1_xboole_0) <- (28) [avatar component clause 301] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
304. 28 | ~7 [avatar split clause 294,184,301]
422. r1_ordinal1(sK0,k1_xboole_0) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3) [superposition 120,166] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
423. r1_ordinal1(sK0,k1_xboole_0) | ~v3_ordinal1(sK0) <- (3, 4) [subsumption resolution 422,171] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
424. r1_ordinal1(sK0,k1_xboole_0) <- (3, 4, 5) [subsumption resolution 423,176] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
426. 43 <=> r1_ordinal1(sK0,k1_xboole_0) [avatar definition]
428. r1_ordinal1(sK0,k1_xboole_0) <- (43) [avatar component clause 426] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
429. 43 | ~3 | ~4 | ~5 [avatar split clause 424,174,169,164,426]
430. r1_ordinal1(sK1,k1_xboole_0) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3) [superposition 121,166] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
431. r1_ordinal1(sK1,k1_xboole_0) | ~v3_ordinal1(sK0) <- (3, 4) [subsumption resolution 430,171] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
432. r1_ordinal1(sK1,k1_xboole_0) <- (3, 4, 5) [subsumption resolution 431,176] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
434. 44 <=> r1_ordinal1(sK1,k1_xboole_0) [avatar definition]
436. r1_ordinal1(sK1,k1_xboole_0) <- (44) [avatar component clause 434] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
437. 44 | ~3 | ~4 | ~5 [avatar split clause 432,174,169,164,434]
450. r1_tarski(sK0,k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(sK0) <- (43) [resolution 142,428] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
452. r1_tarski(sK1,k1_xboole_0) | ~v3_ordinal1(k1_xboole_0) | ~v3_ordinal1(sK1) <- (44) [resolution 142,436] {a:2,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
459. r1_tarski(sK0,k1_xboole_0) | ~v3_ordinal1(sK0) <- (28, 43) [subsumption resolution 450,303] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
460. r1_tarski(sK0,k1_xboole_0) <- (5, 28, 43) [subsumption resolution 459,176] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
462. 45 <=> r1_tarski(sK0,k1_xboole_0) [avatar definition]
464. r1_tarski(sK0,k1_xboole_0) <- (45) [avatar component clause 462] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
465. 45 | ~5 | ~28 | ~43 [avatar split clause 460,426,301,174,462]
468. r1_tarski(sK1,k1_xboole_0) | ~v3_ordinal1(sK1) <- (28, 44) [subsumption resolution 452,303] {a:2,w:5,wCS:500,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
469. r1_tarski(sK1,k1_xboole_0) <- (4, 28, 44) [subsumption resolution 468,171] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
471. 46 <=> r1_tarski(sK1,k1_xboole_0) [avatar definition]
473. r1_tarski(sK1,k1_xboole_0) <- (46) [avatar component clause 471] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
474. 46 | ~4 | ~28 | ~44 [avatar split clause 469,434,301,169,471]
476. k1_xboole_0 = sK0 <- (45) [resolution 464,124] {a:3,w:3,wCS:300,goal:1,sine:3,thAx:0,allAx:6,thDist:-6}
481. 1 | ~45 [avatar split clause 476,462,155]
511. k1_xboole_0 = sK1 <- (46) [resolution 473,124] {a:3,w:3,wCS:300,goal:1,sine:3,thAx:0,allAx:6,thDist:-6}
514. $false <- (~2, 46) [subsumption resolution 511,161] {a:3,w:0,goal:1,sine:3,thAx:0,allAx:6,thDist:-6}
515. 2 | ~46 [avatar contradiction clause 514]
516. $false [avatar sat refutation 162,167,172,177,187,304,429,437,465,474,481,515] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:295,thDist:-295}
% SZS output end Proof for ordinal3__t26_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 41
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 49
% Generated clauses: 214
% Active clauses: 73
% Passive clauses: 107
% Final active clauses: 63
% Final passive clauses: 33
% Main loop iterations started: 74
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 14
% Bw demodulations: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Equational tautologies: 1
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 60
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 15
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 7
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 60
% SAT solver unit clauses: 16
% SAT solver binary clauses: 34
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% sat proof minimization: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

