i: [1,62,0,2,1,1,0,0] t8_wellord2
a: [3,115,0,2,1,62]
s: 115
i: [1,64,0,3,1,1,0,0] t8_wellord2
a: [3,105,0,3,1,64]
s: 105
i: [1,63,0,3,1,1,0,0] t8_wellord2
a: [3,110,0,3,1,63]
s: 110
i: [1,101,0,2,1,0,0,1] rc1_ordinal1
a: [3,120,0,2,1,101]
s: 120
i: [1,65,0,3,1,0,0,1] dt_k1_wellord2
s: 65
i: [1,66,0,3,1,0,0,255] fc1_wellord2
s: 66
i: [1,67,0,3,1,0,0,255] fc1_wellord2
s: 67
i: [1,68,0,3,1,0,0,255] fc2_wellord2
s: 68
i: [1,69,0,3,1,0,0,255] fc2_wellord2
s: 69
i: [1,70,0,3,1,0,0,255] fc3_wellord2
s: 70
i: [1,71,0,3,1,0,0,255] fc3_wellord2
s: 71
i: [1,95,0,3,1,0,0,1] reflexivity_r1_tarski
s: 95
i: [1,72,0,4,2,0,0,255] cc1_ordinal1
s: 72
d: [2,122,1,2,1,1,40,72,115]
a: [3,127,1,2,1,122]
s: 127
i: [1,73,0,4,2,0,0,255] cc1_ordinal1
s: 73
i: [1,94,0,4,1,0,0,255] existence_m1_subset_1
s: 94
d: [2,134,1,2,1,1,40,73,115]
a: [3,139,1,2,1,134]
s: 139
i: [1,96,0,4,1,0,0,255] fc6_relat_1
s: 96
d: [2,123,1,2,1,1,40,72,120]
a: [3,132,1,2,1,123]
s: 132
i: [1,74,0,5,2,0,0,255] fc5_wellord2
s: 74
i: [1,75,0,5,2,0,0,255] fc5_wellord2
s: 75
d: [2,135,1,2,1,1,40,73,120]
a: [3,144,1,2,1,135]
s: 144
i: [1,76,0,5,2,0,0,255] fc4_wellord2
s: 76
i: [1,77,0,5,2,0,0,255] fc4_wellord2
s: 77
i: [1,97,0,5,1,0,0,2] idempotence_k3_xboole_0
s: 97
i: [1,81,0,6,3,0,0,255] cc2_ordinal1
s: 81
i: [1,85,0,6,3,0,0,255] cc1_wellord1
s: 85
i: [1,86,0,6,3,0,0,255] cc1_wellord1
s: 86
i: [1,87,0,6,3,0,0,255] cc1_wellord1
s: 87
i: [1,88,0,6,3,0,0,255] cc1_wellord1
s: 88
i: [1,89,0,6,3,0,0,255] cc1_wellord1
s: 89
i: [1,99,0,6,2,0,0,2] dt_k2_wellord1
s: 99
i: [1,100,0,6,2,0,0,3] fc1_relat_1
s: 100
d: [2,152,1,6,2,0,40,87,65]
s: 152
i: [1,78,0,7,3,0,0,255] cc5_ordinal1
s: 78
i: [1,92,0,7,2,0,0,255] t3_subset
s: 92
d: [2,169,1,5,2,0,40,78,94]
s: 169
i: [1,93,0,7,2,0,0,255] t3_subset
s: 93
d: [2,170,1,5,1,0,40,92,94]
s: 170
i: [1,98,0,7,1,0,0,2] commutativity_k3_xboole_0
s: 98
i: [1,80,0,8,3,0,0,255] cc2_relat_1
s: 80
d: [2,171,2,5,2,0,40,169,73]
s: 171
i: [1,90,0,8,3,0,0,1] t25_wellord1
s: 90
d: [2,172,2,5,2,0,40,169,72]
s: 172
i: [1,91,0,10,2,0,0,1] t7_wellord2
s: 91
i: [1,79,0,11,2,0,0,255] d6_wellord1
s: 79
d: [2,156,1,6,2,0,40,89,65]
s: 156
i: [1,83,0,14,7,0,0,255] cc2_wellord1
s: 83
d: [2,193,1,20,7,0,40,83,75]
d: [2,194,1,17,6,0,21,193,69]
d: [2,195,1,14,5,0,21,194,71]
d: [2,196,1,11,4,0,21,195,67]
d: [2,197,1,8,3,0,21,196,65]
d: [2,198,1,5,2,0,21,197,77]
s: 198
d: [2,175,1,6,2,0,44,100,98]
s: 175
d: [2,176,1,6,2,0,44,100,98]
s: 176
d: [2,199,2,2,1,1,40,198,105]
a: [3,205,2,2,1,199]
s: 205
d: [2,177,1,6,2,0,40,80,94]
s: 177
d: [2,178,1,7,3,0,40,80,93]
s: 178
d: [2,149,1,8,2,0,40,85,96]
s: 149
d: [2,151,1,8,2,0,40,86,96]
s: 151
d: [2,221,2,4,2,1,40,178,110]
a: [3,231,2,2,1,221]
s: 231
d: [2,153,1,8,2,0,40,87,96]
s: 153
d: [2,155,1,8,2,0,40,88,96]
s: 155
d: [2,157,1,8,2,0,40,89,96]
s: 157
d: [2,174,1,8,3,0,40,93,78]
s: 174
d: [2,233,2,5,2,1,40,174,110]
d: [2,236,2,3,1,2,21,233,205]
a: [3,240,2,3,1,236]
s: 240
d: [2,180,1,12,4,0,44,90,91]
d: [2,182,1,9,3,0,21,180,65]
s: 182
d: [2,234,2,5,2,0,40,174,95]
s: 234
d: [2,183,1,13,2,0,44,79,97]
d: [2,191,1,9,1,0,21,183,96]
s: 191
d: [2,186,1,13,2,0,44,97,79]
d: [2,192,1,9,1,0,21,186,96]
s: 192
d: [2,242,2,8,3,0,40,182,198]
s: 242
d: [2,245,3,0,0,3,52,115,110,105,242]
e: 245
d: [2,246,3,5,2,1,40,242,110]
d: [2,250,3,2,1,2,21,246,105]
d: [2,251,3,0,0,3,21,250,115]
e: 251
f: 251,62,63,64
% Refutation found. Thanks to Tanya!
% SZS status Theorem for wellord2__t8_wellord2
% SZS output start Proof for wellord2__t8_wellord2
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (r1_tarski(X1,X0) => v2_wellord1(k1_wellord2(X1)))) [input t8_wellord2]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (r1_tarski(X1,X0) => v2_wellord1(k1_wellord2(X1)))) [negated conjecture 1]
7. ! [X0] : ((v1_wellord1(X0) & v8_relat_2(X0) & v6_relat_2(X0) & v4_relat_2(X0) & v1_relat_2(X0) & v1_relat_1(X0)) => (v2_wellord1(X0) & v1_relat_1(X0))) [input cc2_wellord1]
11. ! [X0] : v1_relat_1(k1_wellord2(X0)) [input dt_k1_wellord2]
19. ! [X0] : (v1_relat_2(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) [input fc1_wellord2]
20. ! [X0] : (v8_relat_2(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) [input fc2_wellord2]
21. ! [X0] : (v4_relat_2(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) [input fc3_wellord2]
22. ! [X0] : (v3_ordinal1(X0) => (v6_relat_2(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0)))) [input fc4_wellord2]
23. ! [X0] : (v3_ordinal1(X0) => (v1_wellord1(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0)))) [input fc5_wellord2]
28. ! [X0] : ! [X1] : (v1_relat_1(X1) => (v2_wellord1(X1) => v2_wellord1(k2_wellord1(X1,X0)))) [input t25_wellord1]
30. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => k1_wellord2(X0) = k2_wellord1(k1_wellord2(X1),X0)) [input t7_wellord2]
31. ! [X0,X1] : (v1_relat_1(X1) => (v2_wellord1(X1) => v2_wellord1(k2_wellord1(X1,X0)))) [flattening 28]
32. ! [X0,X1] : (r1_tarski(X0,X1) => k1_wellord2(X0) = k2_wellord1(k1_wellord2(X1),X0)) [flattening 30]
36. ? [X0] : (? [X1] : (~v2_wellord1(k1_wellord2(X1)) & r1_tarski(X1,X0)) & v3_ordinal1(X0)) [ennf transformation 2]
38. ! [X0] : ((v1_wellord1(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) | ~v3_ordinal1(X0)) [ennf transformation 23]
39. ! [X0] : ((v6_relat_2(k1_wellord2(X0)) & v1_relat_1(k1_wellord2(X0))) | ~v3_ordinal1(X0)) [ennf transformation 22]
45. ! [X0] : ((v2_wellord1(X0) & v1_relat_1(X0)) | (~v1_wellord1(X0) | ~v8_relat_2(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
46. ! [X0] : ((v2_wellord1(X0) & v1_relat_1(X0)) | ~v1_wellord1(X0) | ~v8_relat_2(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0)) [flattening 45]
49. ! [X0,X1] : ((v2_wellord1(k2_wellord1(X1,X0)) | ~v2_wellord1(X1)) | ~v1_relat_1(X1)) [ennf transformation 31]
50. ! [X0,X1] : (v2_wellord1(k2_wellord1(X1,X0)) | ~v2_wellord1(X1) | ~v1_relat_1(X1)) [flattening 49]
51. ! [X0,X1] : (k1_wellord2(X0) = k2_wellord1(k1_wellord2(X1),X0) | ~r1_tarski(X0,X1)) [ennf transformation 32]
54. ? [X0] : (? [X1] : (~v2_wellord1(k1_wellord2(X1)) & r1_tarski(X1,X0)) & v3_ordinal1(X0)) => (? [X1] : (~v2_wellord1(k1_wellord2(X1)) & r1_tarski(X1,sK0)) & v3_ordinal1(sK0)) [choice axiom]
55. ? [X1] : (~v2_wellord1(k1_wellord2(X1)) & r1_tarski(X1,sK0)) => (~v2_wellord1(k1_wellord2(sK1)) & r1_tarski(sK1,sK0)) [choice axiom]
56. (~v2_wellord1(k1_wellord2(sK1)) & r1_tarski(sK1,sK0)) & v3_ordinal1(sK0) [skolemisation 36,55,54]
62. v3_ordinal1(sK0) [cnf transformation 56] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
63. r1_tarski(sK1,sK0) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
64. ~v2_wellord1(k1_wellord2(sK1)) [cnf transformation 56] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
65. v1_relat_1(k1_wellord2(X0)) [cnf transformation 11] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
67. v1_relat_2(k1_wellord2(X0)) [cnf transformation 19] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
69. v8_relat_2(k1_wellord2(X0)) [cnf transformation 20] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
71. v4_relat_2(k1_wellord2(X0)) [cnf transformation 21] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
75. v1_wellord1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [cnf transformation 38] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
77. v6_relat_2(k1_wellord2(X0)) | ~v3_ordinal1(X0) [cnf transformation 39] {a:0,w:5,wCS:2000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
83. ~v1_wellord1(X0) | v2_wellord1(X0) | ~v8_relat_2(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
90. v2_wellord1(k2_wellord1(X1,X0)) | ~v2_wellord1(X1) | ~v1_relat_1(X1) [cnf transformation 50] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
91. k1_wellord2(X0) = k2_wellord1(k1_wellord2(X1),X0) | ~r1_tarski(X0,X1) [cnf transformation 51] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
103. 1 <=> v2_wellord1(k1_wellord2(sK1)) [avatar definition]
105. ~v2_wellord1(k1_wellord2(sK1)) <- (~1) [avatar component clause 103] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
106. ~1 [avatar split clause 64,103]
108. 2 <=> r1_tarski(sK1,sK0) [avatar definition]
110. r1_tarski(sK1,sK0) <- (2) [avatar component clause 108] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
111. 2 [avatar split clause 63,108]
113. 3 <=> v3_ordinal1(sK0) [avatar definition]
115. v3_ordinal1(sK0) <- (3) [avatar component clause 113] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
116. 3 [avatar split clause 62,113]
180. v2_wellord1(k1_wellord2(X1)) | ~v2_wellord1(k1_wellord2(X0)) | ~v1_relat_1(k1_wellord2(X0)) | ~r1_tarski(X1,X0) [superposition 90,91] {a:1,w:12,wCS:4800,sine:1,thAx:0,allAx:2,thDist:-2}
182. ~v2_wellord1(k1_wellord2(X0)) | v2_wellord1(k1_wellord2(X1)) | ~r1_tarski(X1,X0) [subsumption resolution 180,65] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
193. v2_wellord1(k1_wellord2(X0)) | ~v8_relat_2(k1_wellord2(X0)) | ~v6_relat_2(k1_wellord2(X0)) | ~v4_relat_2(k1_wellord2(X0)) | ~v1_relat_2(k1_wellord2(X0)) | ~v1_relat_1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [resolution 83,75] {a:1,w:20,wCS:8000,sine:255,thAx:0,allAx:2,thDist:-2}
194. v2_wellord1(k1_wellord2(X0)) | ~v6_relat_2(k1_wellord2(X0)) | ~v4_relat_2(k1_wellord2(X0)) | ~v1_relat_2(k1_wellord2(X0)) | ~v1_relat_1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [subsumption resolution 193,69] {a:1,w:17,wCS:6800,sine:255,thAx:0,allAx:2,thDist:-2}
195. v2_wellord1(k1_wellord2(X0)) | ~v6_relat_2(k1_wellord2(X0)) | ~v1_relat_2(k1_wellord2(X0)) | ~v1_relat_1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [subsumption resolution 194,71] {a:1,w:14,wCS:5600,sine:255,thAx:0,allAx:2,thDist:-2}
196. v2_wellord1(k1_wellord2(X0)) | ~v6_relat_2(k1_wellord2(X0)) | ~v1_relat_1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [subsumption resolution 195,67] {a:1,w:11,wCS:4400,sine:255,thAx:0,allAx:2,thDist:-2}
197. v2_wellord1(k1_wellord2(X0)) | ~v6_relat_2(k1_wellord2(X0)) | ~v3_ordinal1(X0) [subsumption resolution 196,65] {a:1,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
198. v2_wellord1(k1_wellord2(X0)) | ~v3_ordinal1(X0) [subsumption resolution 197,77] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
242. ~r1_tarski(X0,X1) | v2_wellord1(k1_wellord2(X0)) | ~v3_ordinal1(X1) [resolution 182,198] {a:2,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
246. v2_wellord1(k1_wellord2(sK1)) | ~v3_ordinal1(sK0) <- (2) [resolution 242,110] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
250. ~v3_ordinal1(sK0) <- (~1, 2) [subsumption resolution 246,105] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
251. $false <- (~1, 2, 3) [subsumption resolution 250,115] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
252. 1 | ~2 | ~3 [avatar contradiction clause 251]
253. $false [avatar sat refutation 106,111,116,252] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:80,thDist:-80}
% SZS output end Proof for wellord2__t8_wellord2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 29
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 40
% Generated clauses: 141
% Active clauses: 60
% Passive clauses: 91
% Final active clauses: 60
% Final passive clauses: 24
% Main loop iterations started: 67
% 
% >>> Simplifying Inferences
% Duplicate literals: 1
% Fw subsumption resolutions: 11
% 
% >>> Deletion Inferences
% Simple tautologies: 8
% Forward subsumptions: 23
% 
% >>> Generating Inferences
% Binary resolution: 59
% Unit resulting resolution: 1
% Forward superposition: 8
% Backward superposition: 10
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 15
% SAT solver unit clauses: 4
% SAT solver binary clauses: 5
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

