i: [1,47,0,2,1,1,0,0] t39_orders_1
a: [3,87,0,2,1,47]
s: 87
i: [1,48,0,3,1,1,0,0] t39_orders_1
a: [3,82,0,3,1,48]
s: 82
i: [1,49,0,4,1,1,0,0] t39_orders_1
a: [3,77,0,4,1,49]
s: 77
i: [1,65,0,4,1,0,0,255] fc6_relat_1
s: 65
i: [1,66,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 66
i: [1,50,0,6,3,0,0,1] d5_orders_1
s: 50
i: [1,51,0,6,3,0,0,1] d5_orders_1
s: 51
i: [1,52,0,6,3,0,0,1] d5_orders_1
s: 52
i: [1,53,0,6,3,0,0,1] d5_orders_1
s: 53
i: [1,68,0,6,2,0,0,1] dt_k2_wellord1
s: 68
i: [1,69,0,6,2,0,0,3] fc1_relat_1
s: 69
i: [1,67,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 67
i: [1,60,0,8,3,0,0,1] d8_orders_1
s: 60
d: [2,92,1,5,2,1,40,60,82]
d: [2,93,1,3,1,2,21,92,87]
a: [3,97,1,3,1,93]
s: 97
i: [1,61,0,8,3,0,0,1] d8_orders_1
s: 61
i: [1,62,0,8,3,0,0,1] d8_orders_1
s: 62
d: [2,99,1,5,2,1,40,61,82]
d: [2,100,1,3,1,2,21,99,87]
a: [3,104,1,3,1,100]
s: 104
i: [1,63,0,8,3,0,0,1] d8_orders_1
s: 63
d: [2,106,1,5,2,1,40,62,82]
d: [2,107,1,3,1,2,21,106,87]
a: [3,111,1,3,1,107]
s: 111
i: [1,56,0,9,3,0,0,2] l49_orders_1
s: 56
i: [1,57,0,9,3,0,0,2] l47_orders_1
s: 57
d: [2,113,1,5,2,1,40,63,82]
d: [2,114,1,3,1,2,21,113,87]
a: [3,118,1,3,1,114]
s: 118
i: [1,58,0,9,3,0,0,2] l46_orders_1
s: 58
d: [2,90,1,6,2,0,44,69,67]
s: 90
i: [1,59,0,9,3,0,0,2] l48_orders_1
s: 59
i: [1,72,0,10,4,0,0,2] fc18_relat_2
s: 72
d: [2,91,1,6,2,0,44,69,67]
s: 91
i: [1,73,0,10,4,0,0,2] fc19_relat_2
s: 73
d: [2,124,1,10,4,0,44,72,67]
s: 124
i: [1,55,0,11,2,0,0,255] d6_wellord1
s: 55
i: [1,54,0,12,6,0,0,1] d5_orders_1
s: 54
d: [2,136,1,14,5,0,44,72,55]
d: [2,142,1,12,4,0,18,136]
d: [2,147,1,8,3,0,21,142,65]
s: 147
i: [1,70,0,12,5,0,0,2] fc22_relat_2
s: 70
d: [2,131,1,13,2,0,44,55,66]
d: [2,144,1,9,1,0,21,131,65]
s: 144
i: [1,71,0,12,5,0,0,2] fc5_relat_2
s: 71
i: [1,64,0,17,6,0,0,1] d8_orders_1
s: 64
d: [2,134,1,13,2,0,44,66,55]
d: [2,145,1,9,1,0,21,134,65]
s: 145
d: [2,125,1,10,4,0,44,72,67]
s: 125
d: [2,158,2,13,3,0,44,59,144]
d: [2,163,2,9,2,0,21,158,65]
s: 163
d: [2,135,1,16,5,0,44,73,55]
d: [2,143,1,14,4,0,18,135]
d: [2,146,1,10,3,0,21,143,65]
s: 146
d: [2,132,1,11,2,0,44,55,67]
s: 132
d: [2,159,2,13,3,0,44,58,144]
d: [2,164,2,9,2,0,21,159,65]
s: 164
d: [2,133,1,11,2,0,44,55,67]
s: 133
d: [2,160,2,13,3,0,44,57,144]
d: [2,165,2,9,2,0,21,160,65]
s: 165
d: [2,138,1,11,2,0,44,67,55]
s: 138
d: [2,139,1,11,2,0,44,67,55]
s: 139
d: [2,161,2,13,3,0,44,56,144]
d: [2,166,2,9,2,0,21,161,65]
s: 166
d: [2,153,1,18,6,0,44,70,55]
d: [2,154,1,16,5,0,18,153]
d: [2,156,1,12,4,0,21,154,65]
s: 156
d: [2,178,2,19,3,0,44,132,55]
d: [2,194,2,15,2,0,21,178,65]
d: [2,195,2,11,1,0,21,194,65]
s: 195
d: [2,170,1,18,6,0,44,71,55]
d: [2,171,1,16,5,0,18,170]
d: [2,173,1,12,4,0,21,171,65]
s: 173
d: [2,148,1,25,7,0,40,54,56]
d: [2,149,1,21,6,0,21,148,68]
s: 149
d: [2,205,3,10,2,0,44,68,195]
d: [2,217,3,6,1,0,21,205,65]
s: 217
d: [2,188,2,19,3,0,44,55,132]
d: [2,198,2,15,2,0,21,188,65]
d: [2,199,2,11,1,0,21,198,65]
s: 199
d: [2,211,3,10,2,0,44,68,195]
d: [2,223,3,6,1,0,21,211,65]
s: 223
d: [2,183,2,18,6,0,44,70,132]
d: [2,193,2,16,5,0,18,183]
d: [2,196,2,12,4,0,21,193,65]
s: 196
d: [2,184,2,18,6,0,44,71,132]
d: [2,192,2,16,5,0,18,184]
d: [2,197,2,12,4,0,21,192,65]
s: 197
d: [2,200,3,14,3,0,44,147,195]
d: [2,212,3,10,2,0,21,200,65]
s: 212
d: [2,224,2,21,7,0,40,149,147]
d: [2,230,2,19,6,0,18,224]
s: 230
d: [2,206,3,14,3,0,44,147,195]
d: [2,218,3,10,2,0,21,206,65]
s: 218
d: [2,225,2,22,7,0,40,149,59]
d: [2,229,2,20,6,0,18,225]
s: 229
d: [2,226,2,25,6,0,44,149,144]
d: [2,231,2,21,5,0,21,226,65]
s: 231
d: [2,239,4,10,2,0,44,212,195]
s: 239
d: [2,201,3,15,3,0,44,59,195]
d: [2,213,3,11,2,0,21,201,65]
s: 213
d: [2,240,4,10,2,0,44,212,195]
s: 240
d: [2,202,3,15,3,0,44,58,195]
d: [2,214,3,11,2,0,21,202,65]
s: 214
d: [2,203,3,15,3,0,44,57,195]
d: [2,215,3,11,2,0,21,203,65]
s: 215
d: [2,204,3,15,3,0,44,56,195]
d: [2,216,3,11,2,0,21,204,65]
s: 216
d: [2,207,3,15,3,0,44,59,195]
d: [2,219,3,11,2,0,21,207,65]
s: 219
d: [2,208,3,15,3,0,44,58,195]
d: [2,220,3,11,2,0,21,208,65]
s: 220
d: [2,209,3,15,3,0,44,57,195]
d: [2,221,3,11,2,0,21,209,65]
s: 221
d: [2,210,3,15,3,0,44,56,195]
d: [2,222,3,11,2,0,21,210,65]
s: 222
d: [2,264,4,11,2,0,44,213,195]
s: 264
d: [2,242,3,20,7,0,40,230,57]
d: [2,246,3,18,6,0,18,242]
s: 246
d: [2,265,4,11,2,0,44,213,195]
s: 265
d: [2,250,3,21,7,0,40,229,57]
d: [2,254,3,19,6,0,18,250]
s: 254
d: [2,243,3,25,6,0,44,230,144]
d: [2,247,3,21,5,0,21,243,65]
s: 247
d: [2,268,4,11,2,0,44,214,195]
s: 268
d: [2,251,3,26,6,0,44,229,144]
d: [2,255,3,22,5,0,21,251,65]
s: 255
d: [2,269,4,11,2,0,44,214,195]
s: 269
d: [2,244,3,31,6,0,44,230,195]
d: [2,248,3,27,5,0,21,244,65]
s: 248
d: [2,245,3,31,6,0,44,230,195]
d: [2,249,3,27,5,0,21,245,65]
s: 249
d: [2,273,4,11,2,0,44,215,195]
s: 273
d: [2,252,3,32,6,0,44,229,195]
d: [2,256,3,28,5,0,21,252,65]
s: 256
d: [2,274,4,11,2,0,44,215,195]
s: 274
d: [2,253,3,32,6,0,44,229,195]
d: [2,257,3,28,5,0,21,253,65]
s: 257
d: [2,227,3,33,6,0,44,149,195]
d: [2,232,3,29,5,0,21,227,65]
s: 232
d: [2,279,4,11,2,0,44,216,195]
s: 279
d: [2,228,3,33,6,0,44,149,195]
d: [2,233,3,29,5,0,21,228,65]
s: 233
d: [2,280,4,11,2,0,44,216,195]
s: 280
d: [2,286,4,19,7,0,40,246,58]
d: [2,291,4,17,6,0,18,286]
s: 291
d: [2,296,4,20,7,0,40,254,58]
d: [2,301,4,18,6,0,18,296]
s: 301
d: [2,368,5,0,0,6,52,87,97,104,111,118,77,301]
e: 368
d: [2,369,5,14,5,1,40,301,77]
d: [2,378,5,11,4,2,21,369,118]
d: [2,379,5,9,3,3,21,378,87]
d: [2,380,5,6,2,4,21,379,111]
d: [2,381,5,3,1,5,21,380,104]
d: [2,382,5,0,0,6,21,381,97]
e: 382
f: 382,114,107,100,93,47,48,49
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t39_orders_1
% SZS output start Proof for orders_1__t39_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) => v3_orders_1(k2_wellord1(X0,X1)))) [input t39_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) => v3_orders_1(k2_wellord1(X0,X1)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)))) [input d5_orders_1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r3_orders_1(X0,X1) <=> (r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d8_orders_1]
7. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k2_wellord1(X0,X1))) [input dt_k2_wellord1]
17. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_relat_2(X0,X1) => v1_relat_2(k2_wellord1(X0,X1)))) [input l46_orders_1]
18. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r8_relat_2(X0,X1) => v8_relat_2(k2_wellord1(X0,X1)))) [input l47_orders_1]
19. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r4_relat_2(X0,X1) => v4_relat_2(k2_wellord1(X0,X1)))) [input l48_orders_1]
20. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r6_relat_2(X0,X1) => v6_relat_2(k2_wellord1(X0,X1)))) [input l49_orders_1]
22. ? [X0] : (? [X1] : (~v3_orders_1(k2_wellord1(X0,X1)) & r3_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
23. ! [X0] : ((v3_orders_1(X0) <=> (v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
25. ! [X0] : (! [X1] : (v6_relat_2(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 20]
26. ! [X0] : (! [X1] : (v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 18]
27. ! [X0] : (! [X1] : (v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 17]
28. ! [X0] : (! [X1] : (v4_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 19]
29. ! [X0] : (! [X1] : (r3_orders_1(X0,X1) <=> (r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 6]
30. ! [X0,X1] : (v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
40. ? [X0] : (? [X1] : (~v3_orders_1(k2_wellord1(X0,X1)) & r3_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (~v3_orders_1(k2_wellord1(sK0,X1)) & r3_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
41. ? [X1] : (~v3_orders_1(k2_wellord1(sK0,X1)) & r3_orders_1(sK0,X1)) => (~v3_orders_1(k2_wellord1(sK0,sK1)) & r3_orders_1(sK0,sK1)) [choice axiom]
42. (~v3_orders_1(k2_wellord1(sK0,sK1)) & r3_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 22,41,40]
43. ! [X0] : (((v3_orders_1(X0) | (~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 23]
44. ! [X0] : (((v3_orders_1(X0) | ~v6_relat_2(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v6_relat_2(X0) & v4_relat_2(X0) & v8_relat_2(X0) & v1_relat_2(X0)) | ~v3_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 43]
45. ! [X0] : (! [X1] : ((r3_orders_1(X0,X1) | (~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 29]
46. ! [X0] : (! [X1] : ((r3_orders_1(X0,X1) | ~r6_relat_2(X0,X1) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r6_relat_2(X0,X1) & r4_relat_2(X0,X1) & r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r3_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 45]
47. v1_relat_1(sK0) [cnf transformation 42] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
48. r3_orders_1(sK0,sK1) [cnf transformation 42] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
49. ~v3_orders_1(k2_wellord1(sK0,sK1)) [cnf transformation 42] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
54. ~v6_relat_2(X0) | v3_orders_1(X0) | ~v4_relat_2(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 44] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
56. v6_relat_2(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
57. v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 26] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
58. v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 27] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
59. v4_relat_2(k2_wellord1(X0,X1)) | ~r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 28] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
60. ~r3_orders_1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. ~r3_orders_1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
62. ~r3_orders_1(X0,X1) | r4_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. ~r3_orders_1(X0,X1) | r6_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 46] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
68. v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 30] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. 1 <=> v3_orders_1(k2_wellord1(sK0,sK1)) [avatar definition]
77. ~v3_orders_1(k2_wellord1(sK0,sK1)) <- (~1) [avatar component clause 75] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. ~1 [avatar split clause 49,75]
80. 2 <=> r3_orders_1(sK0,sK1) [avatar definition]
82. r3_orders_1(sK0,sK1) <- (2) [avatar component clause 80] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 2 [avatar split clause 48,80]
85. 3 <=> v1_relat_1(sK0) [avatar definition]
87. v1_relat_1(sK0) <- (3) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 3 [avatar split clause 47,85]
92. r1_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 60,82] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
93. r1_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 92,87] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
95. 4 <=> r1_relat_2(sK0,sK1) [avatar definition]
97. r1_relat_2(sK0,sK1) <- (4) [avatar component clause 95] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
98. 4 | ~2 | ~3 [avatar split clause 93,85,80,95]
99. r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 61,82] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
100. r8_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 99,87] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
102. 5 <=> r8_relat_2(sK0,sK1) [avatar definition]
104. r8_relat_2(sK0,sK1) <- (5) [avatar component clause 102] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
105. 5 | ~2 | ~3 [avatar split clause 100,85,80,102]
106. r4_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 62,82] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
107. r4_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 106,87] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
109. 6 <=> r4_relat_2(sK0,sK1) [avatar definition]
111. r4_relat_2(sK0,sK1) <- (6) [avatar component clause 109] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
112. 6 | ~2 | ~3 [avatar split clause 107,85,80,109]
113. r6_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 63,82] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
114. r6_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 113,87] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
116. 7 <=> r6_relat_2(sK0,sK1) [avatar definition]
118. r6_relat_2(sK0,sK1) <- (7) [avatar component clause 116] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
119. 7 | ~2 | ~3 [avatar split clause 114,85,80,116]
148. v3_orders_1(k2_wellord1(X0,X1)) | ~v4_relat_2(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 54,56] {a:1,w:25,wCS:10000,sine:1,thAx:0,allAx:2,thDist:-2}
149. ~v4_relat_2(k2_wellord1(X0,X1)) | v3_orders_1(k2_wellord1(X0,X1)) | ~v8_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) [subsumption resolution 148,68] {a:1,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
225. v3_orders_1(k2_wellord1(X2,X3)) | ~v8_relat_2(k2_wellord1(X2,X3)) | ~v1_relat_2(k2_wellord1(X2,X3)) | ~r6_relat_2(X2,X3) | ~v1_relat_1(X2) | ~r4_relat_2(X2,X3) | ~v1_relat_1(X2) [resolution 149,59] {a:2,w:22,wCS:8800,sine:1,thAx:0,allAx:3,thDist:-3}
229. ~v8_relat_2(k2_wellord1(X2,X3)) | v3_orders_1(k2_wellord1(X2,X3)) | ~v1_relat_2(k2_wellord1(X2,X3)) | ~r6_relat_2(X2,X3) | ~v1_relat_1(X2) | ~r4_relat_2(X2,X3) [duplicate literal removal 225] {a:2,w:20,wCS:8000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
250. v3_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 229,57] {a:3,w:21,wCS:8400,sine:1,thAx:0,allAx:4,thDist:-4}
254. ~v1_relat_2(k2_wellord1(X0,X1)) | v3_orders_1(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) [duplicate literal removal 250] {a:3,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
296. v3_orders_1(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 254,58] {a:4,w:20,wCS:8000,sine:1,thAx:0,allAx:5,thDist:-5}
301. v3_orders_1(k2_wellord1(X0,X1)) | ~r6_relat_2(X0,X1) | ~v1_relat_1(X0) | ~r4_relat_2(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) [duplicate literal removal 296] {a:4,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
369. ~r6_relat_2(sK0,sK1) | ~v1_relat_1(sK0) | ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1) [resolution 301,77] {a:5,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
378. ~v1_relat_1(sK0) | ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 7) [subsumption resolution 369,118] {a:5,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
379. ~r4_relat_2(sK0,sK1) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3, 7) [subsumption resolution 378,87] {a:5,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
380. ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3, 6, 7) [subsumption resolution 379,111] {a:5,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
381. ~r1_relat_2(sK0,sK1) <- (~1, 3, 5, 6, 7) [subsumption resolution 380,104] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
382. $false <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 381,97] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
383. 1 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar contradiction clause 382]
387. $false [avatar sat refutation 78,83,88,98,105,112,119,383] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:90,thDist:-90}
% SZS output end Proof for orders_1__t39_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 27
% Generated clauses: 306
% Active clauses: 66
% Passive clauses: 127
% Final active clauses: 66
% Final passive clauses: 36
% Main loop iterations started: 91
% 
% >>> Simplifying Inferences
% Duplicate literals: 22
% Fw subsumption resolutions: 77
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Forward subsumptions: 84
% 
% >>> Generating Inferences
% Binary resolution: 48
% Unit resulting resolution: 1
% Forward superposition: 88
% Backward superposition: 37
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 12
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% term sharing: 0.001 s
% forward subsumption: 0.001 s
% superposition: 0.001 s ( own 0.0000 s ) 
% other: 0.004 s ( own 0.002 s ) 

