i: [1,28,0,2,1,1,0,0] t228_xreal_1
a: [3,52,0,2,1,28]
s: 52
i: [1,27,0,2,1,1,0,0] t228_xreal_1
a: [3,57,0,2,1,27]
s: 57
i: [1,26,0,2,1,1,0,0] t228_xreal_1
a: [3,62,0,2,1,26]
s: 62
i: [1,31,0,3,1,1,0,0] t228_xreal_1
a: [3,42,0,3,1,31]
s: 42
i: [1,30,0,3,1,1,0,0] t228_xreal_1
a: [3,47,0,3,1,30]
s: 47
i: [1,38,0,2,1,0,0,1] rc1_xxreal_0
a: [3,74,0,2,1,38]
s: 74
i: [1,29,0,11,4,1,0,0] t228_xreal_1
s: 29
i: [1,36,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,69,0,5,2,36]
s: 69
d: [2,76,1,8,3,1,40,29,42]
d: [2,77,1,5,2,2,21,76,47]
d: [2,78,1,3,1,3,21,77,62]
a: [3,82,1,3,1,78]
s: 82
i: [1,37,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 37
i: [1,32,0,11,4,0,0,1] t227_xreal_1
s: 32
d: [2,84,1,7,3,1,40,37,42]
d: [2,88,1,5,2,2,21,84,52]
d: [2,89,1,3,1,3,21,88,62]
a: [3,93,1,3,1,89]
s: 93
i: [1,33,0,12,4,0,0,1] t227_xreal_1
s: 33
d: [2,85,1,7,3,1,40,37,47]
d: [2,95,1,5,2,2,21,85,57]
d: [2,96,1,3,1,3,21,95,62]
a: [3,100,1,3,1,96]
s: 100
i: [1,34,0,12,4,0,0,1] t227_xreal_1
s: 34
i: [1,35,0,15,6,0,0,1] t2_xxreal_0
s: 35
d: [2,114,2,12,5,1,40,35,100]
d: [2,121,2,10,4,2,21,114,57]
d: [2,122,2,8,3,3,21,121,62]
s: 122
d: [2,110,1,20,8,0,40,35,29]
d: [2,119,1,18,7,0,18,110]
d: [2,120,1,16,6,1,21,119,52]
s: 120
d: [2,115,2,12,5,1,40,35,93]
d: [2,123,2,10,4,2,21,115,52]
d: [2,124,2,8,3,3,21,123,62]
s: 124
d: [2,105,1,21,6,0,40,34,29]
d: [2,108,1,19,5,1,21,105,52]
s: 108
d: [2,102,1,18,6,0,40,33,37]
d: [2,103,1,16,5,0,18,102]
d: [2,104,1,12,4,0,21,103,32]
s: 104
d: [2,149,2,21,7,1,40,108,33]
d: [2,152,2,16,5,1,18,149]
d: [2,155,2,14,4,2,21,152,62]
d: [2,156,2,11,3,3,21,155,42]
d: [2,157,2,9,2,3,21,156,52]
a: [3,161,2,4,1,157]
s: 161
d: [2,173,3,0,0,4,52,62,52,42,161,32]
e: 173
d: [2,174,3,7,3,1,40,161,32]
d: [2,176,3,4,2,2,21,174,42]
d: [2,177,3,2,1,3,21,176,52]
d: [2,178,3,0,0,4,21,177,62]
e: 178
r: 161
d: [2,106,1,18,6,0,40,34,37]
d: [2,107,1,16,5,0,18,106]
d: [2,109,1,12,4,0,21,107,32]
s: 109
a: [3,160,2,4,1,157]
s: 160
d: [2,112,1,19,8,0,40,35,37]
d: [2,117,1,15,6,0,18,112]
s: 117
a: [3,165,2,5,1,157]
s: 165
d: [2,197,2,12,5,1,40,117,100]
d: [2,210,2,10,4,2,21,197,57]
d: [2,211,2,8,3,3,21,210,62]
s: 211
d: [2,198,2,12,5,1,40,117,93]
d: [2,212,2,10,4,2,21,198,52]
d: [2,213,2,8,3,3,21,212,62]
s: 213
d: [2,128,3,12,5,3,40,122,37]
d: [2,129,3,10,4,3,18,128]
d: [2,131,3,8,3,3,21,129,62]
s: 131
d: [2,135,2,20,8,1,40,120,37]
d: [2,136,2,18,7,1,18,135]
d: [2,139,2,16,6,1,21,136,52]
s: 139
d: [2,191,2,20,8,0,40,117,29]
d: [2,206,2,18,7,0,18,191]
d: [2,207,2,16,6,1,21,206,52]
s: 207
d: [2,143,3,12,5,3,40,124,37]
d: [2,144,3,10,4,3,18,143]
d: [2,146,3,8,3,3,21,144,62]
s: 146
d: [2,184,2,24,8,1,40,109,120]
d: [2,190,2,22,7,1,21,184,52]
s: 190
d: [2,215,3,16,5,1,40,165,117]
d: [2,217,3,12,4,2,21,215,160]
d: [2,218,3,10,3,3,21,217,57]
s: 218
d: [2,150,2,30,9,1,40,108,35]
d: [2,151,2,26,8,1,18,150]
d: [2,167,2,24,7,2,21,151,62]
s: 167
d: [2,196,2,30,9,1,40,117,108]
d: [2,201,2,26,8,1,18,196]
d: [2,209,2,24,7,2,21,201,62]
s: 209
d: [2,278,4,12,5,3,40,218,34]
d: [2,284,4,10,4,3,18,278]
d: [2,285,4,8,3,4,21,284,52]
d: [2,286,4,5,2,5,21,285,42]
d: [2,287,4,3,1,6,21,286,62]
a: [3,291,4,3,1,287]
s: 291
d: [2,193,2,19,8,0,40,117,37]
d: [2,204,2,15,6,0,18,193]
s: 204
d: [2,336,5,13,5,2,40,291,120]
d: [2,339,5,10,4,2,18,336]
d: [2,340,5,8,3,3,21,339,57]
s: 340
r: 108
r: 29
r: 120
r: 139
r: 190
d: [2,169,2,23,8,0,40,104,35]
d: [2,170,2,21,7,0,18,169]
d: [2,172,2,17,6,0,21,170,32]
s: 172
d: [2,181,2,23,8,0,40,109,35]
d: [2,185,2,21,7,0,18,181]
d: [2,187,2,17,6,0,21,185,32]
s: 187
d: [2,400,4,17,7,3,40,187,218]
d: [2,416,4,15,6,3,18,400]
d: [2,418,4,13,5,4,21,416,62]
d: [2,419,4,10,4,5,21,418,42]
d: [2,420,4,8,3,6,21,419,52]
s: 420
d: [2,194,2,23,8,0,40,117,109]
d: [2,203,2,21,7,0,18,194]
d: [2,208,2,17,6,0,21,203,32]
s: 208
d: [2,337,5,12,5,1,40,291,117]
d: [2,341,5,10,4,2,21,337,52]
d: [2,342,5,8,3,3,21,341,57]
s: 342
d: [2,199,2,23,8,0,40,117,104]
d: [2,200,2,21,7,0,18,199]
d: [2,214,2,17,6,0,21,200,32]
s: 214
d: [2,216,3,16,5,1,40,165,35]
d: [2,219,3,12,4,2,21,216,160]
d: [2,220,3,10,3,3,21,219,57]
s: 220
d: [2,486,4,12,5,3,40,220,34]
d: [2,494,4,10,4,3,18,486]
d: [2,496,4,8,3,4,21,494,52]
d: [2,497,4,5,2,5,21,496,42]
d: [2,498,4,3,1,6,21,497,62]
a: [3,502,4,3,1,498]
s: 502
d: [2,221,3,17,7,3,40,211,117]
d: [2,230,3,15,6,3,18,221]
d: [2,231,3,13,5,3,21,230,62]
s: 231
d: [2,487,4,17,7,3,40,220,187]
d: [2,493,4,15,6,3,18,487]
d: [2,504,4,13,5,4,21,493,62]
d: [2,505,4,10,4,5,21,504,42]
d: [2,506,4,8,3,6,21,505,52]
s: 506
d: [2,222,3,17,7,3,40,211,35]
d: [2,229,3,15,6,3,18,222]
d: [2,232,3,13,5,3,21,229,62]
s: 232
d: [2,234,3,17,7,3,40,213,117]
d: [2,243,3,15,6,3,18,234]
d: [2,244,3,13,5,3,21,243,62]
s: 244
d: [2,338,5,12,5,1,40,291,35]
d: [2,343,5,10,4,2,21,338,52]
d: [2,344,5,8,3,3,21,343,57]
s: 344
d: [2,235,3,17,7,3,40,213,35]
d: [2,242,3,15,6,3,18,235]
d: [2,245,3,13,5,3,21,242,62]
s: 245
d: [2,350,3,17,7,3,40,204,124]
d: [2,363,3,15,6,3,18,350]
d: [2,368,3,13,5,3,21,363,62]
s: 368
d: [2,351,3,17,7,3,40,204,122]
d: [2,362,3,15,6,3,18,351]
d: [2,369,3,13,5,3,21,362,62]
s: 369
d: [2,182,3,16,5,3,40,109,124]
d: [2,188,3,14,4,3,21,182,62]
s: 188
d: [2,431,5,17,7,6,40,420,117]
d: [2,441,5,15,6,6,18,431]
d: [2,442,5,13,5,6,21,441,57]
s: 442
d: [2,183,3,16,5,3,40,109,122]
d: [2,189,3,14,4,3,21,183,62]
s: 189
d: [2,432,5,17,7,6,40,420,35]
d: [2,440,5,15,6,6,18,432]
d: [2,443,5,13,5,6,21,440,57]
s: 443
d: [2,226,3,16,5,3,40,211,34]
d: [2,233,3,14,4,3,21,226,62]
s: 233
d: [2,239,3,16,5,3,40,213,34]
d: [2,246,3,14,4,3,21,239,62]
s: 246
d: [2,539,5,17,7,6,40,506,204]
d: [2,546,5,15,6,6,18,539]
d: [2,550,5,13,5,6,21,546,57]
s: 550
d: [2,403,3,21,7,3,40,187,213]
d: [2,421,3,19,6,3,21,403,62]
d: [2,422,3,15,5,3,21,421,32]
s: 422
d: [2,372,6,17,7,3,40,340,117]
d: [2,381,6,15,6,3,18,372]
d: [2,382,6,13,5,3,21,381,57]
s: 382
r: 209
d: [2,404,3,21,7,3,40,187,211]
d: [2,423,3,19,6,3,21,404,62]
d: [2,424,3,15,5,3,21,423,32]
s: 424
d: [2,450,3,21,7,3,40,208,124]
d: [2,461,3,19,6,3,21,450,62]
d: [2,462,3,15,5,3,21,461,32]
s: 462
d: [2,373,6,17,7,3,40,340,35]
d: [2,380,6,15,6,3,18,373]
d: [2,383,6,13,5,3,21,380,57]
s: 383
d: [2,451,3,21,7,3,40,208,122]
d: [2,463,3,19,6,3,21,451,62]
d: [2,464,3,15,5,3,21,463,32]
s: 464
d: [2,436,5,16,5,6,40,420,34]
d: [2,444,5,14,4,6,21,436,57]
s: 444
d: [2,296,3,26,9,2,40,167,104]
d: [2,316,3,19,6,2,18,296]
d: [2,317,3,17,5,2,21,316,52]
s: 317
d: [2,318,3,26,9,2,40,209,33]
d: [2,333,3,19,6,2,18,318]
d: [2,334,3,17,5,2,21,333,52]
s: 334
d: [2,541,5,16,5,6,40,506,109]
d: [2,551,5,14,4,6,21,541,57]
s: 551
d: [2,148,3,23,6,4,40,108,122]
d: [2,153,3,19,5,4,18,148]
s: 153
d: [2,377,6,16,5,3,40,340,34]
d: [2,384,6,14,4,3,21,377,57]
s: 384
d: [2,147,3,23,6,3,40,108,124]
d: [2,154,3,19,5,3,18,147]
s: 154
d: [2,247,3,25,10,1,40,139,117]
d: [2,259,3,23,9,1,18,247]
d: [2,260,3,21,8,1,21,259,52]
s: 260
d: [2,800,7,16,6,3,40,384,33]
d: [2,805,7,11,4,3,18,800]
d: [2,810,7,8,3,4,21,805,47]
d: [2,811,7,6,2,5,21,810,62]
d: [2,812,7,4,1,5,21,811,57]
a: [3,816,7,4,1,812]
s: 816
d: [2,820,8,0,0,4,52,62,57,47,816,32]
e: 820
d: [2,821,8,7,3,1,40,816,32]
d: [2,823,8,4,2,2,21,821,47]
d: [2,824,8,2,1,3,21,823,57]
d: [2,825,8,0,0,4,21,824,62]
e: 825
f: 825,812,287,178,157,26,27,28,30,31
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xreal_1__t228_xreal_1
% SZS output start Proof for xreal_1__t228_xreal_1
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (! [X3] : (v1_xxreal_0(X3) => ~(~r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X3) & ~r1_xxreal_0(X3,X0))) => (r1_xxreal_0(X2,X0) | r1_xxreal_0(X1,X0)))))) [input t228_xreal_1]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => (! [X3] : (v1_xxreal_0(X3) => ~(~r1_xxreal_0(X2,X3) & ~r1_xxreal_0(X1,X3) & ~r1_xxreal_0(X3,X0))) => (r1_xxreal_0(X2,X0) | r1_xxreal_0(X1,X0)))))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
6. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ~(! [X2] : (v1_xxreal_0(X2) => ~(~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0))) & ~r1_xxreal_0(X1,X0)))) [input t227_xreal_1]
7. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ! [X2] : (v1_xxreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input t2_xxreal_0]
8. ? [X0] : (? [X1] : (? [X2] : (((~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X1,X0)) & ! [X3] : ((r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0)) | ~v1_xxreal_0(X3))) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
9. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X1,X0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 8]
10. ! [X0] : (! [X1] : ((? [X2] : ((~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0)) & v1_xxreal_0(X2)) | r1_xxreal_0(X1,X0)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 6]
11. ! [X0] : (! [X1] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0) & v1_xxreal_0(X2)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 10]
12. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 7]
13. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2)) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 12]
16. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
17. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 16]
18. ? [X0] : (? [X1] : (? [X2] : (~r1_xxreal_0(X2,X0) & ~r1_xxreal_0(X1,X0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,X0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (? [X2] : (~r1_xxreal_0(X2,sK0) & ~r1_xxreal_0(X1,sK0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
19. ? [X1] : (? [X2] : (~r1_xxreal_0(X2,sK0) & ~r1_xxreal_0(X1,sK0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(X1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(X1)) => (? [X2] : (~r1_xxreal_0(X2,sK0) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) & v1_xxreal_0(sK1)) [choice axiom]
20. ? [X2] : (~r1_xxreal_0(X2,sK0) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(X2,X3) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(X2)) => (~r1_xxreal_0(sK2,sK0) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(sK2,X3) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) [choice axiom]
21. ((~r1_xxreal_0(sK2,sK0) & ~r1_xxreal_0(sK1,sK0) & ! [X3] : (r1_xxreal_0(sK2,X3) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3)) & v1_xxreal_0(sK2)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 9,20,19,18]
22. ! [X1,X0] : (? [X2] : (~r1_xxreal_0(X1,X2) & ~r1_xxreal_0(X2,X0) & v1_xxreal_0(X2)) => (~r1_xxreal_0(X1,sK3(X0,X1)) & ~r1_xxreal_0(sK3(X0,X1),X0) & v1_xxreal_0(sK3(X0,X1)))) [choice axiom]
23. ! [X0] : (! [X1] : ((~r1_xxreal_0(X1,sK3(X0,X1)) & ~r1_xxreal_0(sK3(X0,X1),X0) & v1_xxreal_0(sK3(X0,X1))) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [skolemisation 11,22]
26. v1_xxreal_0(sK0) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
27. v1_xxreal_0(sK1) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
28. v1_xxreal_0(sK2) [cnf transformation 21] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
29. r1_xxreal_0(sK2,X3) | r1_xxreal_0(sK1,X3) | r1_xxreal_0(X3,sK0) | ~v1_xxreal_0(X3) [cnf transformation 21] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
30. ~r1_xxreal_0(sK1,sK0) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
31. ~r1_xxreal_0(sK2,sK0) [cnf transformation 21] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
32. v1_xxreal_0(sK3(X0,X1)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
33. ~r1_xxreal_0(sK3(X0,X1),X0) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
34. ~r1_xxreal_0(X1,sK3(X0,X1)) | r1_xxreal_0(X1,X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
35. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 13] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 17] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
40. 1 <=> r1_xxreal_0(sK2,sK0) [avatar definition]
42. ~r1_xxreal_0(sK2,sK0) <- (~1) [avatar component clause 40] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
43. ~1 [avatar split clause 31,40]
45. 2 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
47. ~r1_xxreal_0(sK1,sK0) <- (~2) [avatar component clause 45] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
48. ~2 [avatar split clause 30,45]
50. 3 <=> v1_xxreal_0(sK2) [avatar definition]
52. v1_xxreal_0(sK2) <- (3) [avatar component clause 50] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
53. 3 [avatar split clause 28,50]
55. 4 <=> v1_xxreal_0(sK1) [avatar definition]
57. v1_xxreal_0(sK1) <- (4) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
58. 4 [avatar split clause 27,55]
60. 5 <=> v1_xxreal_0(sK0) [avatar definition]
62. v1_xxreal_0(sK0) <- (5) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
63. 5 [avatar split clause 26,60]
105. r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,sK3(X0,sK2)) | r1_xxreal_0(sK3(X0,sK2),sK0) | ~v1_xxreal_0(sK3(X0,sK2)) [resolution 34,29] {a:1,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
108. r1_xxreal_0(sK3(X0,sK2),sK0) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,sK3(X0,sK2)) | r1_xxreal_0(sK2,X0) | ~v1_xxreal_0(sK3(X0,sK2)) <- (3) [subsumption resolution 105,52] {a:1,w:19,wCS:1900,nSel:1,goal:1,sine:0,thAx:0,allAx:6,thDist:-6}
110. r1_xxreal_0(X0,X1) | ~r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X1) | r1_xxreal_0(X1,sK0) | ~v1_xxreal_0(X1) [resolution 35,29] {a:1,w:20,wCS:2000,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
112. r1_xxreal_0(X4,X5) | ~r1_xxreal_0(X4,X6) | ~v1_xxreal_0(X5) | ~v1_xxreal_0(X6) | ~v1_xxreal_0(X4) | r1_xxreal_0(X5,X6) | ~v1_xxreal_0(X6) | ~v1_xxreal_0(X5) [resolution 35,37] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
117. ~r1_xxreal_0(X4,X6) | r1_xxreal_0(X4,X5) | ~v1_xxreal_0(X5) | ~v1_xxreal_0(X6) | ~v1_xxreal_0(X4) | r1_xxreal_0(X5,X6) [duplicate literal removal 112] {a:1,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
119. r1_xxreal_0(X0,X1) | ~r1_xxreal_0(X0,sK2) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X1) | r1_xxreal_0(X1,sK0) [duplicate literal removal 110] {a:1,w:18,wCS:1800,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
120. ~r1_xxreal_0(X0,sK2) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X1) | r1_xxreal_0(X1,sK0) <- (3) [subsumption resolution 119,52] {a:1,w:16,wCS:1600,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
149. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK1,sK3(sK0,sK2)) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK3(sK0,sK2)) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (3) [resolution 108,33] {a:2,w:21,wCS:2100,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
152. ~v1_xxreal_0(sK0) | r1_xxreal_0(sK1,sK3(sK0,sK2)) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK3(sK0,sK2)) | ~v1_xxreal_0(sK2) <- (3) [duplicate literal removal 149] {a:2,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
155. r1_xxreal_0(sK1,sK3(sK0,sK2)) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK3(sK0,sK2)) | ~v1_xxreal_0(sK2) <- (3, 5) [subsumption resolution 152,62] {a:2,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
156. r1_xxreal_0(sK1,sK3(sK0,sK2)) | ~v1_xxreal_0(sK3(sK0,sK2)) | ~v1_xxreal_0(sK2) <- (~1, 3, 5) [subsumption resolution 155,42] {a:2,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
157. r1_xxreal_0(sK1,sK3(sK0,sK2)) | ~v1_xxreal_0(sK3(sK0,sK2)) <- (~1, 3, 5) [subsumption resolution 156,52] {a:2,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
159. 12 <=> v1_xxreal_0(sK3(sK0,sK2)) [avatar definition]
160. v1_xxreal_0(sK3(sK0,sK2)) <- (12) [avatar component clause 159] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
161. ~v1_xxreal_0(sK3(sK0,sK2)) <- (~12) [avatar component clause 159] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
163. 13 <=> r1_xxreal_0(sK1,sK3(sK0,sK2)) [avatar definition]
165. r1_xxreal_0(sK1,sK3(sK0,sK2)) <- (13) [avatar component clause 163] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
166. ~12 | 13 | 1 | ~3 | ~5 [avatar split clause 157,60,50,40,163,159]
174. r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~12) [resolution 161,32] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
176. ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (~1, ~12) [subsumption resolution 174,42] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
177. ~v1_xxreal_0(sK0) <- (~1, 3, ~12) [subsumption resolution 176,52] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
178. $false <- (~1, 3, 5, ~12) [subsumption resolution 177,62] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
179. 1 | ~3 | ~5 | 12 [avatar contradiction clause 178]
215. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK3(sK0,sK2)) | ~v1_xxreal_0(sK1) | r1_xxreal_0(X0,sK3(sK0,sK2)) <- (13) [resolution 165,117] {a:3,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
217. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(X0,sK3(sK0,sK2)) <- (12, 13) [subsumption resolution 215,160] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
218. r1_xxreal_0(X0,sK3(sK0,sK2)) | ~v1_xxreal_0(X0) | r1_xxreal_0(sK1,X0) <- (4, 12, 13) [subsumption resolution 217,57] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
278. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK2) | ~v1_xxreal_0(sK0) <- (4, 12, 13) [resolution 218,34] {a:4,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
284. ~v1_xxreal_0(sK2) | r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (4, 12, 13) [duplicate literal removal 278] {a:4,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
285. r1_xxreal_0(sK1,sK2) | r1_xxreal_0(sK2,sK0) | ~v1_xxreal_0(sK0) <- (3, 4, 12, 13) [subsumption resolution 284,52] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
286. r1_xxreal_0(sK1,sK2) | ~v1_xxreal_0(sK0) <- (~1, 3, 4, 12, 13) [subsumption resolution 285,42] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
287. r1_xxreal_0(sK1,sK2) <- (~1, 3, 4, 5, 12, 13) [subsumption resolution 286,62] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:12,thDist:-12}
289. 14 <=> r1_xxreal_0(sK1,sK2) [avatar definition]
291. r1_xxreal_0(sK1,sK2) <- (14) [avatar component clause 289] {a:4,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:12,thDist:-12}
292. 14 | 1 | ~3 | ~4 | ~5 | ~12 | ~13 [avatar split clause 287,163,159,60,55,50,40,289]
336. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(sK1,X0) | r1_xxreal_0(X0,sK0) <- (3, 14) [resolution 291,120] {a:5,w:13,wCS:1300,goal:1,sine:0,thAx:0,allAx:17,thDist:-17}
339. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(X0) | ~v1_xxreal_0(sK1) | r1_xxreal_0(X0,sK0) <- (3, 14) [duplicate literal removal 336] {a:5,w:10,wCS:1000,goal:1,sine:0,thAx:0,allAx:17,thDist:-17}
340. r1_xxreal_0(sK1,X0) | ~v1_xxreal_0(X0) | r1_xxreal_0(X0,sK0) <- (3, 4, 14) [subsumption resolution 339,57] {a:5,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:17,thDist:-17}
377. ~v1_xxreal_0(sK3(X4,sK1)) | r1_xxreal_0(sK3(X4,sK1),sK0) | r1_xxreal_0(sK1,X4) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(X4) <- (3, 4, 14) [resolution 340,34] {a:6,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:19,thDist:-19}
384. r1_xxreal_0(sK3(X4,sK1),sK0) | ~v1_xxreal_0(sK3(X4,sK1)) | r1_xxreal_0(sK1,X4) | ~v1_xxreal_0(X4) <- (3, 4, 14) [subsumption resolution 377,57] {a:6,w:14,wCS:1400,nSel:1,goal:1,sine:0,thAx:0,allAx:19,thDist:-19}
800. ~v1_xxreal_0(sK3(sK0,sK1)) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (3, 4, 14) [resolution 384,33] {a:7,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
805. ~v1_xxreal_0(sK3(sK0,sK1)) | r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (3, 4, 14) [duplicate literal removal 800] {a:7,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
810. ~v1_xxreal_0(sK3(sK0,sK1)) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (~2, 3, 4, 14) [subsumption resolution 805,47] {a:7,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
811. ~v1_xxreal_0(sK3(sK0,sK1)) | ~v1_xxreal_0(sK1) <- (~2, 3, 4, 5, 14) [subsumption resolution 810,62] {a:7,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
812. ~v1_xxreal_0(sK3(sK0,sK1)) <- (~2, 3, 4, 5, 14) [subsumption resolution 811,57] {a:7,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:21,thDist:-21}
814. 16 <=> v1_xxreal_0(sK3(sK0,sK1)) [avatar definition]
816. ~v1_xxreal_0(sK3(sK0,sK1)) <- (~16) [avatar component clause 814] {a:7,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:21,thDist:-21}
817. ~16 | 2 | ~3 | ~4 | ~5 | ~14 [avatar split clause 812,289,60,55,50,45,814]
821. r1_xxreal_0(sK1,sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~16) [resolution 816,32] {a:8,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
823. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, ~16) [subsumption resolution 821,47] {a:8,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
824. ~v1_xxreal_0(sK0) <- (~2, 4, ~16) [subsumption resolution 823,57] {a:8,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
825. $false <- (~2, 4, 5, ~16) [subsumption resolution 824,62] {a:8,w:0,goal:1,sine:1,thAx:0,allAx:23,thDist:-23}
826. 2 | ~4 | ~5 | 16 [avatar contradiction clause 825]
827. $false [avatar sat refutation 43,48,53,58,63,166,179,292,817,826] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:228,thDist:-228}
% SZS output end Proof for xreal_1__t228_xreal_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 6
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 737
% Active clauses: 61
% Passive clauses: 168
% Final active clauses: 54
% Final passive clauses: 91
% Main loop iterations started: 77
% 
% >>> Simplifying Inferences
% Duplicate literals: 272
% Fw subsumption resolutions: 186
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 177
% Backward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 306
% Unit resulting resolution: 2
% Factoring: 13
% 
% >>> AVATAR
% Split clauses: 2
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 20
% SAT solver unit clauses: 6
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 5245
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% backward subsumption index maintenance: 0.001 s
% forward subsumption: 0.003 s
% backward subsumption: 0.001 s
% resolution: 0.001 s
% other: 0.008 s ( own 0.001 s ) 

