i: [1,50,0,11,3,1,0,0] t32_zfmisc_1
a: [3,83,0,3,1,50]
s: 83
a: [3,87,0,3,1,50]
s: 87
a: [3,80,0,5,1,50]
s: 80
i: [1,51,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,95,0,2,1,51]
s: 95
i: [1,52,0,2,1,0,0,255] fc1_xboole_0
a: [3,100,0,2,1,52]
s: 100
i: [1,70,0,2,1,0,0,2] rc2_xboole_0
a: [3,110,0,2,1,70]
s: 110
i: [1,71,0,2,1,0,0,2] rc1_xboole_0
a: [3,115,0,2,1,71]
s: 115
i: [1,66,0,3,1,0,0,1] reflexivity_r1_tarski
s: 66
i: [1,53,0,3,1,0,0,255] d2_xboole_0
a: [3,105,0,3,1,53]
s: 105
r: 95
i: [1,67,0,4,1,0,0,1] fc3_xboole_0
s: 67
i: [1,54,0,5,2,0,0,255] t6_boole
s: 54
i: [1,59,0,5,2,0,0,1] t7_boole
s: 59
i: [1,73,0,5,1,0,0,1] d2_tarski
s: 73
d: [2,120,1,2,1,1,40,59,83]
a: [3,125,1,2,1,120]
s: 125
i: [1,75,0,5,1,0,0,1] d2_tarski
s: 75
i: [1,69,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 69
d: [2,130,1,3,1,1,40,69,83]
a: [3,137,1,3,1,130]
s: 137
i: [1,58,0,7,3,0,0,2] t8_boole
s: 58
d: [2,131,1,3,1,1,40,69,87]
a: [3,142,1,3,1,131]
s: 142
i: [1,68,0,7,1,0,0,1] commutativity_k2_tarski
s: 68
i: [1,56,0,8,2,0,0,1] d3_tarski
s: 56
d: [2,132,1,5,1,0,40,69,73]
s: 132
i: [1,57,0,8,2,0,0,1] d3_tarski
s: 57
d: [2,133,1,5,1,0,40,69,75]
s: 133
i: [1,55,0,9,3,0,0,1] d3_tarski
s: 55
i: [1,76,0,11,3,0,0,1] d2_tarski
s: 76
d: [2,158,1,6,2,1,40,55,83]
s: 158
i: [1,64,0,17,3,0,0,1] d2_tarski
s: 64
d: [2,159,1,6,2,1,40,55,87]
s: 159
i: [1,65,0,17,3,0,0,1] d2_tarski
s: 65
i: [1,63,0,23,4,0,0,1] d2_tarski
s: 63
d: [2,151,1,5,2,0,40,56,59]
s: 151
d: [2,118,1,7,3,0,44,54,54]
s: 118
d: [2,152,2,5,1,0,44,132,68]
s: 152
d: [2,150,1,8,2,0,40,56,69]
s: 150
d: [2,160,1,8,2,0,40,55,73]
s: 160
d: [2,153,2,5,1,0,44,132,68]
s: 153
d: [2,161,1,8,2,0,40,55,75]
s: 161
d: [2,179,2,8,2,0,44,160,68]
s: 179
d: [2,162,1,11,3,0,40,55,56]
s: 162
d: [2,165,1,19,3,0,40,76,56]
s: 165
d: [2,180,2,8,2,0,44,160,68]
s: 180
d: [2,172,1,19,4,0,40,63,59]
s: 172
d: [2,188,2,8,3,0,40,162,59]
s: 188
d: [2,171,1,23,4,0,40,63,69]
s: 171
d: [2,170,1,26,5,0,40,63,55]
s: 170
d: [2,187,2,11,3,0,40,162,69]
s: 187
d: [2,173,1,39,5,0,40,63,76]
s: 173
d: [2,186,2,14,4,0,40,162,55]
s: 186
d: [2,190,2,15,3,0,40,162,150]
s: 190
d: [2,198,2,15,3,0,46,165]
s: 198
d: [2,247,3,12,4,0,40,186,57]
d: [2,257,3,9,3,0,18,247]
s: 257
d: [2,196,2,20,4,0,44,57,165]
d: [2,200,2,15,3,0,18,196]
s: 200
d: [2,250,3,11,4,0,40,186,59]
s: 250
d: [2,206,2,16,4,0,46,172]
s: 206
d: [2,189,2,18,4,0,40,162,76]
s: 189
d: [2,269,3,16,4,0,44,57,200]
d: [2,280,3,11,3,0,18,269]
s: 280
d: [2,290,4,0,0,3,52,87,83,80,280]
e: 290
d: [2,291,4,6,2,1,40,280,80]
d: [2,302,4,3,1,2,21,291,83]
d: [2,303,4,0,0,3,21,302,87]
e: 303
r: 80
r: 159
r: 87
a: [3,88,0,3,1,50]
s: 88
a: [3,79,0,5,1,50]
s: 79
d: [2,305,2,0,0,2,52,88,79,160]
e: 305
d: [2,306,1,0,0,2,52,88,73,79,55]
e: 306
d: [2,308,2,3,1,1,40,79,160]
d: [2,314,2,0,0,2,21,308,88]
e: 314
r: 158
r: 83
r: 88
i: [1,49,0,8,2,1,0,0] t32_zfmisc_1
i: [1,48,0,8,2,1,0,0] t32_zfmisc_1
d: [2,307,2,3,1,1,40,79,161]
f: 308,307,303,48,49,50
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t32_zfmisc_1
% SZS output start Proof for zfmisc_1__t32_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [input t32_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [negated conjecture 1]
5. ! [X0] : ! [X1] : ! [X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [input d2_tarski]
7. ! [X0] : ! [X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [input d3_tarski]
19. ~! [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),X2) <=> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [flattening 2]
20. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X0) => r2_hidden(X2,X1))) [flattening 7]
23. ! [X0,X1,X2] : (k2_tarski(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (X1 = X3 | X0 = X3))) [flattening 5]
25. ? [X0,X1,X2] : (r1_tarski(k2_tarski(X0,X1),X2) <~> (r2_hidden(X1,X2) & r2_hidden(X0,X2))) [ennf transformation 19]
27. ! [X0,X1] : (r1_tarski(X0,X1) <=> ! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0))) [ennf transformation 20]
31. ? [X0,X1,X2] : (((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2)) | ~r1_tarski(k2_tarski(X0,X1),X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | r1_tarski(k2_tarski(X0,X1),X2))) [nnf transformation 25]
32. ? [X0,X1,X2] : ((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~r1_tarski(k2_tarski(X0,X1),X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | r1_tarski(k2_tarski(X0,X1),X2))) [flattening 31]
33. ? [X0,X1,X2] : ((~r2_hidden(X1,X2) | ~r2_hidden(X0,X2) | ~r1_tarski(k2_tarski(X0,X1),X2)) & ((r2_hidden(X1,X2) & r2_hidden(X0,X2)) | r1_tarski(k2_tarski(X0,X1),X2))) => ((~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | ~r1_tarski(k2_tarski(sK0,sK1),sK2)) & ((r2_hidden(sK1,sK2) & r2_hidden(sK0,sK2)) | r1_tarski(k2_tarski(sK0,sK1),sK2))) [choice axiom]
34. (~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | ~r1_tarski(k2_tarski(sK0,sK1),sK2)) & ((r2_hidden(sK1,sK2) & r2_hidden(sK0,sK2)) | r1_tarski(k2_tarski(sK0,sK1),sK2)) [skolemisation 32,33]
35. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X2] : (r2_hidden(X2,X1) | ~r2_hidden(X2,X0)) | ~r1_tarski(X0,X1))) [nnf transformation 27]
36. ! [X0,X1] : ((r1_tarski(X0,X1) | ? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [rectify 35]
37. ! [X1,X0] : (? [X2] : (~r2_hidden(X2,X1) & r2_hidden(X2,X0)) => (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) [choice axiom]
38. ! [X0,X1] : ((r1_tarski(X0,X1) | (~r2_hidden(sK3(X0,X1),X1) & r2_hidden(sK3(X0,X1),X0))) & (! [X3] : (r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r1_tarski(X0,X1))) [skolemisation 36,37]
39. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & ((X1 = X3 | X0 = X3) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & ((X1 = X3 | X0 = X3) | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [nnf transformation 23]
40. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (X1 != X3 & X0 != X3)) & (X1 = X3 | X0 = X3 | ~r2_hidden(X3,X2))) | k2_tarski(X0,X1) != X2)) [flattening 39]
41. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | ? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [rectify 40]
42. ! [X2,X1,X0] : (? [X3] : (((X1 != X3 & X0 != X3) | ~r2_hidden(X3,X2)) & (X1 = X3 | X0 = X3 | r2_hidden(X3,X2))) => (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) [choice axiom]
43. ! [X0,X1,X2] : ((k2_tarski(X0,X1) = X2 | (((sK4(X0,X1,X2) != X1 & sK4(X0,X1,X2) != X0) | ~r2_hidden(sK4(X0,X1,X2),X2)) & (sK4(X0,X1,X2) = X1 | sK4(X0,X1,X2) = X0 | r2_hidden(sK4(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | (X1 != X4 & X0 != X4)) & (X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2))) | k2_tarski(X0,X1) != X2)) [skolemisation 41,42]
48. r2_hidden(sK0,sK2) | r1_tarski(k2_tarski(sK0,sK1),sK2) [cnf transformation 34] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
49. r2_hidden(sK1,sK2) | r1_tarski(k2_tarski(sK0,sK1),sK2) [cnf transformation 34] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
50. ~r2_hidden(sK1,sK2) | ~r2_hidden(sK0,sK2) | ~r1_tarski(k2_tarski(sK0,sK1),sK2) [cnf transformation 34] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
55. ~r2_hidden(X3,X0) | r2_hidden(X3,X1) | ~r1_tarski(X0,X1) [cnf transformation 38] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
56. r2_hidden(sK3(X0,X1),X0) | r1_tarski(X0,X1) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
57. ~r2_hidden(sK3(X0,X1),X1) | r1_tarski(X0,X1) [cnf transformation 38] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
60. X1 = X4 | X0 = X4 | ~r2_hidden(X4,X2) | k2_tarski(X0,X1) != X2 [cnf transformation 43] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
61. r2_hidden(X4,X2) | X0 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 43] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
62. r2_hidden(X4,X2) | X1 != X4 | k2_tarski(X0,X1) != X2 [cnf transformation 43] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
72. r2_hidden(X4,X2) | k2_tarski(X0,X4) != X2 [equality resolution 62] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
73. r2_hidden(X4,k2_tarski(X0,X4)) [equality resolution 72] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
74. r2_hidden(X4,X2) | k2_tarski(X4,X1) != X2 [equality resolution 61] {a:0,w:8,wCS:3200,sine:1,thAx:0,allAx:2,thDist:-2}
75. r2_hidden(X4,k2_tarski(X4,X1)) [equality resolution 74] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
76. ~r2_hidden(X4,k2_tarski(X0,X1)) | X0 = X4 | X1 = X4 [equality resolution 60] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
78. 1 <=> r1_tarski(k2_tarski(sK0,sK1),sK2) [avatar definition]
79. r1_tarski(k2_tarski(sK0,sK1),sK2) <- (1) [avatar component clause 78] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
80. ~r1_tarski(k2_tarski(sK0,sK1),sK2) <- (~1) [avatar component clause 78] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
82. 2 <=> r2_hidden(sK0,sK2) [avatar definition]
83. r2_hidden(sK0,sK2) <- (2) [avatar component clause 82] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
86. 3 <=> r2_hidden(sK1,sK2) [avatar definition]
87. r2_hidden(sK1,sK2) <- (3) [avatar component clause 86] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
89. ~1 | ~2 | ~3 [avatar split clause 50,86,82,78]
90. 1 | 3 [avatar split clause 49,86,78]
91. 1 | 2 [avatar split clause 48,82,78]
160. ~r1_tarski(k2_tarski(X4,X2),X3) | r2_hidden(X2,X3) [resolution 55,73] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. ~r1_tarski(k2_tarski(X5,X7),X6) | r2_hidden(X5,X6) [resolution 55,75] {a:1,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. sK3(k2_tarski(X4,X5),X6) = X5 | sK3(k2_tarski(X4,X5),X6) = X4 | r1_tarski(k2_tarski(X4,X5),X6) [resolution 76,56] {a:1,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
196. ~r2_hidden(X8,X9) | r1_tarski(k2_tarski(X7,X8),X9) | sK3(k2_tarski(X7,X8),X9) = X7 | r1_tarski(k2_tarski(X7,X8),X9) [superposition 57,165] {a:2,w:20,wCS:8000,sine:1,thAx:0,allAx:6,thDist:-6}
200. sK3(k2_tarski(X7,X8),X9) = X7 | r1_tarski(k2_tarski(X7,X8),X9) | ~r2_hidden(X8,X9) [duplicate literal removal 196] {a:2,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
269. ~r2_hidden(X6,X8) | r1_tarski(k2_tarski(X6,X7),X8) | r1_tarski(k2_tarski(X6,X7),X8) | ~r2_hidden(X7,X8) [superposition 57,200] {a:3,w:16,wCS:6400,sine:1,thAx:0,allAx:8,thDist:-8}
280. r1_tarski(k2_tarski(X6,X7),X8) | ~r2_hidden(X6,X8) | ~r2_hidden(X7,X8) [duplicate literal removal 269] {a:3,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:8,thDist:-8}
291. ~r2_hidden(sK0,sK2) | ~r2_hidden(sK1,sK2) <- (~1) [resolution 280,80] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
302. ~r2_hidden(sK1,sK2) <- (~1, 2) [subsumption resolution 291,83] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
303. $false <- (~1, 2, 3) [subsumption resolution 302,87] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
304. 1 | ~2 | ~3 [avatar contradiction clause 303]
307. r2_hidden(sK0,sK2) <- (1) [resolution 79,161] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
308. r2_hidden(sK1,sK2) <- (1) [resolution 79,160] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
316. 2 | ~1 [avatar split clause 307,78,82]
317. 3 | ~1 [avatar split clause 308,78,86]
318. $false [avatar sat refutation 89,90,91,304,316,317] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:107,thDist:-107}
% SZS output end Proof for zfmisc_1__t32_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 17
% Introduced skolems: 7
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 218
% Active clauses: 54
% Passive clauses: 97
% Final active clauses: 47
% Final passive clauses: 38
% Main loop iterations started: 59
% 
% >>> Simplifying Inferences
% Duplicate literals: 26
% Trivial inequalities: 1
% Fw subsumption resolutions: 5
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Equational tautologies: 2
% Forward subsumptions: 64
% 
% >>> Generating Inferences
% Binary resolution: 79
% Unit resulting resolution: 3
% Forward superposition: 24
% Backward superposition: 26
% Self superposition: 1
% Equality factoring: 5
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 7
% Sat splitting refutations: 5
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 5
% SAT solver binary clauses: 11
% 
% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% non unit clause index maintenance: 0.001 s
% forward subsumption: 0.001 s
% other: 0.003 s ( own 0.001 s ) 

