i: [1,33,0,2,1,1,0,0] t50_xxreal_0
a: [3,57,0,2,1,33]
s: 57
i: [1,32,0,2,1,1,0,0] t50_xxreal_0
a: [3,62,0,2,1,32]
s: 62
i: [1,35,0,5,1,1,0,0] t50_xxreal_0
a: [3,47,0,5,1,35]
s: 47
i: [1,34,0,5,1,1,0,0] t50_xxreal_0
a: [3,52,0,5,1,34]
s: 52
i: [1,43,0,2,1,0,0,1] rc1_xxreal_0
a: [3,78,0,2,1,43]
s: 78
i: [1,38,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,69,0,5,2,38]
s: 69
i: [1,40,0,9,3,0,0,1] idempotence_k3_xxreal_0
a: [3,73,0,7,2,40]
s: 73
i: [1,39,0,8,3,0,0,1] fc5_xxreal_0
s: 39
i: [1,36,0,9,3,0,0,1] t17_xxreal_0
s: 36
i: [1,42,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 42
i: [1,41,0,11,3,0,0,1] commutativity_k3_xxreal_0
s: 41
d: [2,88,1,13,5,0,44,36,41]
d: [2,91,1,9,3,0,18,88]
s: 91
i: [1,37,0,13,5,0,0,1] t1_xxreal_0
s: 37
d: [2,100,1,16,4,1,40,37,52]
d: [2,109,1,11,3,2,21,100,47]
d: [2,110,1,9,2,3,21,109,62]
a: [3,114,1,4,1,110]
s: 114
d: [2,122,2,0,0,3,52,62,57,114,39]
e: 122
d: [2,123,2,4,2,1,40,114,39]
d: [2,125,2,2,1,2,21,123,57]
d: [2,126,2,0,0,3,21,125,62]
e: 126
r: 114
a: [3,113,1,4,1,110]
s: 113
a: [3,118,1,5,1,110]
s: 118
d: [2,128,2,0,0,3,52,62,57,118,36]
e: 128
d: [2,129,2,4,2,1,40,118,36]
d: [2,132,2,2,1,2,21,129,57]
d: [2,133,2,0,0,3,21,132,62]
e: 133
f: 133,126,110,32,33,34,35
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_0__t50_xxreal_0
% SZS output start Proof for xxreal_0__t50_xxreal_0
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => (r1_xxreal_0(X0,k3_xxreal_0(X0,X1)) => k3_xxreal_0(X0,X1) = X0))) [input t50_xxreal_0]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => (r1_xxreal_0(X0,k3_xxreal_0(X0,X1)) => k3_xxreal_0(X0,X1) = X0))) [negated conjecture 1]
6. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => v1_xxreal_0(k3_xxreal_0(X0,X1))) [input fc5_xxreal_0]
10. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => r1_xxreal_0(k3_xxreal_0(X0,X1),X0))) [input t17_xxreal_0]
11. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input t1_xxreal_0]
12. ? [X0] : (? [X1] : ((k3_xxreal_0(X0,X1) != X0 & r1_xxreal_0(X0,k3_xxreal_0(X0,X1))) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
13. ? [X0] : (? [X1] : (k3_xxreal_0(X0,X1) != X0 & r1_xxreal_0(X0,k3_xxreal_0(X0,X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [flattening 12]
14. ! [X0] : (! [X1] : (r1_xxreal_0(k3_xxreal_0(X0,X1),X0) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 10]
15. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 11]
16. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [flattening 15]
19. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 6]
20. ! [X0,X1] : (v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 19]
27. ? [X0] : (? [X1] : (k3_xxreal_0(X0,X1) != X0 & r1_xxreal_0(X0,k3_xxreal_0(X0,X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (sK0 != k3_xxreal_0(sK0,X1) & r1_xxreal_0(sK0,k3_xxreal_0(sK0,X1)) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
28. ? [X1] : (sK0 != k3_xxreal_0(sK0,X1) & r1_xxreal_0(sK0,k3_xxreal_0(sK0,X1)) & v1_xxreal_0(X1)) => (sK0 != k3_xxreal_0(sK0,sK1) & r1_xxreal_0(sK0,k3_xxreal_0(sK0,sK1)) & v1_xxreal_0(sK1)) [choice axiom]
29. (sK0 != k3_xxreal_0(sK0,sK1) & r1_xxreal_0(sK0,k3_xxreal_0(sK0,sK1)) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 13,28,27]
32. v1_xxreal_0(sK0) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
33. v1_xxreal_0(sK1) [cnf transformation 29] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
34. r1_xxreal_0(sK0,k3_xxreal_0(sK0,sK1)) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
35. sK0 != k3_xxreal_0(sK0,sK1) [cnf transformation 29] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
36. r1_xxreal_0(k3_xxreal_0(X0,X1),X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 14] {a:0,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 16] {a:0,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. v1_xxreal_0(k3_xxreal_0(X0,X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 20] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. 1 <=> sK0 = k3_xxreal_0(sK0,sK1) [avatar definition]
47. sK0 != k3_xxreal_0(sK0,sK1) <- (~1) [avatar component clause 45] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
48. ~1 [avatar split clause 35,45]
50. 2 <=> r1_xxreal_0(sK0,k3_xxreal_0(sK0,sK1)) [avatar definition]
52. r1_xxreal_0(sK0,k3_xxreal_0(sK0,sK1)) <- (2) [avatar component clause 50] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
53. 2 [avatar split clause 34,50]
55. 3 <=> v1_xxreal_0(sK1) [avatar definition]
57. v1_xxreal_0(sK1) <- (3) [avatar component clause 55] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
58. 3 [avatar split clause 33,55]
60. 4 <=> v1_xxreal_0(sK0) [avatar definition]
62. v1_xxreal_0(sK0) <- (4) [avatar component clause 60] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
63. 4 [avatar split clause 32,60]
100. sK0 = k3_xxreal_0(sK0,sK1) | ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (2) [resolution 37,52] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
109. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (~1, 2) [subsumption resolution 100,47] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
110. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) | ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (~1, 2, 4) [subsumption resolution 109,62] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
112. 9 <=> v1_xxreal_0(k3_xxreal_0(sK0,sK1)) [avatar definition]
114. ~v1_xxreal_0(k3_xxreal_0(sK0,sK1)) <- (~9) [avatar component clause 112] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. 10 <=> r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) [avatar definition]
118. ~r1_xxreal_0(k3_xxreal_0(sK0,sK1),sK0) <- (~10) [avatar component clause 116] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
119. ~9 | ~10 | 1 | ~2 | ~4 [avatar split clause 110,60,50,45,116,112]
123. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~9) [resolution 114,39] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
125. ~v1_xxreal_0(sK0) <- (3, ~9) [subsumption resolution 123,57] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
126. $false <- (3, 4, ~9) [subsumption resolution 125,62] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
127. ~3 | ~4 | 9 [avatar contradiction clause 126]
129. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~10) [resolution 118,36] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
132. ~v1_xxreal_0(sK0) <- (3, ~10) [subsumption resolution 129,57] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
133. $false <- (3, 4, ~10) [subsumption resolution 132,62] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
134. ~3 | ~4 | 10 [avatar contradiction clause 133]
135. $false [avatar sat refutation 48,53,58,63,119,127,134] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:66,thDist:-66}
% SZS output end Proof for xxreal_0__t50_xxreal_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 10
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 64
% Active clauses: 16
% Passive clauses: 19
% Final active clauses: 15
% Final passive clauses: 3
% Main loop iterations started: 16
% 
% >>> Simplifying Inferences
% Duplicate literals: 25
% Fw subsumption resolutions: 8
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 8
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% Factoring: 1
% Forward superposition: 5
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 14
% SAT solver unit clauses: 5
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

