i: [1,132,0,2,1,1,0,0] t16_ordinal3
a: [3,212,0,2,1,132]
s: 212
i: [1,135,0,3,1,1,0,0] t16_ordinal3
a: [3,197,0,3,1,135]
s: 197
i: [1,134,0,3,1,1,0,0] t16_ordinal3
a: [3,202,0,3,1,134]
s: 202
i: [1,133,0,3,1,1,0,0] t16_ordinal3
a: [3,207,0,3,1,133]
s: 207
i: [1,136,0,2,1,0,0,1] spc1_boole
a: [3,217,0,2,1,136]
s: 217
i: [1,137,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,222,0,2,1,137]
s: 222
i: [1,138,0,2,1,0,0,1] fc1_xboole_0
a: [3,227,0,2,1,138]
s: 227
i: [1,144,0,2,1,0,0,255] fc6_ordinal1
a: [3,252,0,2,1,144]
s: 252
i: [1,143,0,2,1,0,0,255] fc6_ordinal1
a: [3,257,0,2,1,143]
s: 257
i: [1,184,0,2,1,0,0,2] rc2_xboole_0
a: [3,284,0,2,1,184]
s: 284
i: [1,185,0,2,1,0,0,255] rc4_ordinal1
a: [3,289,0,2,1,185]
s: 289
i: [1,186,0,2,1,0,0,1] rc1_ordinal1
a: [3,294,0,2,1,186]
s: 294
i: [1,187,0,2,1,0,0,2] rc1_xboole_0
a: [3,299,0,2,1,187]
s: 299
i: [1,191,0,2,1,0,0,255] rc2_ordinal1
a: [3,304,0,2,1,191]
s: 304
i: [1,190,0,2,1,0,0,255] rc2_ordinal1
a: [3,309,0,2,1,190]
s: 309
i: [1,189,0,2,1,0,0,255] rc2_ordinal1
a: [3,314,0,2,1,189]
s: 314
i: [1,188,0,2,1,0,0,255] rc2_ordinal1
a: [3,319,0,2,1,188]
s: 319
i: [1,193,0,2,1,0,0,255] rc5_ordinal1
a: [3,324,0,2,1,193]
s: 324
i: [1,192,0,2,1,0,0,255] rc5_ordinal1
a: [3,329,0,2,1,192]
s: 329
i: [1,172,0,3,1,0,0,3] irreflexivity_r2_xboole_0
s: 172
i: [1,173,0,3,1,0,0,2] reflexivity_r1_tarski
s: 173
i: [1,139,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,232,0,3,1,139]
s: 232
i: [1,140,0,3,1,0,0,255] d2_xboole_0
a: [3,237,0,3,1,140]
s: 237
r: 222
i: [1,141,0,3,1,0,0,255] t1_numerals
a: [3,242,0,3,1,141]
s: 242
i: [1,147,0,3,1,0,0,255] spc1_numerals
a: [3,262,0,3,1,147]
s: 262
i: [1,146,0,3,1,0,0,255] spc1_numerals
a: [3,267,0,3,1,146]
s: 267
d: [2,332,0,3,1,2,22,267,232]
a: [3,336,0,3,1,332]
s: 336
i: [1,150,0,4,2,0,0,255] cc6_ordinal1
s: 150
d: [2,338,1,2,1,1,40,150,289]
a: [3,343,1,2,1,338]
s: 343
i: [1,151,0,4,2,0,0,255] cc1_ordinal1
s: 151
i: [1,152,0,4,2,0,0,255] cc1_ordinal1
s: 152
d: [2,351,1,2,1,1,40,151,212]
a: [3,363,1,2,1,351]
s: 363
i: [1,156,0,4,2,0,0,255] cc9_ordinal1
s: 156
d: [2,376,1,2,1,1,40,152,212]
a: [3,388,1,2,1,376]
s: 388
i: [1,157,0,4,2,0,0,255] cc7_ordinal1
s: 157
i: [1,158,0,4,2,0,0,1] cc3_ordinal1
s: 158
d: [2,339,1,2,1,1,40,150,324]
a: [3,348,1,2,1,339]
s: 348
i: [1,171,0,4,1,0,0,2] existence_m1_subset_1
s: 171
d: [2,350,1,2,1,1,40,151,252]
a: [3,358,1,2,1,350]
s: 358
i: [1,142,0,4,1,0,0,255] dt_k5_numbers
a: [3,247,0,4,1,142]
s: 247
i: [1,145,0,4,1,0,0,255] spc1_numerals
a: [3,272,0,4,1,145]
s: 272
d: [2,353,1,2,1,1,40,151,294]
a: [3,373,1,2,1,353]
s: 373
d: [2,448,0,4,1,2,22,247,232]
a: [3,452,0,4,1,448]
s: 452
d: [2,375,1,2,1,1,40,152,252]
a: [3,383,1,2,1,375]
s: 383
d: [2,454,0,4,1,2,22,272,232]
a: [3,458,0,4,1,454]
s: 458
i: [1,159,0,5,2,0,0,1] t6_boole
s: 159
d: [2,378,1,2,1,1,40,152,294]
a: [3,398,1,2,1,378]
s: 398
i: [1,160,0,5,2,0,0,255] cc8_ordinal1
s: 160
d: [2,400,1,2,1,1,40,156,299]
a: [3,405,1,2,1,400]
s: 405
i: [1,168,0,5,2,0,0,2] t7_boole
s: 168
i: [1,180,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,279,0,5,2,180]
s: 279
d: [2,401,1,2,1,1,40,156,227]
a: [3,410,1,2,1,401]
s: 410
i: [1,161,0,6,3,0,0,255] cc2_ordinal1
s: 161
d: [2,412,1,2,1,1,40,157,299]
a: [3,417,1,2,1,412]
s: 417
i: [1,162,0,6,2,0,0,2] t1_subset
s: 162
i: [1,174,0,6,2,0,0,3] antisymmetry_r2_xboole_0
s: 174
d: [2,413,1,2,1,1,40,157,227]
a: [3,422,1,2,1,413]
s: 422
i: [1,175,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 175
d: [2,424,1,2,1,1,40,158,299]
a: [3,429,1,2,1,424]
s: 429
i: [1,154,0,7,3,0,0,1] cc2_ordinal2
s: 154
i: [1,155,0,7,3,0,0,1] cc5_ordinal1
s: 155
d: [2,425,1,2,1,1,40,158,227]
a: [3,434,1,2,1,425]
s: 434
i: [1,165,0,7,2,0,0,2] t3_subset
s: 165
d: [2,463,1,2,1,1,40,160,336]
a: [3,469,1,2,1,463]
s: 469
i: [1,166,0,7,2,0,0,2] t3_subset
s: 166
i: [1,167,0,7,3,0,0,1] t8_boole
s: 167
d: [2,499,1,4,2,1,40,154,262]
a: [3,511,1,2,1,499]
s: 511
i: [1,148,0,8,3,0,0,255] cc10_ordinal1
s: 148
d: [2,554,1,4,2,1,40,148,452]
a: [3,559,1,2,1,554]
s: 559
i: [1,153,0,8,3,0,0,1] t14_ordinal3
s: 153
i: [1,163,0,8,3,0,0,2] t2_subset
s: 163
d: [2,352,2,2,1,1,40,151,343]
a: [3,368,2,2,1,352]
s: 368
i: [1,164,0,9,3,0,0,2] d8_xboole_0
s: 164
d: [2,377,2,2,1,1,40,152,343]
a: [3,393,2,2,1,377]
s: 393
i: [1,170,0,9,3,0,0,2] t5_subset
s: 170
i: [1,149,0,10,4,0,0,3] t11_ordinal1
s: 149
d: [2,566,1,5,2,1,40,163,262]
a: [3,577,1,2,1,566]
s: 577
i: [1,169,0,10,3,0,0,2] t4_subset
s: 169
d: [2,436,2,2,1,1,40,348,152]
a: [3,441,2,2,1,436]
s: 441
i: [1,181,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 181
i: [1,182,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 182
d: [2,616,1,7,3,1,40,182,207]
d: [2,621,1,5,2,2,21,616,511]
d: [2,622,1,3,1,3,21,621,212]
a: [3,626,1,3,1,622]
s: 626
i: [1,183,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 183
d: [2,628,2,6,2,1,40,626,164]
d: [2,629,2,3,1,2,21,628,197]
a: [3,633,2,3,1,629]
s: 633
i: [1,179,0,14,4,0,0,255] existence_m2_subset_1
s: 179
i: [1,176,0,15,5,0,0,255] dt_m2_subset_1
s: 176
d: [2,638,3,7,3,1,40,633,149]
d: [2,640,3,5,2,2,21,638,511]
d: [2,641,3,3,1,3,21,640,363]
a: [3,645,3,3,1,641]
s: 645
d: [2,655,4,0,0,3,52,212,202,645,153]
e: 655
d: [2,656,4,5,2,1,40,645,153]
d: [2,661,4,2,1,2,21,656,202]
d: [2,662,4,0,0,3,21,661,212]
e: 662
d: [2,500,1,4,2,1,40,154,336]
d: [2,513,1,2,1,2,21,500,252]
f: 662,641,629,622,513,351,332,146,144,139,132,133,134,135
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t16_ordinal3
% SZS output start Proof for ordinal3__t16_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ~(1 != X0 & k1_xboole_0 != X0 & r1_ordinal1(X0,1))) [input t16_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ~(1 != X0 & k1_xboole_0 != X0 & r1_ordinal1(X0,1))) [negated conjecture 1]
6. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
8. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (m1_subset_1(X1,X0) => v3_ordinal1(X1))) [input cc2_ordinal2]
18. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
30. v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1) [input fc6_ordinal1]
38. k4_ordinal1 = k5_numbers [input redefinition_k5_numbers]
40. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
44. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) & v2_xxreal_0(1) [input spc1_numerals]
45. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
46. ! [X0] : (v3_ordinal1(X0) => (r2_hidden(X0,1) => k1_xboole_0 = X0)) [input t14_ordinal3]
58. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 18]
66. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 58]
67. m1_subset_1(1,k1_numbers) & m1_subset_1(1,k5_numbers) & m2_subset_1(1,k1_numbers,k5_numbers) [pure predicate removal 44]
69. ? [X0] : ((1 != X0 & k1_xboole_0 != X0 & r1_ordinal1(X0,1)) & v3_ordinal1(X0)) [ennf transformation 2]
70. ? [X0] : (1 != X0 & k1_xboole_0 != X0 & r1_ordinal1(X0,1) & v3_ordinal1(X0)) [flattening 69]
72. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 45]
73. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 72]
75. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 6]
76. ! [X0] : ((k1_xboole_0 = X0 | ~r2_hidden(X0,1)) | ~v3_ordinal1(X0)) [ennf transformation 46]
77. ! [X0] : (k1_xboole_0 = X0 | ~r2_hidden(X0,1) | ~v3_ordinal1(X0)) [flattening 76]
78. ! [X0] : (! [X1] : (v3_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v3_ordinal1(X0)) [ennf transformation 8]
90. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 66]
91. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 90]
109. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 40]
110. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 109]
111. ? [X0] : (1 != X0 & k1_xboole_0 != X0 & r1_ordinal1(X0,1) & v3_ordinal1(X0)) => (1 != sK0 & k1_xboole_0 != sK0 & r1_ordinal1(sK0,1) & v3_ordinal1(sK0)) [choice axiom]
112. 1 != sK0 & k1_xboole_0 != sK0 & r1_ordinal1(sK0,1) & v3_ordinal1(sK0) [skolemisation 70,111]
119. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 110]
132. v3_ordinal1(sK0) [cnf transformation 112] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
133. r1_ordinal1(sK0,1) [cnf transformation 112] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
134. k1_xboole_0 != sK0 [cnf transformation 112] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
135. 1 != sK0 [cnf transformation 112] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
139. k4_ordinal1 = k5_numbers [cnf transformation 38] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
144. v3_ordinal1(k4_ordinal1) [cnf transformation 30] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
146. m1_subset_1(1,k5_numbers) [cnf transformation 67] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
149. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 73] {a:0,w:10,wCS:4000,nSel:1,sine:3,thAx:0,allAx:1,thDist:-1}
151. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 75] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
153. ~r2_hidden(X0,1) | k1_xboole_0 = X0 | ~v3_ordinal1(X0) [cnf transformation 77] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
154. ~m1_subset_1(X1,X0) | v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 78] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
164. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 91] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
182. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 119] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
195. 1 <=> 1 = sK0 [avatar definition]
197. 1 != sK0 <- (~1) [avatar component clause 195] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
198. ~1 [avatar split clause 135,195]
200. 2 <=> k1_xboole_0 = sK0 [avatar definition]
202. k1_xboole_0 != sK0 <- (~2) [avatar component clause 200] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
203. ~2 [avatar split clause 134,200]
205. 3 <=> r1_ordinal1(sK0,1) [avatar definition]
207. r1_ordinal1(sK0,1) <- (3) [avatar component clause 205] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
208. 3 [avatar split clause 133,205]
210. 4 <=> v3_ordinal1(sK0) [avatar definition]
212. v3_ordinal1(sK0) <- (4) [avatar component clause 210] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
213. 4 [avatar split clause 132,210]
230. 8 <=> k4_ordinal1 = k5_numbers [avatar definition]
232. k4_ordinal1 = k5_numbers <- (8) [avatar component clause 230] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
233. 8 [avatar split clause 139,230]
250. 12 <=> v3_ordinal1(k4_ordinal1) [avatar definition]
252. v3_ordinal1(k4_ordinal1) <- (12) [avatar component clause 250] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
253. 12 [avatar split clause 144,250]
265. 15 <=> m1_subset_1(1,k5_numbers) [avatar definition]
267. m1_subset_1(1,k5_numbers) <- (15) [avatar component clause 265] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
268. 15 [avatar split clause 146,265]
332. m1_subset_1(1,k4_ordinal1) <- (8, 15) [forward demodulation 267,232] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
334. 29 <=> m1_subset_1(1,k4_ordinal1) [avatar definition]
336. m1_subset_1(1,k4_ordinal1) <- (29) [avatar component clause 334] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
337. 29 | ~8 | ~15 [avatar split clause 332,265,230,334]
351. v1_ordinal1(sK0) <- (4) [resolution 151,212] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
361. 33 <=> v1_ordinal1(sK0) [avatar definition]
363. v1_ordinal1(sK0) <- (33) [avatar component clause 361] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
364. 33 | ~4 [avatar split clause 351,210,361]
500. v3_ordinal1(1) | ~v3_ordinal1(k4_ordinal1) <- (29) [resolution 154,336] {a:1,w:4,wCS:1600,sine:1,thAx:0,allAx:2,thDist:-2}
509. 55 <=> v3_ordinal1(1) [avatar definition]
511. v3_ordinal1(1) <- (55) [avatar component clause 509] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
513. v3_ordinal1(1) <- (12, 29) [subsumption resolution 500,252] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
514. 55 | ~12 | ~29 [avatar split clause 513,334,250,509]
616. r1_tarski(sK0,1) | ~v3_ordinal1(1) | ~v3_ordinal1(sK0) <- (3) [resolution 182,207] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
621. r1_tarski(sK0,1) | ~v3_ordinal1(sK0) <- (3, 55) [subsumption resolution 616,511] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
622. r1_tarski(sK0,1) <- (3, 4, 55) [subsumption resolution 621,212] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
624. 70 <=> r1_tarski(sK0,1) [avatar definition]
626. r1_tarski(sK0,1) <- (70) [avatar component clause 624] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
627. 70 | ~3 | ~4 | ~55 [avatar split clause 622,509,210,205,624]
628. 1 = sK0 | r2_xboole_0(sK0,1) <- (70) [resolution 626,164] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
629. r2_xboole_0(sK0,1) <- (~1, 70) [subsumption resolution 628,197] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
631. 71 <=> r2_xboole_0(sK0,1) [avatar definition]
633. r2_xboole_0(sK0,1) <- (71) [avatar component clause 631] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
634. 71 | 1 | ~70 [avatar split clause 629,624,195,631]
638. r2_hidden(sK0,1) | ~v3_ordinal1(1) | ~v1_ordinal1(sK0) <- (71) [resolution 633,149] {a:3,w:7,wCS:700,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
640. r2_hidden(sK0,1) | ~v1_ordinal1(sK0) <- (55, 71) [subsumption resolution 638,511] {a:3,w:5,wCS:500,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
641. r2_hidden(sK0,1) <- (33, 55, 71) [subsumption resolution 640,363] {a:3,w:3,wCS:300,goal:1,sine:3,thAx:0,allAx:5,thDist:-5}
643. 72 <=> r2_hidden(sK0,1) [avatar definition]
645. r2_hidden(sK0,1) <- (72) [avatar component clause 643] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
646. 72 | ~33 | ~55 | ~71 [avatar split clause 641,631,509,361,643]
656. k1_xboole_0 = sK0 | ~v3_ordinal1(sK0) <- (72) [resolution 645,153] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
661. ~v3_ordinal1(sK0) <- (~2, 72) [subsumption resolution 656,202] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
662. $false <- (~2, 4, 72) [subsumption resolution 661,212] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
663. 2 | ~4 | ~72 [avatar contradiction clause 662]
664. $false [avatar sat refutation 198,203,208,213,233,253,268,337,364,514,627,634,646,663] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:325,thDist:-325}
% SZS output end Proof for ordinal3__t16_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 54
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 62
% Generated clauses: 235
% Active clauses: 83
% Passive clauses: 119
% Final active clauses: 82
% Final passive clauses: 32
% Main loop iterations started: 87
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 11
% Fw demodulations: 3
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 77
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 3
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 78
% SAT solver unit clauses: 27
% SAT solver binary clauses: 32
% 
% Memory used [KB]: 5245
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% congruence closure: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

