i: [1,149,0,2,1,1,0,0] t6_real
a: [3,216,0,2,1,149]
s: 216
i: [1,148,0,2,1,1,0,0] t6_real
a: [3,221,0,2,1,148]
s: 221
i: [1,147,0,2,1,1,0,0] t6_real
a: [3,226,0,2,1,147]
s: 226
i: [1,145,0,2,1,1,0,0] t6_real
a: [3,236,0,2,1,145]
s: 236
i: [1,144,0,2,1,1,0,0] t6_real
a: [3,241,0,2,1,144]
s: 241
i: [1,146,0,3,1,1,0,0] t6_real
a: [3,231,0,3,1,146]
s: 231
i: [1,150,0,2,1,0,0,255] fc1_xboole_0
a: [3,246,0,2,1,150]
s: 246
i: [1,156,0,2,1,0,0,255] fc6_ordinal1
a: [3,271,0,2,1,156]
s: 271
i: [1,155,0,2,1,0,0,255] fc6_ordinal1
a: [3,276,0,2,1,155]
s: 276
i: [1,193,0,2,1,0,0,1] rc2_xboole_0
a: [3,288,0,2,1,193]
s: 288
i: [1,194,0,2,1,0,0,255] rc4_ordinal1
a: [3,293,0,2,1,194]
s: 293
i: [1,195,0,2,1,0,0,255] rc1_ordinal1
a: [3,298,0,2,1,195]
s: 298
i: [1,196,0,2,1,0,0,1] rc1_xreal_0
a: [3,303,0,2,1,196]
s: 303
i: [1,197,0,2,1,0,0,1] rc1_xboole_0
a: [3,308,0,2,1,197]
s: 308
i: [1,201,0,2,1,0,0,255] rc2_ordinal1
a: [3,313,0,2,1,201]
s: 313
i: [1,200,0,2,1,0,0,255] rc2_ordinal1
a: [3,318,0,2,1,200]
s: 318
i: [1,199,0,2,1,0,0,255] rc2_ordinal1
a: [3,323,0,2,1,199]
s: 323
i: [1,198,0,2,1,0,0,255] rc2_ordinal1
a: [3,328,0,2,1,198]
s: 328
i: [1,203,0,2,1,0,0,255] rc5_ordinal1
a: [3,333,0,2,1,203]
s: 333
i: [1,202,0,2,1,0,0,255] rc5_ordinal1
a: [3,338,0,2,1,202]
s: 338
i: [1,206,0,2,1,0,0,1] rc2_xreal_0
a: [3,343,0,2,1,206]
s: 343
i: [1,205,0,2,1,0,0,1] rc2_xreal_0
a: [3,348,0,2,1,205]
s: 348
i: [1,204,0,2,1,0,0,1] rc2_xreal_0
a: [3,353,0,2,1,204]
s: 353
i: [1,209,0,2,1,0,0,1] rc3_xreal_0
a: [3,358,0,2,1,209]
s: 358
i: [1,208,0,2,1,0,0,1] rc3_xreal_0
a: [3,363,0,2,1,208]
s: 363
i: [1,207,0,2,1,0,0,1] rc3_xreal_0
a: [3,368,0,2,1,207]
s: 368
i: [1,212,0,2,1,0,0,2] rc4_xreal_0
a: [3,373,0,2,1,212]
s: 373
i: [1,211,0,2,1,0,0,2] rc4_xreal_0
a: [3,378,0,2,1,211]
s: 378
i: [1,210,0,2,1,0,0,2] rc4_xreal_0
a: [3,383,0,2,1,210]
s: 383
i: [1,185,0,3,1,0,0,255] reflexivity_r1_tarski
s: 185
i: [1,151,0,3,1,0,0,255] redefinition_k6_numbers
a: [3,251,0,3,1,151]
s: 251
r: 246
d: [2,385,0,2,1,2,23,246,251]
a: [3,389,0,2,1,385]
s: 389
i: [1,152,0,3,1,0,0,255] redefinition_k5_numbers
a: [3,256,0,3,1,152]
s: 256
i: [1,158,0,4,2,0,0,255] cc6_ordinal1
s: 158
i: [1,159,0,4,2,0,0,255] cc2_xreal_0
s: 159
i: [1,164,0,4,2,0,0,255] cc1_ordinal1
s: 164
d: [2,391,1,2,1,1,40,158,293]
a: [3,396,1,2,1,391]
s: 396
i: [1,165,0,4,2,0,0,255] cc1_ordinal1
s: 165
d: [2,392,1,2,1,1,40,158,333]
a: [3,401,1,2,1,392]
s: 401
i: [1,167,0,4,2,0,0,2] cc4_xreal_0
s: 167
i: [1,170,0,4,2,0,0,255] cc9_ordinal1
s: 170
d: [2,465,1,2,1,1,40,167,241]
a: [3,474,1,2,1,465]
s: 474
i: [1,171,0,4,2,0,0,255] cc7_ordinal1
s: 171
d: [2,466,1,2,1,1,40,167,236]
a: [3,479,1,2,1,466]
s: 479
i: [1,172,0,4,2,0,0,255] cc3_ordinal1
s: 172
i: [1,184,0,4,1,0,0,255] existence_m1_subset_1
s: 184
d: [2,403,1,2,1,1,40,159,293]
a: [3,408,1,2,1,403]
s: 408
i: [1,153,0,4,1,0,0,255] dt_k5_numbers
a: [3,261,0,4,1,153]
s: 261
d: [2,404,1,2,1,1,40,159,333]
a: [3,413,1,2,1,404]
s: 413
i: [1,154,0,4,1,0,0,255] dt_k6_numbers
a: [3,266,0,4,1,154]
s: 266
d: [2,546,0,4,1,2,22,261,256]
a: [3,550,0,4,1,546]
s: 550
d: [2,415,1,2,1,1,40,164,271]
a: [3,421,1,2,1,415]
s: 421
d: [2,558,0,4,1,2,22,266,256]
a: [3,562,0,4,1,558]
s: 562
d: [2,416,1,2,1,1,40,164,298]
a: [3,426,1,2,1,416]
s: 426
i: [1,173,0,5,2,0,0,255] t6_boole
s: 173
i: [1,174,0,5,2,0,0,255] cc1_xreal_0
s: 174
d: [2,434,1,2,1,1,40,165,271]
a: [3,441,1,2,1,434]
s: 441
i: [1,175,0,5,2,0,0,255] cc8_ordinal1
s: 175
d: [2,436,1,2,1,1,40,165,298]
a: [3,451,1,2,1,436]
s: 451
i: [1,181,0,5,2,0,0,255] t7_boole
s: 181
i: [1,191,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,283,0,5,2,191]
s: 283
d: [2,467,1,2,1,1,40,167,303]
a: [3,484,1,2,1,467]
s: 484
d: [2,564,0,5,2,1,22,173,251]
s: 564
d: [2,486,1,2,1,1,40,170,383]
a: [3,492,1,2,1,486]
s: 492
i: [1,176,0,6,3,0,0,255] cc2_ordinal1
s: 176
i: [1,177,0,6,2,0,0,255] t1_subset
s: 177
d: [2,487,1,2,1,1,40,170,308]
a: [3,497,1,2,1,487]
s: 497
i: [1,186,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 186
d: [2,488,1,2,1,1,40,170,389]
a: [3,502,1,2,1,488]
s: 502
i: [1,160,0,7,3,0,0,1] d6_xxreal_0
s: 160
i: [1,161,0,7,3,0,0,1] d6_xxreal_0
s: 161
d: [2,504,1,2,1,1,40,171,383]
a: [3,510,1,2,1,504]
s: 510
i: [1,162,0,7,3,0,0,1] d7_xxreal_0
s: 162
d: [2,505,1,2,1,1,40,171,308]
a: [3,515,1,2,1,505]
s: 515
i: [1,163,0,7,3,0,0,1] d7_xxreal_0
s: 163
i: [1,166,0,7,3,0,0,255] cc5_ordinal1
s: 166
d: [2,506,1,2,1,1,40,171,389]
a: [3,520,1,2,1,506]
s: 520
i: [1,179,0,7,2,0,0,255] t3_subset
s: 179
d: [2,522,1,2,1,1,40,172,383]
a: [3,528,1,2,1,522]
s: 528
i: [1,180,0,7,3,0,0,255] t8_boole
s: 180
i: [1,157,0,8,3,0,0,255] cc10_ordinal1
s: 157
d: [2,523,1,2,1,1,40,172,308]
a: [3,533,1,2,1,523]
s: 533
i: [1,178,0,8,3,0,0,255] t2_subset
s: 178
d: [2,524,1,2,1,1,40,172,389]
a: [3,538,1,2,1,524]
s: 538
i: [1,183,0,9,3,0,0,255] t5_subset
s: 183
i: [1,182,0,10,3,0,0,255] t4_subset
s: 182
d: [2,584,1,6,3,1,40,160,283]
d: [2,586,1,4,2,1,18,584]
a: [3,590,1,2,1,586]
s: 590
i: [1,192,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 192
d: [2,652,1,4,2,1,40,157,550]
a: [3,657,1,2,1,652]
s: 657
i: [1,168,0,13,5,0,0,255] l1_real
s: 168
i: [1,190,0,14,4,0,0,255] existence_m2_subset_1
s: 190
d: [2,721,1,10,4,1,40,168,231]
d: [2,727,1,8,3,2,21,721,241]
d: [2,728,1,6,2,3,21,727,236]
a: [3,732,1,3,1,728]
s: 732
i: [1,169,0,15,6,0,0,1] l2_real
s: 169
d: [2,749,1,12,5,1,40,169,231]
d: [2,755,1,10,4,2,21,749,236]
d: [2,756,1,8,3,3,21,755,241]
s: 756
i: [1,187,0,15,5,0,0,255] dt_m2_subset_1
s: 187
i: [1,188,0,15,5,0,0,255] redefinition_m2_subset_1
s: 188
d: [2,771,2,9,4,3,40,756,163]
d: [2,774,2,7,3,4,21,771,216]
d: [2,775,2,5,2,5,21,774,474]
a: [3,779,2,3,1,775]
s: 779
i: [1,189,0,15,5,0,0,255] redefinition_m2_subset_1
s: 189
d: [2,801,3,4,2,1,40,779,162]
d: [2,805,3,2,1,2,21,801,479]
a: [3,809,3,2,1,805]
s: 809
d: [2,701,1,5,2,1,40,183,550]
a: [3,708,1,2,1,701]
s: 708
d: [2,834,2,0,0,2,52,657,708,170]
e: 834
d: [2,837,2,2,1,1,40,708,170]
d: [2,839,2,0,0,2,21,837,657]
e: 839
r: 657
r: 708
a: [3,661,1,2,1,652]
s: 661
d: [2,803,3,10,4,1,40,779,168]
d: [2,816,3,8,3,2,21,803,236]
a: [3,820,3,3,1,816]
s: 820
d: [2,852,4,0,0,3,52,479,221,820,161]
e: 852
d: [2,853,4,4,2,1,40,820,161]
d: [2,857,4,2,1,2,21,853,221]
d: [2,858,4,0,0,3,21,857,479]
e: 858
r: 590
r: 820
a: [3,589,1,2,1,586]
s: 589
a: [3,819,3,3,1,816]
s: 819
a: [3,594,1,2,1,586]
s: 594
d: [2,602,1,6,3,1,40,162,283]
d: [2,606,1,4,2,1,18,602]
a: [3,610,1,2,1,606]
s: 610
a: [3,824,3,3,1,816]
s: 824
r: 221
r: 231
r: 236
r: 479
r: 732
r: 756
r: 779
r: 809
r: 819
d: [2,885,0,3,1,2,23,231,824]
a: [3,896,0,3,1,885]
s: 896
d: [2,910,1,4,2,1,40,896,160]
d: [2,914,1,2,1,2,21,910,474]
a: [3,918,1,2,1,914]
s: 918
d: [2,888,1,3,1,2,23,732,824]
a: [3,902,1,3,1,888]
s: 902
d: [2,931,2,0,0,3,52,474,216,902,163]
e: 931
d: [2,932,2,4,2,1,40,902,163]
d: [2,936,2,2,1,2,21,932,216]
d: [2,937,2,0,0,3,21,936,474]
e: 937
d: [2,629,2,2,1,1,40,520,159]
d: [2,872,4,10,4,1,40,819,168]
d: [2,879,4,7,3,2,21,872,779]
d: [2,880,4,5,2,3,21,879,236]
f: 937,888,880,858,775,728,629,506,466,465,385,151,150,144,145,146,147,148,149
% Refutation found. Thanks to Tanya!
% SZS status Theorem for real__t6_real
% SZS output start Proof for real__t6_real
1. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [input t6_real]
2. ~! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => (r1_xxreal_0(X0,X1) => (v3_xxreal_0(X0) | v2_xxreal_0(X1) | v1_xboole_0(X0))))) [negated conjecture 1]
8. ! [X0] : (v7_ordinal1(X0) => v1_xreal_0(X0)) [input cc2_xreal_0]
12. ! [X0] : (v1_xreal_0(X0) => v1_xxreal_0(X0)) [input cc4_xreal_0]
15. ! [X0] : (v1_xboole_0(X0) => v7_ordinal1(X0)) [input cc7_ordinal1]
19. ! [X0] : (v1_xxreal_0(X0) => (v2_xxreal_0(X0) <=> ~r1_xxreal_0(X0,k6_numbers))) [input d6_xxreal_0]
20. ! [X0] : (v1_xxreal_0(X0) => (v3_xxreal_0(X0) <=> ~r1_xxreal_0(k6_numbers,X0))) [input d7_xxreal_0]
31. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
33. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ((r1_xxreal_0(X1,X0) & r1_xxreal_0(X0,X1)) => X0 = X1))) [input l1_real]
34. ! [X0] : (v1_xreal_0(X0) => ! [X1] : (v1_xreal_0(X1) => ! [X2] : (v1_xreal_0(X2) => ((r1_xxreal_0(X1,X2) & r1_xxreal_0(X0,X1)) => r1_xxreal_0(X0,X2))))) [input l2_real]
46. k6_numbers = k1_xboole_0 [input redefinition_k6_numbers]
72. ? [X0] : (? [X1] : (((~v3_xxreal_0(X0) & ~v2_xxreal_0(X1) & ~v1_xboole_0(X0)) & r1_xxreal_0(X0,X1)) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [ennf transformation 2]
73. ? [X0] : (? [X1] : (~v3_xxreal_0(X0) & ~v2_xxreal_0(X1) & ~v1_xboole_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) [flattening 72]
76. ! [X0] : (v1_xreal_0(X0) | ~v7_ordinal1(X0)) [ennf transformation 8]
77. ! [X0] : ((v2_xxreal_0(X0) <=> ~r1_xxreal_0(X0,k6_numbers)) | ~v1_xxreal_0(X0)) [ennf transformation 19]
78. ! [X0] : ((v3_xxreal_0(X0) <=> ~r1_xxreal_0(k6_numbers,X0)) | ~v1_xxreal_0(X0)) [ennf transformation 20]
81. ! [X0] : (v1_xxreal_0(X0) | ~v1_xreal_0(X0)) [ennf transformation 12]
82. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 33]
83. ! [X0] : (! [X1] : (X0 = X1 | ~r1_xxreal_0(X1,X0) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 82]
84. ! [X0] : (! [X1] : (! [X2] : ((r1_xxreal_0(X0,X2) | (~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1))) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [ennf transformation 34]
85. ! [X0] : (! [X1] : (! [X2] : (r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X1,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2)) | ~v1_xreal_0(X1)) | ~v1_xreal_0(X0)) [flattening 84]
87. ! [X0] : (v7_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 15]
114. ? [X0] : (? [X1] : (~v3_xxreal_0(X0) & ~v2_xxreal_0(X1) & ~v1_xboole_0(X0) & r1_xxreal_0(X0,X1) & v1_xreal_0(X1)) & v1_xreal_0(X0)) => (? [X1] : (~v3_xxreal_0(sK0) & ~v2_xxreal_0(X1) & ~v1_xboole_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) & v1_xreal_0(sK0)) [choice axiom]
115. ? [X1] : (~v3_xxreal_0(sK0) & ~v2_xxreal_0(X1) & ~v1_xboole_0(sK0) & r1_xxreal_0(sK0,X1) & v1_xreal_0(X1)) => (~v3_xxreal_0(sK0) & ~v2_xxreal_0(sK1) & ~v1_xboole_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) [choice axiom]
116. (~v3_xxreal_0(sK0) & ~v2_xxreal_0(sK1) & ~v1_xboole_0(sK0) & r1_xxreal_0(sK0,sK1) & v1_xreal_0(sK1)) & v1_xreal_0(sK0) [skolemisation 73,115,114]
117. ! [X0] : (((v2_xxreal_0(X0) | r1_xxreal_0(X0,k6_numbers)) & (~r1_xxreal_0(X0,k6_numbers) | ~v2_xxreal_0(X0))) | ~v1_xxreal_0(X0)) [nnf transformation 77]
118. ! [X0] : (((v3_xxreal_0(X0) | r1_xxreal_0(k6_numbers,X0)) & (~r1_xxreal_0(k6_numbers,X0) | ~v3_xxreal_0(X0))) | ~v1_xxreal_0(X0)) [nnf transformation 78]
144. v1_xreal_0(sK0) [cnf transformation 116] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
145. v1_xreal_0(sK1) [cnf transformation 116] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
146. r1_xxreal_0(sK0,sK1) [cnf transformation 116] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
147. ~v1_xboole_0(sK0) [cnf transformation 116] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
148. ~v2_xxreal_0(sK1) [cnf transformation 116] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
149. ~v3_xxreal_0(sK0) [cnf transformation 116] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
150. v1_xboole_0(k1_xboole_0) [cnf transformation 31] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
151. k6_numbers = k1_xboole_0 [cnf transformation 46] {a:0,w:3,wCS:1200,sine:255,thAx:0,allAx:1,thDist:-1}
159. ~v7_ordinal1(X0) | v1_xreal_0(X0) [cnf transformation 76] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
161. r1_xxreal_0(X0,k6_numbers) | v2_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 117] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
163. r1_xxreal_0(k6_numbers,X0) | v3_xxreal_0(X0) | ~v1_xxreal_0(X0) [cnf transformation 118] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
167. ~v1_xreal_0(X0) | v1_xxreal_0(X0) [cnf transformation 81] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
168. ~r1_xxreal_0(X1,X0) | X0 = X1 | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 83] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
169. ~r1_xxreal_0(X1,X2) | r1_xxreal_0(X0,X2) | ~r1_xxreal_0(X0,X1) | ~v1_xreal_0(X2) | ~v1_xreal_0(X1) | ~v1_xreal_0(X0) [cnf transformation 85] {a:0,w:15,wCS:6000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
171. ~v1_xboole_0(X0) | v7_ordinal1(X0) [cnf transformation 87] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
214. 1 <=> v3_xxreal_0(sK0) [avatar definition]
216. ~v3_xxreal_0(sK0) <- (~1) [avatar component clause 214] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
217. ~1 [avatar split clause 149,214]
219. 2 <=> v2_xxreal_0(sK1) [avatar definition]
221. ~v2_xxreal_0(sK1) <- (~2) [avatar component clause 219] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
222. ~2 [avatar split clause 148,219]
224. 3 <=> v1_xboole_0(sK0) [avatar definition]
227. ~3 [avatar split clause 147,224]
229. 4 <=> r1_xxreal_0(sK0,sK1) [avatar definition]
231. r1_xxreal_0(sK0,sK1) <- (4) [avatar component clause 229] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
232. 4 [avatar split clause 146,229]
234. 5 <=> v1_xreal_0(sK1) [avatar definition]
236. v1_xreal_0(sK1) <- (5) [avatar component clause 234] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
237. 5 [avatar split clause 145,234]
239. 6 <=> v1_xreal_0(sK0) [avatar definition]
241. v1_xreal_0(sK0) <- (6) [avatar component clause 239] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
242. 6 [avatar split clause 144,239]
244. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
246. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 244] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
247. 7 [avatar split clause 150,244]
249. 8 <=> k6_numbers = k1_xboole_0 [avatar definition]
251. k6_numbers = k1_xboole_0 <- (8) [avatar component clause 249] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
252. 8 [avatar split clause 151,249]
385. v1_xboole_0(k6_numbers) <- (7, 8) [backward demodulation 246,251] {a:0,w:2,wCS:800,sine:255,thAx:0,allAx:1,thDist:-1}
387. 36 <=> v1_xboole_0(k6_numbers) [avatar definition]
389. v1_xboole_0(k6_numbers) <- (36) [avatar component clause 387] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
390. 36 | ~7 | ~8 [avatar split clause 385,249,244,387]
465. v1_xxreal_0(sK0) <- (6) [resolution 167,241] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
466. v1_xxreal_0(sK1) <- (5) [resolution 167,236] {a:1,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:4,thDist:-4}
472. 49 <=> v1_xxreal_0(sK0) [avatar definition]
474. v1_xxreal_0(sK0) <- (49) [avatar component clause 472] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
475. 49 | ~6 [avatar split clause 465,239,472]
477. 50 <=> v1_xxreal_0(sK1) [avatar definition]
479. v1_xxreal_0(sK1) <- (50) [avatar component clause 477] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
480. 50 | ~5 [avatar split clause 466,234,477]
506. v7_ordinal1(k6_numbers) <- (36) [resolution 171,389] {a:1,w:2,wCS:800,sine:255,thAx:0,allAx:2,thDist:-2}
518. 57 <=> v7_ordinal1(k6_numbers) [avatar definition]
520. v7_ordinal1(k6_numbers) <- (57) [avatar component clause 518] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
521. 57 | ~36 [avatar split clause 506,387,518]
629. v1_xreal_0(k6_numbers) <- (57) [resolution 520,159] {a:2,w:2,wCS:800,sine:255,thAx:0,allAx:3,thDist:-3}
633. 71 <=> v1_xreal_0(k6_numbers) [avatar definition]
636. 71 | ~57 [avatar split clause 629,518,633]
721. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(sK1) <- (4) [resolution 168,231] {a:1,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
727. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) | ~v1_xreal_0(sK1) <- (4, 6) [subsumption resolution 721,241] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
728. sK0 = sK1 | ~r1_xxreal_0(sK1,sK0) <- (4, 5, 6) [subsumption resolution 727,236] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
730. 84 <=> r1_xxreal_0(sK1,sK0) [avatar definition]
732. ~r1_xxreal_0(sK1,sK0) <- (~84) [avatar component clause 730] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
734. 85 <=> sK0 = sK1 [avatar definition]
737. ~84 | 85 | ~4 | ~5 | ~6 [avatar split clause 728,239,234,229,734,730]
749. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xreal_0(sK1) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X0) <- (4) [resolution 169,231] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
755. r1_xxreal_0(X0,sK1) | ~r1_xxreal_0(X0,sK0) | ~v1_xreal_0(sK0) | ~v1_xreal_0(X0) <- (4, 5) [subsumption resolution 749,236] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
756. ~r1_xxreal_0(X0,sK0) | r1_xxreal_0(X0,sK1) | ~v1_xreal_0(X0) <- (4, 5, 6) [subsumption resolution 755,241] {a:1,w:8,wCS:800,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
771. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) | v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (4, 5, 6) [resolution 756,163] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
774. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) | ~v1_xxreal_0(sK0) <- (~1, 4, 5, 6) [subsumption resolution 771,216] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
775. r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(k6_numbers) <- (~1, 4, 5, 6, 49) [subsumption resolution 774,474] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
777. 90 <=> r1_xxreal_0(k6_numbers,sK1) [avatar definition]
779. r1_xxreal_0(k6_numbers,sK1) <- (90) [avatar component clause 777] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
780. ~71 | 90 | 1 | ~4 | ~5 | ~6 | ~49 [avatar split clause 775,472,239,234,229,214,777,633]
818. 95 <=> r1_xxreal_0(sK1,k6_numbers) [avatar definition]
819. r1_xxreal_0(sK1,k6_numbers) <- (95) [avatar component clause 818] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
820. ~r1_xxreal_0(sK1,k6_numbers) <- (~95) [avatar component clause 818] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
822. 96 <=> k6_numbers = sK1 [avatar definition]
824. k6_numbers = sK1 <- (96) [avatar component clause 822] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
853. v2_xxreal_0(sK1) | ~v1_xxreal_0(sK1) <- (~95) [resolution 820,161] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
857. ~v1_xxreal_0(sK1) <- (~2, ~95) [subsumption resolution 853,221] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
858. $false <- (~2, 50, ~95) [subsumption resolution 857,479] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
859. 2 | ~50 | 95 [avatar contradiction clause 858]
872. k6_numbers = sK1 | ~r1_xxreal_0(k6_numbers,sK1) | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (95) [resolution 819,168] {a:4,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
879. k6_numbers = sK1 | ~v1_xreal_0(sK1) | ~v1_xreal_0(k6_numbers) <- (90, 95) [subsumption resolution 872,779] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
880. k6_numbers = sK1 | ~v1_xreal_0(k6_numbers) <- (5, 90, 95) [subsumption resolution 879,236] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
881. ~71 | 96 | ~5 | ~90 | ~95 [avatar split clause 880,818,777,234,822,633]
888. ~r1_xxreal_0(k6_numbers,sK0) <- (~84, 96) [backward demodulation 732,824] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
900. 104 <=> r1_xxreal_0(k6_numbers,sK0) [avatar definition]
902. ~r1_xxreal_0(k6_numbers,sK0) <- (~104) [avatar component clause 900] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
903. ~104 | 84 | ~96 [avatar split clause 888,822,730,900]
932. v3_xxreal_0(sK0) | ~v1_xxreal_0(sK0) <- (~104) [resolution 902,163] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
936. ~v1_xxreal_0(sK0) <- (~1, ~104) [subsumption resolution 932,216] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
937. $false <- (~1, 49, ~104) [subsumption resolution 936,474] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
938. 1 | ~49 | 104 [avatar contradiction clause 937]
939. k6_numbers != k1_xboole_0 | k6_numbers != sK1 | sK0 != sK1 | v1_xboole_0(sK0) | ~v1_xboole_0(k1_xboole_0) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
940. $false [avatar sat refutation 217,222,227,232,237,242,247,252,390,475,480,521,636,737,780,859,881,903,938,939] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:592,thDist:-592}
% SZS output end Proof for real__t6_real
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 56
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 2
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 69
% Generated clauses: 370
% Active clauses: 107
% Passive clauses: 174
% Final active clauses: 93
% Final passive clauses: 62
% Main loop iterations started: 110
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 33
% Fw demodulations: 3
% Bw demodulations: 10
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Equational tautologies: 1
% Forward subsumptions: 24
% 
% >>> Generating Inferences
% Binary resolution: 108
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 16
% Backward superposition: 2
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 25
% Split components: 54
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 120
% SAT solver unit clauses: 35
% SAT solver binary clauses: 49
% 
% Memory used [KB]: 9978
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% preprocessing: 0.001 s
% sat proof minimization: 0.001 s
% forward subsumption: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.0000 s ) 

