i: [1,19,0,12,4,1,0,0] t60_mcart_1
a: [3,31,0,3,1,19]
s: 31
i: [1,18,0,11,1,1,0,0] t60_mcart_1
a: [3,44,0,11,1,18]
s: 44
i: [1,23,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 23
d: [2,46,1,3,1,1,40,23,31]
a: [3,50,1,3,1,46]
s: 50
i: [1,21,0,10,2,0,0,1] t104_zfmisc_1
s: 21
i: [1,22,0,10,2,0,0,1] t104_zfmisc_1
s: 22
d: [2,53,1,10,2,0,40,21,23]
s: 53
i: [1,20,0,13,1,0,0,255] t49_mcart_1
s: 20
d: [2,54,1,10,2,0,40,22,23]
s: 54
d: [2,60,1,12,2,0,44,22,20]
s: 60
d: [2,61,1,12,2,0,44,22,20]
s: 61
d: [2,65,2,12,2,0,44,54,20]
s: 65
d: [2,57,1,15,1,0,44,20,20]
s: 57
d: [2,66,2,12,2,0,44,54,20]
s: 66
d: [2,58,1,16,2,0,44,21,20]
s: 58
d: [2,59,1,16,2,0,44,21,20]
s: 59
d: [2,67,2,12,2,0,40,60,23]
s: 67
d: [2,56,1,19,1,0,44,20,20]
s: 56
d: [2,69,2,12,2,0,40,61,23]
s: 69
d: [2,68,2,14,2,0,44,60,20]
s: 68
a: [3,38,0,3,1,19]
s: 38
d: [2,70,2,14,2,0,44,61,20]
s: 70
d: [2,62,2,16,2,0,44,53,20]
s: 62
d: [2,72,3,14,2,0,44,65,20]
s: 72
d: [2,63,2,16,2,0,44,53,20]
s: 63
d: [2,82,2,16,2,0,40,58,23]
s: 82
d: [2,122,4,3,1,1,40,72,44]
a: [3,131,4,3,1,122]
s: 131
d: [2,86,2,16,2,0,40,59,23]
s: 86
d: [2,80,3,14,2,0,44,66,20]
s: 80
d: [2,84,2,22,2,0,44,58,20]
s: 84
d: [2,87,2,22,2,0,44,59,20]
s: 87
d: [2,109,3,14,2,0,40,68,23]
s: 109
d: [2,140,3,14,2,0,40,84,54]
s: 140
d: [2,154,4,3,1,1,40,140,44]
a: [3,163,4,3,1,154]
s: 163
d: [2,142,3,14,2,0,40,84,22]
s: 142
a: [3,34,0,3,1,19]
s: 34
d: [2,155,4,14,2,0,40,140,23]
s: 155
d: [2,141,3,18,2,0,40,84,53]
s: 141
d: [2,172,4,7,1,1,40,141,44]
a: [3,181,4,7,1,172]
s: 181
d: [2,185,5,0,0,2,52,31,181,54]
e: 185
d: [2,186,5,0,0,2,52,50,181,22]
e: 186
d: [2,187,5,3,1,1,40,181,54]
d: [2,193,5,0,0,2,21,187,31]
e: 193
d: [2,189,5,3,1,1,40,181,22]
d: [2,195,5,0,0,2,21,189,50]
e: 195
r: 31
r: 50
d: [2,108,3,3,1,1,40,68,44]
d: [2,165,4,3,1,1,40,142,44]
d: [2,188,5,3,1,1,40,181,53]
f: 188,187,172,165,108,18,19
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t60_mcart_1
% SZS output start Proof for mcart_1__t60_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k4_zfmisc_1(X4,X5,X6,X7)) => (~r1_xboole_0(X3,X7) & ~r1_xboole_0(X2,X6) & ~r1_xboole_0(X1,X5) & ~r1_xboole_0(X0,X4))) [input t60_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : ! [X6] : ! [X7] : (~r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k4_zfmisc_1(X4,X5,X6,X7)) => (~r1_xboole_0(X3,X7) & ~r1_xboole_0(X2,X6) & ~r1_xboole_0(X1,X5) & ~r1_xboole_0(X0,X4))) [negated conjecture 1]
6. ! [X0,X1] : (r1_xboole_0(X0,X1) => r1_xboole_0(X1,X0)) [input symmetry_r1_xboole_0]
7. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_xboole_0(X2,X3) | r1_xboole_0(X0,X1)) => r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [input t104_zfmisc_1]
8. ! [X0] : ! [X1] : ! [X2] : ! [X3] : k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2),X3) [input t49_mcart_1]
9. ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (~r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k4_zfmisc_1(X4,X5,X6,X7)) => (~r1_xboole_0(X3,X7) & ~r1_xboole_0(X2,X6) & ~r1_xboole_0(X1,X5) & ~r1_xboole_0(X0,X4))) [flattening 2]
10. ! [X0,X1,X2,X3] : k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2),X3) [flattening 8]
11. ! [X0,X1,X2,X3] : ((r1_xboole_0(X2,X3) | r1_xboole_0(X0,X1)) => r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [flattening 7]
13. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((r1_xboole_0(X3,X7) | r1_xboole_0(X2,X6) | r1_xboole_0(X1,X5) | r1_xboole_0(X0,X4)) & ~r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k4_zfmisc_1(X4,X5,X6,X7))) [ennf transformation 9]
14. ! [X0,X1,X2,X3] : (r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | (~r1_xboole_0(X2,X3) & ~r1_xboole_0(X0,X1))) [ennf transformation 11]
15. ! [X0,X1] : (r1_xboole_0(X1,X0) | ~r1_xboole_0(X0,X1)) [ennf transformation 6]
16. ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((r1_xboole_0(X3,X7) | r1_xboole_0(X2,X6) | r1_xboole_0(X1,X5) | r1_xboole_0(X0,X4)) & ~r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k4_zfmisc_1(X4,X5,X6,X7))) => ((r1_xboole_0(sK3,sK7) | r1_xboole_0(sK2,sK6) | r1_xboole_0(sK1,sK5) | r1_xboole_0(sK0,sK4)) & ~r1_xboole_0(k4_zfmisc_1(sK0,sK1,sK2,sK3),k4_zfmisc_1(sK4,sK5,sK6,sK7))) [choice axiom]
17. (r1_xboole_0(sK3,sK7) | r1_xboole_0(sK2,sK6) | r1_xboole_0(sK1,sK5) | r1_xboole_0(sK0,sK4)) & ~r1_xboole_0(k4_zfmisc_1(sK0,sK1,sK2,sK3),k4_zfmisc_1(sK4,sK5,sK6,sK7)) [skolemisation 13,16]
18. ~r1_xboole_0(k4_zfmisc_1(sK0,sK1,sK2,sK3),k4_zfmisc_1(sK4,sK5,sK6,sK7)) [cnf transformation 17] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
19. r1_xboole_0(sK3,sK7) | r1_xboole_0(sK2,sK6) | r1_xboole_0(sK1,sK5) | r1_xboole_0(sK0,sK4) [cnf transformation 17] {a:0,w:12,wCS:1200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. k4_zfmisc_1(X0,X1,X2,X3) = k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2),X3) [cnf transformation 10] {a:0,w:13,wCS:5200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
21. r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_xboole_0(X0,X1) [cnf transformation 14] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
22. r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_xboole_0(X2,X3) [cnf transformation 14] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
23. ~r1_xboole_0(X0,X1) | r1_xboole_0(X1,X0) [cnf transformation 15] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. 1 <=> r1_xboole_0(sK0,sK4) [avatar definition]
29. 2 <=> r1_xboole_0(sK1,sK5) [avatar definition]
33. 3 <=> r1_xboole_0(sK2,sK6) [avatar definition]
37. 4 <=> r1_xboole_0(sK3,sK7) [avatar definition]
40. 1 | 2 | 3 | 4 [avatar split clause 19,37,33,29,25]
42. 5 <=> r1_xboole_0(k4_zfmisc_1(sK0,sK1,sK2,sK3),k4_zfmisc_1(sK4,sK5,sK6,sK7)) [avatar definition]
44. ~r1_xboole_0(k4_zfmisc_1(sK0,sK1,sK2,sK3),k4_zfmisc_1(sK4,sK5,sK6,sK7)) <- (~5) [avatar component clause 42] {a:0,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
45. ~5 [avatar split clause 18,42]
53. r1_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X0,X3)) | ~r1_xboole_0(X0,X1) [resolution 21,23] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. r1_xboole_0(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X0)) | ~r1_xboole_0(X0,X1) [resolution 22,23] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
58. r1_xboole_0(k4_zfmisc_1(X0,X1,X2,X3),k2_zfmisc_1(X4,X5)) | ~r1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2),X4) [superposition 21,20] {a:1,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
60. r1_xboole_0(k4_zfmisc_1(X12,X13,X14,X15),k2_zfmisc_1(X16,X17)) | ~r1_xboole_0(X15,X17) [superposition 22,20] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
68. r1_xboole_0(k4_zfmisc_1(X4,X5,X6,X7),k4_zfmisc_1(X0,X1,X2,X3)) | ~r1_xboole_0(X7,X3) [superposition 60,20] {a:2,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
84. ~r1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(X4,X5),X6),k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2)) | r1_xboole_0(k4_zfmisc_1(X4,X5,X6,X7),k4_zfmisc_1(X0,X1,X2,X3)) [superposition 58,20] {a:2,w:22,wCS:8800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
108. ~r1_xboole_0(sK3,sK7) <- (~5) [resolution 68,44] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
114. ~4 | 5 [avatar split clause 108,42,37]
141. r1_xboole_0(k4_zfmisc_1(X8,X9,X10,X11),k4_zfmisc_1(X12,X13,X14,X15)) | ~r1_xboole_0(k2_zfmisc_1(X12,X13),k2_zfmisc_1(X8,X9)) [resolution 84,53] {a:3,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:5,thDist:-5}
142. r1_xboole_0(k4_zfmisc_1(X16,X17,X18,X19),k4_zfmisc_1(X20,X21,X22,X23)) | ~r1_xboole_0(X18,X22) [resolution 84,22] {a:3,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
165. ~r1_xboole_0(sK2,sK6) <- (~5) [resolution 142,44] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
171. ~3 | 5 [avatar split clause 165,42,33]
172. ~r1_xboole_0(k2_zfmisc_1(sK4,sK5),k2_zfmisc_1(sK0,sK1)) <- (~5) [resolution 141,44] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
179. 9 <=> r1_xboole_0(k2_zfmisc_1(sK4,sK5),k2_zfmisc_1(sK0,sK1)) [avatar definition]
181. ~r1_xboole_0(k2_zfmisc_1(sK4,sK5),k2_zfmisc_1(sK0,sK1)) <- (~9) [avatar component clause 179] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
182. ~9 | 5 [avatar split clause 172,42,179]
187. ~r1_xboole_0(sK1,sK5) <- (~9) [resolution 181,54] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
188. ~r1_xboole_0(sK0,sK4) <- (~9) [resolution 181,53] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
197. ~2 | 9 [avatar split clause 187,179,29]
204. ~1 | 9 [avatar split clause 188,179,25]
205. $false [avatar sat refutation 40,45,114,171,182,197,204] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:136,thDist:-136}
% SZS output end Proof for mcart_1__t60_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% >>> Preprocessing
% Introduced skolems: 8
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 6
% Generated clauses: 146
% Active clauses: 30
% Passive clauses: 58
% Final active clauses: 28
% Final passive clauses: 19
% Main loop iterations started: 39
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% Fw demodulations: 26
% 
% >>> Deletion Inferences
% Forward subsumptions: 48
% Fw demodulations to eq. taut.: 3
% 
% >>> Generating Inferences
% Binary resolution: 32
% Unit resulting resolution: 2
% Forward superposition: 45
% Backward superposition: 16
% Self superposition: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 4
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 1
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

