i: [1,53,0,2,1,1,0,0] t21_yellow_5
a: [3,104,0,2,1,53]
s: 104
i: [1,52,0,2,1,1,0,0] t21_yellow_5
a: [3,109,0,2,1,52]
s: 109
i: [1,51,0,2,1,1,0,0] t21_yellow_5
a: [3,114,0,2,1,51]
s: 114
i: [1,50,0,2,1,1,0,0] t21_yellow_5
a: [3,119,0,2,1,50]
s: 119
i: [1,57,0,8,2,1,0,0] t21_yellow_5
a: [3,80,0,4,1,57]
s: 80
i: [1,55,0,4,1,1,0,0] t21_yellow_5
a: [3,94,0,4,1,55]
s: 94
i: [1,54,0,4,1,1,0,0] t21_yellow_5
a: [3,99,0,4,1,54]
s: 99
i: [1,56,0,7,1,1,0,0] t21_yellow_5
a: [3,89,0,7,1,56]
s: 89
i: [1,73,0,2,1,0,0,1] existence_l1_orders_2
a: [3,124,0,2,1,73]
s: 124
i: [1,68,0,4,1,0,0,1] existence_m1_subset_1
s: 68
i: [1,59,0,6,3,0,0,2] cc1_lattice3
s: 59
i: [1,58,0,7,2,0,0,1] dt_k3_yellow_0
s: 58
i: [1,72,0,17,4,0,0,255] dt_k10_lattice3
s: 72
i: [1,60,0,21,7,0,0,1] t19_yellow_5
s: 60
i: [1,69,0,21,6,0,0,1] dt_k13_lattice3
s: 69
i: [1,70,0,23,6,0,0,255] redefinition_k13_lattice3
s: 70
i: [1,71,0,23,6,0,0,1] commutativity_k13_lattice3
s: 71
i: [1,75,0,28,7,0,0,2] t22_yellow_0
s: 75
d: [2,133,0,21,6,0,21,75,69]
s: 133
i: [1,76,0,28,7,0,0,2] t22_yellow_0
s: 76
d: [2,149,0,21,6,0,21,76,69]
s: 149
d: [2,134,1,19,6,1,44,133,89]
d: [2,139,1,15,5,2,21,134,94]
d: [2,140,1,11,4,3,21,139,99]
d: [2,141,1,9,3,4,21,140,104]
d: [2,142,1,7,2,5,21,141,114]
d: [2,143,1,5,1,6,21,142,119]
a: [3,147,1,5,1,143]
s: 147
i: [1,64,0,40,10,0,0,2] t22_yellow_0
s: 64
d: [2,165,2,16,6,1,40,147,60]
d: [2,166,2,12,5,2,21,165,94]
d: [2,167,2,10,4,3,21,166,104]
d: [2,168,2,8,3,4,21,167,109]
d: [2,169,2,6,2,5,21,168,119]
a: [3,173,2,2,1,169]
s: 173
d: [2,175,3,0,0,3,52,104,114,173,59]
e: 175
d: [2,176,3,4,2,1,40,173,59]
d: [2,178,3,2,1,2,21,176,114]
d: [2,179,3,0,0,3,21,178,104]
e: 179
r: 173
a: [3,83,0,4,1,57]
s: 83
r: 147
r: 89
d: [2,182,0,6,1,2,23,89,83]
a: [3,194,0,6,1,182]
s: 194
a: [3,172,2,2,1,169]
s: 172
i: [1,65,0,40,10,0,0,2] t22_yellow_0
s: 65
d: [2,183,1,3,1,2,44,80,83]
a: [3,199,1,3,1,183]
s: 199
i: [1,66,0,40,10,0,0,2] t22_yellow_0
s: 66
i: [1,67,0,40,10,0,0,2] t22_yellow_0
s: 67
d: [2,181,1,4,1,2,23,147,83]
a: [3,189,1,4,1,181]
s: 189
i: [1,74,0,40,10,0,0,2] t22_yellow_0
s: 74
d: [2,208,1,18,6,1,44,149,194]
d: [2,211,1,14,5,2,21,208,94]
d: [2,212,1,10,4,3,21,211,99]
d: [2,213,1,8,3,4,21,212,104]
d: [2,214,1,6,2,5,21,213,114]
d: [2,215,1,4,1,6,21,214,119]
a: [3,219,1,4,1,215]
s: 219
d: [2,231,0,33,9,0,21,74,69]
s: 231
d: [2,150,1,19,6,1,44,149,89]
d: [2,155,1,15,5,2,21,150,94]
d: [2,156,1,11,4,3,21,155,99]
d: [2,157,1,9,3,4,21,156,104]
d: [2,158,1,7,2,5,21,157,114]
d: [2,159,1,5,1,6,21,158,119]
a: [3,163,1,5,1,159]
s: 163
d: [2,184,1,19,7,1,44,60,83]
d: [2,201,1,17,6,2,21,184,104]
d: [2,202,1,15,5,3,21,201,109]
d: [2,203,1,13,4,4,21,202,119]
a: [3,206,1,11,3,203]
s: 206
d: [2,247,2,0,0,4,52,199,99,219,206]
e: 247
d: [2,249,2,7,2,2,40,206,219]
d: [2,252,2,3,1,3,21,249,99]
d: [2,253,2,0,0,4,21,252,199]
e: 253
f: 253,215,203,183,182,179,169,143,50,51,52,53,54,55,56,57
% Refutation found. Thanks to Tanya!
% SZS status Theorem for yellow_5__t21_yellow_5
% SZS output start Proof for yellow_5__t21_yellow_5
1. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k13_lattice3(X0,X1,X2) = k3_yellow_0(X0) => (k3_yellow_0(X0) = X2 & k3_yellow_0(X0) = X1))))) [input t21_yellow_5]
2. ~! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (k13_lattice3(X0,X1,X2) = k3_yellow_0(X0) => (k3_yellow_0(X0) = X2 & k3_yellow_0(X0) = X1))))) [negated conjecture 1]
3. ! [X0] : (l1_orders_2(X0) => (v1_lattice3(X0) => ~v2_struct_0(X0))) [input cc1_lattice3]
6. ! [X0,X1,X2] : ((m1_subset_1(X2,u1_struct_0(X0)) & m1_subset_1(X1,u1_struct_0(X0)) & l1_orders_2(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0))) [input dt_k13_lattice3]
16. ! [X0] : ((l1_orders_2(X0) & v1_yellow_0(X0) & v5_orders_2(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => (r1_orders_2(X0,X1,k3_yellow_0(X0)) => k3_yellow_0(X0) = X1))) [input t19_yellow_5]
17. ! [X0] : ((l1_orders_2(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (k13_lattice3(X0,X1,X2) = X3 <=> (! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => ((r1_orders_2(X0,X2,X4) & r1_orders_2(X0,X1,X4)) => r1_orders_2(X0,X3,X4))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3))))))) [input t22_yellow_0]
20. ? [X0] : (? [X1] : (? [X2] : (((k3_yellow_0(X0) != X2 | k3_yellow_0(X0) != X1) & k13_lattice3(X0,X1,X2) = k3_yellow_0(X0)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_orders_2(X0) & v1_yellow_0(X0) & v1_lattice3(X0) & v5_orders_2(X0))) [ennf transformation 2]
21. ? [X0] : (? [X1] : (? [X2] : ((k3_yellow_0(X0) != X2 | k3_yellow_0(X0) != X1) & k13_lattice3(X0,X1,X2) = k3_yellow_0(X0) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v1_yellow_0(X0) & v1_lattice3(X0) & v5_orders_2(X0)) [flattening 20]
23. ! [X0] : ((~v2_struct_0(X0) | ~v1_lattice3(X0)) | ~l1_orders_2(X0)) [ennf transformation 3]
24. ! [X0] : (~v2_struct_0(X0) | ~v1_lattice3(X0) | ~l1_orders_2(X0)) [flattening 23]
25. ! [X0] : (! [X1] : ((k3_yellow_0(X0) = X1 | ~r1_orders_2(X0,X1,k3_yellow_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0))) [ennf transformation 16]
26. ! [X0] : (! [X1] : (k3_yellow_0(X0) = X1 | ~r1_orders_2(X0,X1,k3_yellow_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0)) [flattening 25]
27. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k13_lattice3(X0,X1,X2) = X3 <=> (! [X4] : ((r1_orders_2(X0,X3,X4) | (~r1_orders_2(X0,X2,X4) | ~r1_orders_2(X0,X1,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 17]
28. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((k13_lattice3(X0,X1,X2) = X3 <=> (! [X4] : (r1_orders_2(X0,X3,X4) | ~r1_orders_2(X0,X2,X4) | ~r1_orders_2(X0,X1,X4) | ~m1_subset_1(X4,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [flattening 27]
29. ! [X0,X1,X2] : (m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0)) | (~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0))) [ennf transformation 6]
30. ! [X0,X1,X2] : (m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [flattening 29]
37. ? [X0] : (? [X1] : (? [X2] : ((k3_yellow_0(X0) != X2 | k3_yellow_0(X0) != X1) & k13_lattice3(X0,X1,X2) = k3_yellow_0(X0) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_orders_2(X0) & v1_yellow_0(X0) & v1_lattice3(X0) & v5_orders_2(X0)) => (? [X1] : (? [X2] : ((k3_yellow_0(sK0) != X2 | k3_yellow_0(sK0) != X1) & k3_yellow_0(sK0) = k13_lattice3(sK0,X1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v1_yellow_0(sK0) & v1_lattice3(sK0) & v5_orders_2(sK0)) [choice axiom]
38. ? [X1] : (? [X2] : ((k3_yellow_0(sK0) != X2 | k3_yellow_0(sK0) != X1) & k3_yellow_0(sK0) = k13_lattice3(sK0,X1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : ((k3_yellow_0(sK0) != X2 | k3_yellow_0(sK0) != sK1) & k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
39. ? [X2] : ((k3_yellow_0(sK0) != X2 | k3_yellow_0(sK0) != sK1) & k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,X2) & m1_subset_1(X2,u1_struct_0(sK0))) => ((k3_yellow_0(sK0) != sK2 | k3_yellow_0(sK0) != sK1) & k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
40. (((k3_yellow_0(sK0) != sK2 | k3_yellow_0(sK0) != sK1) & k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,sK2) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_orders_2(sK0) & v1_yellow_0(sK0) & v1_lattice3(sK0) & v5_orders_2(sK0) [skolemisation 21,39,38,37]
41. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k13_lattice3(X0,X1,X2) = X3 | (? [X4] : (~r1_orders_2(X0,X3,X4) & r1_orders_2(X0,X2,X4) & r1_orders_2(X0,X1,X4) & m1_subset_1(X4,u1_struct_0(X0))) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X3))) & ((! [X4] : (r1_orders_2(X0,X3,X4) | ~r1_orders_2(X0,X2,X4) | ~r1_orders_2(X0,X1,X4) | ~m1_subset_1(X4,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3)) | k13_lattice3(X0,X1,X2) != X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [nnf transformation 28]
42. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k13_lattice3(X0,X1,X2) = X3 | ? [X4] : (~r1_orders_2(X0,X3,X4) & r1_orders_2(X0,X2,X4) & r1_orders_2(X0,X1,X4) & m1_subset_1(X4,u1_struct_0(X0))) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X3)) & ((! [X4] : (r1_orders_2(X0,X3,X4) | ~r1_orders_2(X0,X2,X4) | ~r1_orders_2(X0,X1,X4) | ~m1_subset_1(X4,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3)) | k13_lattice3(X0,X1,X2) != X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [flattening 41]
43. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k13_lattice3(X0,X1,X2) = X3 | ? [X4] : (~r1_orders_2(X0,X3,X4) & r1_orders_2(X0,X2,X4) & r1_orders_2(X0,X1,X4) & m1_subset_1(X4,u1_struct_0(X0))) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X3)) & ((! [X5] : (r1_orders_2(X0,X3,X5) | ~r1_orders_2(X0,X2,X5) | ~r1_orders_2(X0,X1,X5) | ~m1_subset_1(X5,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3)) | k13_lattice3(X0,X1,X2) != X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [rectify 42]
44. ! [X3,X2,X1,X0] : (? [X4] : (~r1_orders_2(X0,X3,X4) & r1_orders_2(X0,X2,X4) & r1_orders_2(X0,X1,X4) & m1_subset_1(X4,u1_struct_0(X0))) => (~r1_orders_2(X0,X3,sK3(X0,X1,X2,X3)) & r1_orders_2(X0,X2,sK3(X0,X1,X2,X3)) & r1_orders_2(X0,X1,sK3(X0,X1,X2,X3)) & m1_subset_1(sK3(X0,X1,X2,X3),u1_struct_0(X0)))) [choice axiom]
45. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((k13_lattice3(X0,X1,X2) = X3 | (~r1_orders_2(X0,X3,sK3(X0,X1,X2,X3)) & r1_orders_2(X0,X2,sK3(X0,X1,X2,X3)) & r1_orders_2(X0,X1,sK3(X0,X1,X2,X3)) & m1_subset_1(sK3(X0,X1,X2,X3),u1_struct_0(X0))) | ~r1_orders_2(X0,X2,X3) | ~r1_orders_2(X0,X1,X3)) & ((! [X5] : (r1_orders_2(X0,X3,X5) | ~r1_orders_2(X0,X2,X5) | ~r1_orders_2(X0,X1,X5) | ~m1_subset_1(X5,u1_struct_0(X0))) & r1_orders_2(X0,X2,X3) & r1_orders_2(X0,X1,X3)) | k13_lattice3(X0,X1,X2) != X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0)) [skolemisation 43,44]
50. v5_orders_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
51. v1_lattice3(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
52. v1_yellow_0(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
53. l1_orders_2(sK0) [cnf transformation 40] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
54. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 40] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
55. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 40] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
56. k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,sK2) [cnf transformation 40] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
57. k3_yellow_0(sK0) != sK2 | k3_yellow_0(sK0) != sK1 [cnf transformation 40] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
59. ~v2_struct_0(X0) | ~v1_lattice3(X0) | ~l1_orders_2(X0) [cnf transformation 24] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
60. ~r1_orders_2(X0,X1,k3_yellow_0(X0)) | k3_yellow_0(X0) = X1 | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_yellow_0(X0) | ~v5_orders_2(X0) | v2_struct_0(X0) [cnf transformation 26] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
61. r1_orders_2(X0,X1,X3) | k13_lattice3(X0,X1,X2) != X3 | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 45] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:2,thDist:-2}
62. r1_orders_2(X0,X2,X3) | k13_lattice3(X0,X1,X2) != X3 | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 45] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:2,thDist:-2}
69. m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [cnf transformation 30] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
75. r1_orders_2(X0,X2,k13_lattice3(X0,X1,X2)) | ~m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [equality resolution 62] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:2,thDist:-2}
76. r1_orders_2(X0,X1,k13_lattice3(X0,X1,X2)) | ~m1_subset_1(k13_lattice3(X0,X1,X2),u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [equality resolution 61] {a:0,w:28,wCS:11200,sine:2,thAx:0,allAx:2,thDist:-2}
78. 1 <=> k3_yellow_0(sK0) = sK1 [avatar definition]
80. k3_yellow_0(sK0) != sK1 <- (~1) [avatar component clause 78] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
82. 2 <=> k3_yellow_0(sK0) = sK2 [avatar definition]
83. k3_yellow_0(sK0) = sK2 <- (2) [avatar component clause 82] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
85. ~1 | ~2 [avatar split clause 57,82,78]
87. 3 <=> k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,sK2) [avatar definition]
89. k3_yellow_0(sK0) = k13_lattice3(sK0,sK1,sK2) <- (3) [avatar component clause 87] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
90. 3 [avatar split clause 56,87]
92. 4 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
94. m1_subset_1(sK2,u1_struct_0(sK0)) <- (4) [avatar component clause 92] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
95. 4 [avatar split clause 55,92]
97. 5 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
99. m1_subset_1(sK1,u1_struct_0(sK0)) <- (5) [avatar component clause 97] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
100. 5 [avatar split clause 54,97]
102. 6 <=> l1_orders_2(sK0) [avatar definition]
104. l1_orders_2(sK0) <- (6) [avatar component clause 102] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
105. 6 [avatar split clause 53,102]
107. 7 <=> v1_yellow_0(sK0) [avatar definition]
109. v1_yellow_0(sK0) <- (7) [avatar component clause 107] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
110. 7 [avatar split clause 52,107]
112. 8 <=> v1_lattice3(sK0) [avatar definition]
114. v1_lattice3(sK0) <- (8) [avatar component clause 112] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
115. 8 [avatar split clause 51,112]
117. 9 <=> v5_orders_2(sK0) [avatar definition]
119. v5_orders_2(sK0) <- (9) [avatar component clause 117] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
120. 9 [avatar split clause 50,117]
133. r1_orders_2(X0,X2,k13_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [subsumption resolution 75,69] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
134. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (3) [superposition 133,89] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
139. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (3, 4) [subsumption resolution 134,94] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
140. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (3, 4, 5) [subsumption resolution 139,99] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
141. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (3, 4, 5, 6) [subsumption resolution 140,104] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
142. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) | ~v5_orders_2(sK0) <- (3, 4, 5, 6, 8) [subsumption resolution 141,114] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
143. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) <- (3, 4, 5, 6, 8, 9) [subsumption resolution 142,119] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
145. 11 <=> r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) [avatar definition]
147. r1_orders_2(sK0,sK2,k3_yellow_0(sK0)) <- (11) [avatar component clause 145] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
148. 11 | ~3 | ~4 | ~5 | ~6 | ~8 | ~9 [avatar split clause 143,117,112,102,97,92,87,145]
149. r1_orders_2(X0,X1,k13_lattice3(X0,X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_orders_2(X0) | ~v1_lattice3(X0) | ~v5_orders_2(X0) [subsumption resolution 76,69] {a:0,w:21,wCS:8400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
165. k3_yellow_0(sK0) = sK2 | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (11) [resolution 147,60] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
166. k3_yellow_0(sK0) = sK2 | ~l1_orders_2(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (4, 11) [subsumption resolution 165,94] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
167. k3_yellow_0(sK0) = sK2 | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (4, 6, 11) [subsumption resolution 166,104] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
168. k3_yellow_0(sK0) = sK2 | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (4, 6, 7, 11) [subsumption resolution 167,109] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
169. k3_yellow_0(sK0) = sK2 | v2_struct_0(sK0) <- (4, 6, 7, 9, 11) [subsumption resolution 168,119] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
171. 13 <=> v2_struct_0(sK0) [avatar definition]
173. v2_struct_0(sK0) <- (13) [avatar component clause 171] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
174. 13 | 2 | ~4 | ~6 | ~7 | ~9 | ~11 [avatar split clause 169,145,117,107,102,92,82,171]
176. ~v1_lattice3(sK0) | ~l1_orders_2(sK0) <- (13) [resolution 173,59] {a:3,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
178. ~l1_orders_2(sK0) <- (8, 13) [subsumption resolution 176,114] {a:3,w:2,wCS:200,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
179. $false <- (6, 8, 13) [subsumption resolution 178,104] {a:3,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
180. ~6 | ~8 | ~13 [avatar contradiction clause 179]
182. sK2 = k13_lattice3(sK0,sK1,sK2) <- (2, 3) [backward demodulation 89,83] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
183. sK1 != sK2 <- (~1, 2) [superposition 80,83] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
184. ~r1_orders_2(sK0,X0,sK2) | sK2 = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2) [superposition 60,83] {a:1,w:19,wCS:1900,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
192. 15 <=> sK2 = k13_lattice3(sK0,sK1,sK2) [avatar definition]
194. sK2 = k13_lattice3(sK0,sK1,sK2) <- (15) [avatar component clause 192] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
195. 15 | ~2 | ~3 [avatar split clause 182,87,82,192]
197. 16 <=> sK1 = sK2 [avatar definition]
199. sK1 != sK2 <- (~16) [avatar component clause 197] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
200. ~16 | 1 | ~2 [avatar split clause 183,82,78,197]
201. ~r1_orders_2(sK0,X0,sK2) | sK2 = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v1_yellow_0(sK0) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2, 6) [subsumption resolution 184,104] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
202. ~r1_orders_2(sK0,X0,sK2) | sK2 = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | ~v5_orders_2(sK0) | v2_struct_0(sK0) <- (2, 6, 7) [subsumption resolution 201,109] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
203. ~r1_orders_2(sK0,X0,sK2) | sK2 = X0 | ~m1_subset_1(X0,u1_struct_0(sK0)) | v2_struct_0(sK0) <- (2, 6, 7, 9) [subsumption resolution 202,119] {a:1,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
205. 17 <=> ! [X0] : (~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | sK2 = X0) [avatar definition]
206. ~r1_orders_2(sK0,X0,sK2) | ~m1_subset_1(X0,u1_struct_0(sK0)) | sK2 = X0 <- (17) [avatar component clause 205] {a:1,w:11,wCS:1100,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
207. 13 | 17 | ~2 | ~6 | ~7 | ~9 [avatar split clause 203,117,107,102,82,205,171]
208. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (15) [superposition 149,194] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
211. r1_orders_2(sK0,sK1,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (4, 15) [subsumption resolution 208,94] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
212. r1_orders_2(sK0,sK1,sK2) | ~l1_orders_2(sK0) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (4, 5, 15) [subsumption resolution 211,99] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
213. r1_orders_2(sK0,sK1,sK2) | ~v1_lattice3(sK0) | ~v5_orders_2(sK0) <- (4, 5, 6, 15) [subsumption resolution 212,104] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
214. r1_orders_2(sK0,sK1,sK2) | ~v5_orders_2(sK0) <- (4, 5, 6, 8, 15) [subsumption resolution 213,114] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
215. r1_orders_2(sK0,sK1,sK2) <- (4, 5, 6, 8, 9, 15) [subsumption resolution 214,119] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
217. 18 <=> r1_orders_2(sK0,sK1,sK2) [avatar definition]
219. r1_orders_2(sK0,sK1,sK2) <- (18) [avatar component clause 217] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
220. 18 | ~4 | ~5 | ~6 | ~8 | ~9 | ~15 [avatar split clause 215,192,117,112,102,97,92,217]
249. ~m1_subset_1(sK1,u1_struct_0(sK0)) | sK1 = sK2 <- (17, 18) [resolution 206,219] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
252. sK1 = sK2 <- (5, 17, 18) [subsumption resolution 249,99] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
253. $false <- (5, ~16, 17, 18) [subsumption resolution 252,199] {a:2,w:0,goal:1,sine:255,thAx:0,allAx:11,thDist:-11}
254. ~5 | 16 | ~17 | ~18 [avatar contradiction clause 253]
259. $false [avatar sat refutation 85,90,95,100,105,110,115,120,148,174,180,195,200,207,220,254] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:206,thDist:-206}
% SZS output end Proof for yellow_5__t21_yellow_5
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 24
% Generated clauses: 132
% Active clauses: 33
% Passive clauses: 43
% Final active clauses: 30
% Final passive clauses: 6
% Main loop iterations started: 37
% 
% >>> Simplifying Inferences
% Duplicate literals: 68
% Fw subsumption resolutions: 43
% Fw demodulations: 1
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 10
% 
% >>> Generating Inferences
% Binary resolution: 9
% Unit resulting resolution: 2
% Forward superposition: 10
% Backward superposition: 9
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 24
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% resolution: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

