i: [1,814,0,47,10,1,0,0] t3_gate_2
d: [2,1256,0,35,7,0,18,814]
a: [3,1280,0,4,1,1256]
s: 1280
a: [3,1271,0,7,1,1256]
s: 1271
i: [1,815,0,5,2,0,0,1] fc2_gate_1
s: 815
i: [1,816,0,5,2,0,0,1] fc1_gate_1
s: 816
i: [1,817,0,5,2,0,0,255] t6_boole
s: 817
i: [1,821,0,6,2,0,0,1] fc5_gate_1
s: 821
r: 1280
a: [3,1279,0,4,1,1256]
s: 1279
a: [3,1284,0,4,1,1256]
s: 1284
r: 1284
a: [3,1283,0,4,1,1256]
s: 1283
a: [3,1268,0,6,1,1256]
s: 1268
i: [1,822,0,6,2,0,0,1] fc4_gate_1
s: 822
d: [2,1347,1,2,1,1,40,821,1280]
a: [3,1351,1,2,1,1347]
s: 1351
i: [1,818,0,7,3,0,0,2] t8_boole
s: 818
i: [1,819,0,7,1,0,0,1] commutativity_k2_gate_1
s: 819
i: [1,820,0,7,1,0,0,1] commutativity_k3_gate_1
s: 820
r: 1268
r: 1271
r: 1283
r: 1279
a: [3,1263,0,4,1,1256]
s: 1263
d: [2,1361,0,4,1,1,23,1283,820]
a: [3,1378,0,4,1,1361]
s: 1378
d: [2,1362,0,4,1,1,23,1279,820]
a: [3,1383,0,4,1,1362]
s: 1383
a: [3,1260,0,6,1,1256]
s: 1260
d: [2,1359,0,6,1,1,23,1268,820]
a: [3,1367,0,6,1,1359]
s: 1367
d: [2,1360,0,7,1,1,23,1271,820]
a: [3,1373,0,7,1,1360]
s: 1373
d: [2,1387,1,3,1,1,40,1260,821]
a: [3,1391,1,3,1,1387]
s: 1391
d: [2,1393,2,0,0,2,52,1351,1391,816]
e: 1393
d: [2,1394,2,2,1,1,40,1391,816]
d: [2,1396,2,0,0,2,21,1394,1351]
e: 1396
r: 1351
i: [1,825,0,7,2,0,0,1] fc18_gate_1
s: 825
a: [3,1350,1,2,1,1347]
s: 1350
i: [1,826,0,7,2,0,0,1] fc20_gate_1
s: 826
i: [1,827,0,7,2,0,0,1] fc19_gate_1
s: 827
d: [2,1357,1,6,2,0,44,822,819]
s: 1357
i: [1,823,0,8,3,0,0,1] fc3_gate_1
s: 823
i: [1,622,0,47,11,1,0,0] t3_gate_2
d: [2,1104,0,39,9,0,18,622]
a: [3,1299,0,2,1,1104]
s: 1299
i: [1,824,0,8,3,0,0,1] fc6_gate_1
s: 824
i: [1,828,0,11,4,0,0,1] fc17_gate_1
s: 828
d: [2,1358,1,6,2,0,44,822,819]
s: 1358
d: [2,1363,1,6,2,0,44,821,820]
s: 1363
r: 1260
r: 1350
r: 1391
d: [2,1443,0,4,1,1,22,1284,820]
d: [2,1444,0,0,0,2,21,1443,1378]
e: 1444
r: 1378
a: [3,1295,0,2,1,1104]
s: 1295
a: [3,1377,0,4,1,1361]
s: 1377
d: [2,1447,1,0,0,2,52,1295,1377,821]
e: 1447
d: [2,1449,1,2,1,1,40,1377,821]
d: [2,1451,1,0,0,2,21,1449,1295]
e: 1451
r: 1295
a: [3,1296,0,2,1,1104]
s: 1296
s: 1351
a: [3,1259,0,6,1,1256]
s: 1259
d: [2,1412,1,4,2,1,40,824,1383]
a: [3,1430,1,2,1,1412]
s: 1430
a: [3,1390,1,3,1,1387]
s: 1390
d: [2,1409,1,6,2,1,40,824,1367]
a: [3,1441,1,4,1,1409]
s: 1441
d: [2,1458,2,0,0,2,52,1296,1441,1357]
e: 1458
d: [2,1459,2,0,0,2,52,1299,1441,822]
e: 1459
d: [2,1460,2,2,1,1,40,1441,1357]
d: [2,1464,2,0,0,2,21,1460,1296]
e: 1464
d: [2,1461,2,2,1,1,40,1441,822]
d: [2,1466,2,0,0,2,21,1461,1299]
e: 1466
r: 1299
r: 1367
r: 1441
d: [2,1410,1,7,2,1,40,824,1373]
a: [3,1436,1,5,1,1410]
s: 1436
d: [2,1470,2,0,0,2,52,1296,1436,827]
e: 1470
d: [2,1471,2,2,1,1,40,1436,827]
d: [2,1475,2,0,0,2,21,1471,1296]
e: 1475
r: 1296
r: 1377
s: 1295
a: [3,1300,0,2,1,1104]
s: 1300
s: 1378
a: [3,1368,0,6,1,1359]
s: 1368
a: [3,1440,1,4,1,1409]
s: 1440
d: [2,1485,2,0,0,3,52,1300,1295,1440,823]
e: 1485
d: [2,1486,2,4,2,1,40,1440,823]
d: [2,1488,2,2,1,2,21,1486,1295]
d: [2,1489,2,0,0,3,21,1488,1300]
e: 1489
r: 1368
r: 1440
d: [2,1491,0,4,1,1,22,1280,820]
d: [2,1492,0,0,0,2,21,1491,1383]
e: 1492
r: 1383
a: [3,1382,0,4,1,1362]
s: 1382
d: [2,1497,2,0,0,2,52,1430,1382,821]
e: 1497
d: [2,1499,1,2,1,1,40,1382,821]
d: [2,1501,1,0,0,2,21,1499,1430]
e: 1501
r: 1430
r: 1300
r: 1373
r: 1436
a: [3,1429,1,2,1,1412]
s: 1429
a: [3,1275,0,4,1,1256]
s: 1275
r: 1263
r: 1275
a: [3,1264,0,4,1,1256]
s: 1264
d: [2,1512,2,0,0,2,52,1351,1264,1357]
e: 1512
d: [2,1513,2,2,1,1,40,1264,1357]
d: [2,1516,2,0,0,2,21,1513,1351]
e: 1516
r: 1264
a: [3,1276,0,4,1,1256]
s: 1276
r: 1276
r: 1382
r: 1429
s: 1300
s: 1430
s: 1275
s: 1383
s: 1441
s: 1367
a: [3,1435,1,5,1,1410]
s: 1435
d: [2,1530,2,0,0,4,52,1300,1295,1430,1435,828]
e: 1530
d: [2,1531,2,6,3,1,40,1435,828]
d: [2,1533,2,4,2,2,21,1531,1430]
d: [2,1534,2,2,1,3,21,1533,1295]
d: [2,1535,2,0,0,4,21,1534,1300]
e: 1535
r: 1435
a: [3,1372,0,7,1,1360]
s: 1372
r: 1300
r: 1351
r: 1441
r: 1275
r: 1372
r: 1390
s: 1299
s: 1350
s: 1276
d: [2,1546,2,0,0,2,52,1350,1276,821]
e: 1546
d: [2,1548,1,2,1,1,40,1276,821]
d: [2,1550,1,0,0,2,21,1548,1350]
e: 1550
r: 1259
r: 1276
s: 1391
s: 1263
s: 1260
d: [2,1418,2,3,1,1,40,1363,1260]
a: [3,1425,2,3,1,1418]
s: 1425
d: [2,1557,3,0,0,2,52,1299,1425,816]
e: 1557
d: [2,1558,3,2,1,1,40,1425,816]
d: [2,1560,3,0,0,2,21,1558,1299]
e: 1560
r: 1260
r: 1263
r: 1295
r: 1299
r: 1425
s: 1264
s: 1300
s: 1275
s: 1259
a: [3,1424,2,3,1,1418]
s: 1424
d: [2,1569,3,0,0,2,52,1300,1424,815]
e: 1569
d: [2,1570,3,2,1,1,40,1424,815]
d: [2,1572,3,0,0,2,21,1570,1300]
e: 1572
i: [1,811,0,47,10,1,0,0] t3_gate_2
d: [2,1253,0,35,7,0,18,811]
i: [1,694,0,39,8,1,0,0] t3_gate_2
d: [2,1148,0,35,7,0,18,694]
i: [1,270,0,35,8,1,0,0] t3_gate_2
d: [2,972,0,31,7,0,18,270]
i: [1,50,0,35,7,1,0,0] t3_gate_2
d: [2,832,0,27,5,0,18,50]
d: [2,1353,1,2,1,1,40,1284,821]
d: [2,1355,1,2,1,1,40,1268,821]
d: [2,1403,1,4,2,1,40,823,1263]
d: [2,1406,1,2,1,2,21,1403,1350]
d: [2,1411,1,4,2,1,40,824,1378]
d: [2,1478,0,4,1,1,22,1283,820]
d: [2,1473,2,2,1,1,40,1436,825]
d: [2,1483,1,4,1,1,40,1368,821]
d: [2,1472,2,2,1,1,40,1436,826]
a: [3,1272,0,7,1,1256]
d: [2,1505,0,7,1,1,22,1272,820]
d: [2,1507,0,6,1,1,22,1268,820]
a: [3,1267,0,6,1,1256]
d: [2,1518,0,6,1,1,22,1267,820]
d: [2,1520,2,2,1,1,40,1276,1363]
d: [2,1523,0,4,1,1,22,1279,820]
d: [2,1538,1,5,1,1,40,1372,821]
d: [2,1525,1,4,2,1,40,1275,824]
d: [2,1537,2,2,1,1,40,1372,1363]
d: [2,1457,2,2,1,1,40,1390,815]
d: [2,1529,1,6,2,1,40,1367,824]
d: [2,1544,0,7,1,1,22,1271,820]
d: [2,1554,2,3,1,1,40,1260,1363]
d: [2,1563,1,2,1,1,40,1264,822]
d: [2,1566,1,6,2,1,40,1259,824]
d: [2,1567,1,3,1,2,21,1566,1391]
f: 1572,1567,1563,1560,1554,1550,1544,1529,1457,1537,1525,1538,1535,1523,1520,1518,1516,1507,1505,1472,1499,1491,1489,1483,1473,1478,1471,1466,1464,1449,1443,1410,1411,1412,1406,1394,1387,1355,1353,1347,832,972,1148,1253,1256
% Refutation found. Thanks to Tanya!
% SZS status Theorem for gate_2__t3_gate_2
% SZS output start Proof for gate_2__t3_gate_2
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k3_gate_1(X2,X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) => (~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) => (~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) => v1_xboole_0(k2_gate_1(X0,X1))) & ~(v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) & v1_xboole_0(k2_gate_1(X0,X1)))) [input t3_gate_2]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : (~(v1_xboole_0(k3_gate_1(X2,X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) => (~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) => (~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) => v1_xboole_0(k2_gate_1(X0,X1))) & ~(v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) & v1_xboole_0(k2_gate_1(X0,X1)))) [negated conjecture 1]
4. ! [X0,X1] : k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [input commutativity_k2_gate_1]
5. ! [X0,X1] : k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [input commutativity_k3_gate_1]
11. ! [X0,X1,X2] : ((~v1_xboole_0(X2) & ~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc17_gate_1]
12. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X0,X1,X2))) [input fc18_gate_1]
13. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X0,X2))) [input fc19_gate_1]
14. ! [X0] : (v1_xboole_0(X0) => ~v1_xboole_0(k1_gate_1(X0))) [input fc1_gate_1]
15. ! [X0,X1,X2] : (v1_xboole_0(X0) => v1_xboole_0(k8_gate_1(X1,X2,X0))) [input fc20_gate_1]
16. ! [X0] : (~v1_xboole_0(X0) => v1_xboole_0(k1_gate_1(X0))) [input fc2_gate_1]
17. ! [X0,X1] : ((~v1_xboole_0(X1) & ~v1_xboole_0(X0)) => ~v1_xboole_0(k2_gate_1(X0,X1))) [input fc3_gate_1]
18. ! [X0,X1] : (v1_xboole_0(X0) => v1_xboole_0(k2_gate_1(X0,X1))) [input fc4_gate_1]
19. ! [X0,X1] : (~v1_xboole_0(X1) => ~v1_xboole_0(k3_gate_1(X0,X1))) [input fc5_gate_1]
20. ! [X0,X1] : ((v1_xboole_0(X1) & v1_xboole_0(X0)) => v1_xboole_0(k3_gate_1(X0,X1))) [input fc6_gate_1]
24. ~! [X0,X1,X2,X3] : (~(v1_xboole_0(k3_gate_1(X2,X1)) & ~(v1_xboole_0(X0) & ~v1_xboole_0(X2)) & ~(v1_xboole_0(X2) & ~v1_xboole_0(X0)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) => (~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) => (~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1)))) & ~(v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) & (~v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) => v1_xboole_0(k2_gate_1(X0,X1))) & ~(v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) & v1_xboole_0(k2_gate_1(X0,X1)))) [flattening 2]
29. ? [X0,X1,X2,X3] : ((v1_xboole_0(k3_gate_1(X2,X1)) & (~v1_xboole_0(X0) | v1_xboole_0(X2)) & (~v1_xboole_0(X2) | v1_xboole_0(X0)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | ((v1_xboole_0(k3_gate_1(X3,X1)) | v1_xboole_0(k3_gate_1(X2,X1)) | v1_xboole_0(k3_gate_1(X0,X1))) & ~v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1))) | (v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | ((v1_xboole_0(k3_gate_1(X2,X1)) | v1_xboole_0(k3_gate_1(X0,X1))) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | (~v1_xboole_0(k2_gate_1(X0,X1)) & ~v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1)))) | (v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) & v1_xboole_0(k2_gate_1(X0,X1)))) [ennf transformation 24]
30. ! [X0] : (v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0)) [ennf transformation 16]
31. ! [X0] : (~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0)) [ennf transformation 14]
34. ! [X0,X1] : (~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1)) [ennf transformation 19]
35. ! [X0,X1] : (v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0)) [ennf transformation 18]
36. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | (v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 17]
37. ! [X0,X1] : (~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 36]
38. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | (~v1_xboole_0(X1) | ~v1_xboole_0(X0))) [ennf transformation 20]
39. ! [X0,X1] : (v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0)) [flattening 38]
40. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0)) [ennf transformation 12]
41. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0)) [ennf transformation 15]
42. ! [X0,X1,X2] : (v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0)) [ennf transformation 13]
43. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | (v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0))) [ennf transformation 11]
44. ! [X0,X1,X2] : (~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0)) [flattening 43]
45. ? [X0,X1,X2,X3] : ((v1_xboole_0(k3_gate_1(X2,X1)) & (~v1_xboole_0(X0) | v1_xboole_0(X2)) & (~v1_xboole_0(X2) | v1_xboole_0(X0)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | ((v1_xboole_0(k3_gate_1(X3,X1)) | v1_xboole_0(k3_gate_1(X2,X1)) | v1_xboole_0(k3_gate_1(X0,X1))) & ~v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1))) | (v1_xboole_0(k3_gate_1(k8_gate_1(X0,X2,X3),X1)) & ~v1_xboole_0(k3_gate_1(X3,X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | ((v1_xboole_0(k3_gate_1(X2,X1)) | v1_xboole_0(k3_gate_1(X0,X1))) & ~v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1))) | (v1_xboole_0(k3_gate_1(k2_gate_1(X0,X2),X1)) & ~v1_xboole_0(k3_gate_1(X2,X1)) & ~v1_xboole_0(k3_gate_1(X0,X1))) | (~v1_xboole_0(k2_gate_1(X0,X1)) & ~v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1)))) | (v1_xboole_0(k3_gate_1(k1_gate_1(X0),k1_gate_1(X1))) & v1_xboole_0(k2_gate_1(X0,X1)))) => ((v1_xboole_0(k3_gate_1(sK2,sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK2) | v1_xboole_0(sK0)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | ((v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1))) & ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1))) | (v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) & ~v1_xboole_0(k3_gate_1(sK3,sK1)) & ~v1_xboole_0(k3_gate_1(sK2,sK1)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | ((v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1))) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) & ~v1_xboole_0(k3_gate_1(sK2,sK1)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | (~v1_xboole_0(k2_gate_1(sK0,sK1)) & ~v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1)))) | (v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) & v1_xboole_0(k2_gate_1(sK0,sK1)))) [choice axiom]
46. (v1_xboole_0(k3_gate_1(sK2,sK1)) & (~v1_xboole_0(sK0) | v1_xboole_0(sK2)) & (~v1_xboole_0(sK2) | v1_xboole_0(sK0)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | ((v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1))) & ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1))) | (v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) & ~v1_xboole_0(k3_gate_1(sK3,sK1)) & ~v1_xboole_0(k3_gate_1(sK2,sK1)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | ((v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1))) & ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1))) | (v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) & ~v1_xboole_0(k3_gate_1(sK2,sK1)) & ~v1_xboole_0(k3_gate_1(sK0,sK1))) | (~v1_xboole_0(k2_gate_1(sK0,sK1)) & ~v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1)))) | (v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) & v1_xboole_0(k2_gate_1(sK0,sK1))) [skolemisation 29,45]
50. ~v1_xboole_0(k3_gate_1(sK0,sK1)) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(sK0,sK1)) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k3_gate_1(sK0,sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [cnf transformation 46] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
270. ~v1_xboole_0(sK2) | v1_xboole_0(sK0) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(sK2,sK1)) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k3_gate_1(sK2,sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [cnf transformation 46] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
694. v1_xboole_0(k3_gate_1(sK2,sK1)) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [cnf transformation 46] {a:0,w:39,wCS:3900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
811. v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) | v1_xboole_0(k2_gate_1(sK0,sK1)) [cnf transformation 46] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
814. v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [cnf transformation 46] {a:0,w:47,wCS:4700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
815. v1_xboole_0(k1_gate_1(X0)) | v1_xboole_0(X0) [cnf transformation 30] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
816. ~v1_xboole_0(k1_gate_1(X0)) | ~v1_xboole_0(X0) [cnf transformation 31] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
819. k2_gate_1(X0,X1) = k2_gate_1(X1,X0) [cnf transformation 4] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
820. k3_gate_1(X0,X1) = k3_gate_1(X1,X0) [cnf transformation 5] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
821. ~v1_xboole_0(k3_gate_1(X0,X1)) | v1_xboole_0(X1) [cnf transformation 34] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
822. v1_xboole_0(k2_gate_1(X0,X1)) | ~v1_xboole_0(X0) [cnf transformation 35] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
823. ~v1_xboole_0(k2_gate_1(X0,X1)) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 37] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
824. v1_xboole_0(k3_gate_1(X0,X1)) | ~v1_xboole_0(X1) | ~v1_xboole_0(X0) [cnf transformation 39] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
825. v1_xboole_0(k8_gate_1(X0,X1,X2)) | ~v1_xboole_0(X0) [cnf transformation 40] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
826. v1_xboole_0(k8_gate_1(X1,X2,X0)) | ~v1_xboole_0(X0) [cnf transformation 41] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
827. v1_xboole_0(k8_gate_1(X1,X0,X2)) | ~v1_xboole_0(X0) [cnf transformation 42] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
828. ~v1_xboole_0(k8_gate_1(X0,X1,X2)) | v1_xboole_0(X2) | v1_xboole_0(X1) | v1_xboole_0(X0) [cnf transformation 44] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
832. ~v1_xboole_0(k3_gate_1(sK0,sK1)) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [duplicate literal removal 50] {a:0,w:27,wCS:2700,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
972. ~v1_xboole_0(sK2) | v1_xboole_0(sK0) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(sK2,sK1)) | ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [duplicate literal removal 270] {a:0,w:31,wCS:3100,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1148. v1_xboole_0(k3_gate_1(sK2,sK1)) | ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | ~v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [duplicate literal removal 694] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1253. v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) | v1_xboole_0(k2_gate_1(sK0,sK1)) [duplicate literal removal 811] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1256. v1_xboole_0(k3_gate_1(sK2,sK1)) | v1_xboole_0(k3_gate_1(sK3,sK1)) | v1_xboole_0(k3_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) | v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) | ~v1_xboole_0(k2_gate_1(sK0,sK1)) | v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [duplicate literal removal 814] {a:0,w:35,wCS:3500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
1258. 1 <=> v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) [avatar definition]
1259. ~v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) <- (~1) [avatar component clause 1258] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1260. v1_xboole_0(k3_gate_1(k1_gate_1(sK0),k1_gate_1(sK1))) <- (1) [avatar component clause 1258] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1262. 2 <=> v1_xboole_0(k2_gate_1(sK0,sK1)) [avatar definition]
1263. v1_xboole_0(k2_gate_1(sK0,sK1)) <- (2) [avatar component clause 1262] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1264. ~v1_xboole_0(k2_gate_1(sK0,sK1)) <- (~2) [avatar component clause 1262] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1266. 3 <=> v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) [avatar definition]
1267. ~v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) <- (~3) [avatar component clause 1266] {a:0,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1268. v1_xboole_0(k3_gate_1(k2_gate_1(sK0,sK2),sK1)) <- (3) [avatar component clause 1266] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1270. 4 <=> v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) [avatar definition]
1271. ~v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) <- (~4) [avatar component clause 1270] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1272. v1_xboole_0(k3_gate_1(k8_gate_1(sK0,sK2,sK3),sK1)) <- (4) [avatar component clause 1270] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1274. 5 <=> v1_xboole_0(k3_gate_1(sK0,sK1)) [avatar definition]
1275. ~v1_xboole_0(k3_gate_1(sK0,sK1)) <- (~5) [avatar component clause 1274] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1276. v1_xboole_0(k3_gate_1(sK0,sK1)) <- (5) [avatar component clause 1274] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1278. 6 <=> v1_xboole_0(k3_gate_1(sK3,sK1)) [avatar definition]
1279. ~v1_xboole_0(k3_gate_1(sK3,sK1)) <- (~6) [avatar component clause 1278] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1280. v1_xboole_0(k3_gate_1(sK3,sK1)) <- (6) [avatar component clause 1278] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1282. 7 <=> v1_xboole_0(k3_gate_1(sK2,sK1)) [avatar definition]
1283. ~v1_xboole_0(k3_gate_1(sK2,sK1)) <- (~7) [avatar component clause 1282] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1284. v1_xboole_0(k3_gate_1(sK2,sK1)) <- (7) [avatar component clause 1282] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1285. 1 | ~2 | 3 | 4 | 5 | 6 | 7 [avatar split clause 1256,1282,1278,1274,1270,1266,1262,1258]
1286. 2 | ~1 | 3 | 4 | 5 | 6 | 7 [avatar split clause 1253,1282,1278,1274,1270,1266,1258,1262]
1287. 1 | ~2 | 3 | 5 | ~6 | ~4 | 7 [avatar split clause 1148,1282,1270,1278,1274,1266,1262,1258]
1294. 8 <=> v1_xboole_0(sK2) [avatar definition]
1295. ~v1_xboole_0(sK2) <- (~8) [avatar component clause 1294] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1296. v1_xboole_0(sK2) <- (8) [avatar component clause 1294] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1298. 9 <=> v1_xboole_0(sK0) [avatar definition]
1299. v1_xboole_0(sK0) <- (9) [avatar component clause 1298] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1300. ~v1_xboole_0(sK0) <- (~9) [avatar component clause 1298] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1325. 1 | ~2 | ~3 | ~7 | ~4 | 9 | ~8 [avatar split clause 972,1294,1298,1270,1282,1266,1262,1258]
1343. 1 | ~2 | ~3 | ~4 | ~5 [avatar split clause 832,1274,1270,1266,1262,1258]
1347. v1_xboole_0(sK1) <- (6) [resolution 821,1280] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1349. 10 <=> v1_xboole_0(sK1) [avatar definition]
1350. ~v1_xboole_0(sK1) <- (~10) [avatar component clause 1349] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1351. v1_xboole_0(sK1) <- (10) [avatar component clause 1349] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1352. 10 | ~6 [avatar split clause 1347,1278,1349]
1353. v1_xboole_0(sK1) <- (7) [resolution 1284,821] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1354. 10 | ~7 [avatar split clause 1353,1282,1349]
1355. v1_xboole_0(sK1) <- (3) [resolution 1268,821] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1356. 10 | ~3 [avatar split clause 1355,1266,1349]
1357. v1_xboole_0(k2_gate_1(X1,X0)) | ~v1_xboole_0(X0) [superposition 822,819] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1363. ~v1_xboole_0(k3_gate_1(X1,X0)) | v1_xboole_0(X1) [superposition 821,820] {a:1,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1366. 11 <=> v1_xboole_0(k3_gate_1(sK1,k2_gate_1(sK0,sK2))) [avatar definition]
1367. ~v1_xboole_0(k3_gate_1(sK1,k2_gate_1(sK0,sK2))) <- (~11) [avatar component clause 1366] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1368. v1_xboole_0(k3_gate_1(sK1,k2_gate_1(sK0,sK2))) <- (11) [avatar component clause 1366] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1371. 12 <=> v1_xboole_0(k3_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) [avatar definition]
1372. v1_xboole_0(k3_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (12) [avatar component clause 1371] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1373. ~v1_xboole_0(k3_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (~12) [avatar component clause 1371] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1376. 13 <=> v1_xboole_0(k3_gate_1(sK1,sK2)) [avatar definition]
1377. v1_xboole_0(k3_gate_1(sK1,sK2)) <- (13) [avatar component clause 1376] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1378. ~v1_xboole_0(k3_gate_1(sK1,sK2)) <- (~13) [avatar component clause 1376] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1381. 14 <=> v1_xboole_0(k3_gate_1(sK1,sK3)) [avatar definition]
1382. v1_xboole_0(k3_gate_1(sK1,sK3)) <- (14) [avatar component clause 1381] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1383. ~v1_xboole_0(k3_gate_1(sK1,sK3)) <- (~14) [avatar component clause 1381] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
1387. v1_xboole_0(k1_gate_1(sK1)) <- (1) [resolution 1260,821] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1389. 15 <=> v1_xboole_0(k1_gate_1(sK1)) [avatar definition]
1390. ~v1_xboole_0(k1_gate_1(sK1)) <- (~15) [avatar component clause 1389] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1391. v1_xboole_0(k1_gate_1(sK1)) <- (15) [avatar component clause 1389] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1392. 15 | ~1 [avatar split clause 1387,1258,1389]
1394. ~v1_xboole_0(sK1) <- (15) [resolution 1391,816] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1398. ~10 | ~15 [avatar split clause 1394,1389,1349]
1403. v1_xboole_0(sK1) | v1_xboole_0(sK0) <- (2) [resolution 823,1263] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1406. v1_xboole_0(sK0) <- (2, ~10) [subsumption resolution 1403,1350] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1407. 9 | ~2 | 10 [avatar split clause 1406,1349,1262,1298]
1410. ~v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) | ~v1_xboole_0(sK1) <- (~12) [resolution 824,1373] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1411. ~v1_xboole_0(sK2) | ~v1_xboole_0(sK1) <- (~13) [resolution 824,1378] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1412. ~v1_xboole_0(sK3) | ~v1_xboole_0(sK1) <- (~14) [resolution 824,1383] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1423. 16 <=> v1_xboole_0(k1_gate_1(sK0)) [avatar definition]
1424. ~v1_xboole_0(k1_gate_1(sK0)) <- (~16) [avatar component clause 1423] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1425. v1_xboole_0(k1_gate_1(sK0)) <- (16) [avatar component clause 1423] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1428. 17 <=> v1_xboole_0(sK3) [avatar definition]
1430. ~v1_xboole_0(sK3) <- (~17) [avatar component clause 1428] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1431. ~10 | ~17 | 14 [avatar split clause 1412,1381,1428,1349]
1432. ~10 | ~8 | 13 [avatar split clause 1411,1376,1294,1349]
1434. 18 <=> v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) [avatar definition]
1435. v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (18) [avatar component clause 1434] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1436. ~v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (~18) [avatar component clause 1434] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1437. ~10 | ~18 | 12 [avatar split clause 1410,1371,1434,1349]
1439. 19 <=> v1_xboole_0(k2_gate_1(sK0,sK2)) [avatar definition]
1440. v1_xboole_0(k2_gate_1(sK0,sK2)) <- (19) [avatar component clause 1439] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1441. ~v1_xboole_0(k2_gate_1(sK0,sK2)) <- (~19) [avatar component clause 1439] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
1443. v1_xboole_0(k3_gate_1(sK1,sK2)) <- (7) [forward demodulation 1284,820] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1446. 13 | ~7 [avatar split clause 1443,1282,1376]
1449. v1_xboole_0(sK2) <- (13) [resolution 1377,821] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1453. 8 | ~13 [avatar split clause 1449,1376,1294]
1457. v1_xboole_0(sK1) <- (~15) [resolution 1390,815] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1460. ~v1_xboole_0(sK2) <- (~19) [resolution 1441,1357] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1461. ~v1_xboole_0(sK0) <- (~19) [resolution 1441,822] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1464. $false <- (8, ~19) [subsumption resolution 1460,1296] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1465. ~8 | 19 [avatar contradiction clause 1464]
1466. $false <- (9, ~19) [subsumption resolution 1461,1299] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1467. ~9 | 19 [avatar contradiction clause 1466]
1471. ~v1_xboole_0(sK2) <- (~18) [resolution 1436,827] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1472. ~v1_xboole_0(sK3) <- (~18) [resolution 1436,826] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1473. ~v1_xboole_0(sK0) <- (~18) [resolution 1436,825] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1477. ~8 | 18 [avatar split clause 1471,1434,1294]
1478. ~v1_xboole_0(k3_gate_1(sK1,sK2)) <- (~7) [forward demodulation 1283,820] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1479. ~13 | 7 [avatar split clause 1478,1282,1376]
1480. ~9 | 18 [avatar split clause 1473,1434,1298]
1483. v1_xboole_0(k2_gate_1(sK0,sK2)) <- (11) [resolution 1368,821] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1484. 19 | ~11 [avatar split clause 1483,1366,1439]
1486. v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (19) [resolution 1440,823] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1488. v1_xboole_0(sK0) <- (~8, 19) [subsumption resolution 1486,1295] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1489. $false <- (~8, ~9, 19) [subsumption resolution 1488,1300] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1490. 8 | 9 | ~19 [avatar contradiction clause 1489]
1491. v1_xboole_0(k3_gate_1(sK1,sK3)) <- (6) [forward demodulation 1280,820] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1495. 14 | ~6 [avatar split clause 1491,1278,1381]
1499. v1_xboole_0(sK3) <- (14) [resolution 1382,821] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1503. 17 | ~14 [avatar split clause 1499,1381,1428]
1504. ~17 | 18 [avatar split clause 1472,1434,1428]
1505. v1_xboole_0(k3_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (4) [forward demodulation 1272,820] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1506. 12 | ~4 [avatar split clause 1505,1270,1371]
1507. v1_xboole_0(k3_gate_1(sK1,k2_gate_1(sK0,sK2))) <- (3) [forward demodulation 1268,820] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1508. 11 | ~3 [avatar split clause 1507,1266,1366]
1513. ~v1_xboole_0(sK1) <- (~2) [resolution 1264,1357] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1516. $false <- (~2, 10) [subsumption resolution 1513,1351] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1517. 2 | ~10 [avatar contradiction clause 1516]
1518. ~v1_xboole_0(k3_gate_1(sK1,k2_gate_1(sK0,sK2))) <- (~3) [forward demodulation 1267,820] {a:0,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1519. ~11 | 3 [avatar split clause 1518,1266,1366]
1520. v1_xboole_0(sK0) <- (5) [resolution 1276,1363] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1522. 9 | ~5 [avatar split clause 1520,1274,1298]
1523. ~v1_xboole_0(k3_gate_1(sK1,sK3)) <- (~6) [forward demodulation 1279,820] {a:0,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1524. ~14 | 6 [avatar split clause 1523,1278,1381]
1525. ~v1_xboole_0(sK1) | ~v1_xboole_0(sK0) <- (~5) [resolution 1275,824] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1529. ~v1_xboole_0(k2_gate_1(sK0,sK2)) | ~v1_xboole_0(sK1) <- (~11) [resolution 1367,824] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1531. v1_xboole_0(sK3) | v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (18) [resolution 1435,828] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1533. v1_xboole_0(sK2) | v1_xboole_0(sK0) <- (~17, 18) [subsumption resolution 1531,1430] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1534. v1_xboole_0(sK0) <- (~8, ~17, 18) [subsumption resolution 1533,1295] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1535. $false <- (~8, ~9, ~17, 18) [subsumption resolution 1534,1300] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1536. 8 | 9 | 17 | ~18 [avatar contradiction clause 1535]
1537. v1_xboole_0(sK1) <- (12) [resolution 1372,1363] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1538. v1_xboole_0(k8_gate_1(sK0,sK2,sK3)) <- (12) [resolution 1372,821] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1539. 18 | ~12 [avatar split clause 1538,1371,1434]
1540. ~9 | ~10 | 5 [avatar split clause 1525,1274,1349,1298]
1541. 10 | ~12 [avatar split clause 1537,1371,1349]
1542. 10 | 15 [avatar split clause 1457,1389,1349]
1543. ~10 | ~19 | 11 [avatar split clause 1529,1366,1439,1349]
1544. ~v1_xboole_0(k3_gate_1(sK1,k8_gate_1(sK0,sK2,sK3))) <- (~4) [forward demodulation 1271,820] {a:0,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:2,thDist:-2}
1545. ~12 | 4 [avatar split clause 1544,1270,1371]
1548. v1_xboole_0(sK1) <- (5) [resolution 1276,821] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1550. $false <- (5, ~10) [subsumption resolution 1548,1350] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1551. ~5 | 10 [avatar contradiction clause 1550]
1554. v1_xboole_0(k1_gate_1(sK0)) <- (1) [resolution 1260,1363] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1556. 16 | ~1 [avatar split clause 1554,1258,1423]
1558. ~v1_xboole_0(sK0) <- (16) [resolution 1425,816] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1560. $false <- (9, 16) [subsumption resolution 1558,1299] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1561. ~9 | ~16 [avatar contradiction clause 1560]
1563. ~v1_xboole_0(sK0) <- (~2) [resolution 1264,822] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1564. ~9 | 2 [avatar split clause 1563,1262,1298]
1566. ~v1_xboole_0(k1_gate_1(sK1)) | ~v1_xboole_0(k1_gate_1(sK0)) <- (~1) [resolution 1259,824] {a:1,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1567. ~v1_xboole_0(k1_gate_1(sK0)) <- (~1, 15) [subsumption resolution 1566,1391] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
1568. ~16 | 1 | ~15 [avatar split clause 1567,1389,1258,1423]
1570. v1_xboole_0(sK0) <- (~16) [resolution 1424,815] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1572. $false <- (~9, ~16) [subsumption resolution 1570,1300] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1573. 9 | 16 [avatar contradiction clause 1572]
1574. $false [avatar sat refutation 1285,1286,1287,1325,1343,1352,1354,1356,1392,1398,1407,1431,1432,1437,1446,1453,1465,1467,1477,1479,1480,1484,1490,1495,1503,1504,1506,1508,1517,1519,1522,1524,1536,1539,1540,1541,1542,1543,1545,1551,1556,1561,1564,1568,1573] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:846,thDist:-846}
% SZS output end Proof for gate_2__t3_gate_2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 22
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 782
% Generated clauses: 1425
% Active clauses: 71
% Passive clauses: 83
% Final active clauses: 27
% Final passive clauses: 4
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 752
% Fw subsumption resolutions: 20
% Fw demodulations: 11
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 724
% Forward subsumptions: 34
% 
% >>> Generating Inferences
% Binary resolution: 59
% Unit resulting resolution: 12
% Forward superposition: 8
% Backward superposition: 4
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 59
% Split components: 396
% Sat splitting refutations: 26
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 124
% SAT solver binary clauses: 56
% 
% Memory used [KB]: 5500
% Time elapsed: 0.004 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% forward subsumption resolution: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.004 s ( own 0.001 s ) 

