i: [1,114,0,2,1,1,0,0] t7_diraf
a: [3,193,0,2,1,114]
s: 193
i: [1,113,0,2,1,1,0,0] t7_diraf
a: [3,198,0,2,1,113]
s: 198
i: [1,112,0,2,1,1,0,0] t7_diraf
a: [3,203,0,2,1,112]
s: 203
i: [1,117,0,4,1,1,0,0] t7_diraf
a: [3,178,0,4,1,117]
s: 178
i: [1,116,0,4,1,1,0,0] t7_diraf
a: [3,183,0,4,1,116]
s: 183
i: [1,115,0,4,1,1,0,0] t7_diraf
a: [3,188,0,4,1,115]
s: 188
i: [1,120,0,11,2,1,0,0] t7_diraf
a: [3,166,0,5,1,120]
s: 166
a: [3,163,0,6,1,120]
s: 163
i: [1,121,0,2,1,0,0,255] spc1_boole
a: [3,208,0,2,1,121]
s: 208
i: [1,122,0,2,1,0,0,255] fc1_xboole_0
a: [3,213,0,2,1,122]
s: 213
i: [1,156,0,2,1,0,0,2] rc2_xboole_0
a: [3,218,0,2,1,156]
s: 218
i: [1,157,0,2,1,0,0,1] existence_l1_analoaf
a: [3,223,0,2,1,157]
s: 223
i: [1,158,0,2,1,0,0,2] existence_l1_struct_0
a: [3,228,0,2,1,158]
s: 228
i: [1,159,0,2,1,0,0,2] rc1_xboole_0
a: [3,233,0,2,1,159]
s: 233
i: [1,123,0,4,2,0,0,1] dt_l1_analoaf
s: 123
i: [1,154,0,4,1,0,0,1] existence_m1_subset_1
s: 154
d: [2,235,1,2,1,1,40,123,193]
a: [3,240,1,2,1,235]
s: 240
i: [1,138,0,5,2,0,0,255] t6_boole
s: 138
d: [2,236,1,2,1,1,40,123,223]
a: [3,245,1,2,1,236]
s: 245
i: [1,153,0,5,2,0,0,255] t7_boole
s: 153
i: [1,124,0,6,3,0,0,255] cc7_struct_0
s: 124
d: [2,249,1,4,2,1,40,124,228]
a: [3,255,1,2,1,249]
s: 255
i: [1,125,0,6,3,0,0,255] cc5_struct_0
s: 125
d: [2,251,2,4,2,1,40,124,245]
a: [3,264,2,2,1,251]
s: 264
i: [1,127,0,6,3,0,0,1] cc17_struct_0
s: 127
i: [1,128,0,6,3,0,0,1] cc2_struct_0
s: 128
i: [1,129,0,6,3,0,0,1] cc1_struct_0
s: 129
i: [1,132,0,6,3,0,0,255] cc4_struct_0
s: 132
i: [1,133,0,6,3,0,0,255] cc6_struct_0
s: 133
i: [1,150,0,6,2,0,0,255] t1_subset
s: 150
i: [1,155,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 155
i: [1,130,0,7,3,0,0,255] cc8_struct_0
s: 130
i: [1,134,0,7,3,0,0,255] cc11_struct_0
s: 134
i: [1,135,0,7,3,0,0,255] cc11_struct_0
s: 135
i: [1,136,0,7,3,0,0,255] cc9_struct_0
s: 136
i: [1,139,0,7,3,0,0,255] fc9_struct_0
s: 139
i: [1,142,0,7,3,0,0,1] fc2_struct_0
s: 142
i: [1,146,0,7,3,0,0,255] fc6_struct_0
s: 146
i: [1,147,0,7,3,0,0,255] fc8_struct_0
s: 147
i: [1,148,0,7,3,0,0,1] fc1_struct_0
s: 148
i: [1,149,0,7,3,0,0,255] fc7_struct_0
s: 149
i: [1,152,0,7,3,0,0,255] t8_boole
s: 152
i: [1,151,0,8,3,0,0,255] t2_subset
s: 151
d: [2,280,1,7,2,1,40,151,178]
a: [3,291,1,3,1,280]
s: 291
r: 291
i: [1,137,0,9,4,0,0,255] cc10_struct_0
s: 137
i: [1,140,0,27,7,0,0,1] d3_diraf
s: 140
a: [3,290,1,3,1,280]
s: 290
i: [1,141,0,27,7,0,0,1] d3_diraf
s: 141
d: [2,303,2,4,2,1,40,291,142]
d: [2,304,2,2,1,2,21,303,240]
a: [3,307,2,2,1,304]
s: 307
i: [1,144,0,30,8,0,0,1] t6_diraf
s: 144
i: [1,145,0,30,8,0,0,1] t6_diraf
s: 145
a: [3,287,1,4,1,280]
s: 287
i: [1,143,0,36,9,0,0,1] t6_diraf
s: 143
d: [2,281,1,7,2,1,40,151,183]
a: [3,296,1,4,1,281]
s: 296
d: [2,282,1,7,2,1,40,151,188]
a: [3,301,1,4,1,282]
s: 301
d: [2,283,1,6,2,0,40,151,154]
s: 283
d: [2,323,2,4,1,1,40,287,155]
a: [3,329,2,4,1,323]
s: 329
d: [2,247,1,7,3,0,44,138,138]
s: 247
d: [2,333,2,4,1,1,40,296,155]
a: [3,339,2,4,1,333]
s: 339
d: [2,270,1,9,4,0,44,130,138]
s: 270
d: [2,272,1,9,4,0,44,136,138]
s: 272
d: [2,341,2,4,1,1,40,301,155]
a: [3,347,2,4,1,341]
s: 347
d: [2,319,1,45,13,0,40,144,140]
d: [2,320,1,31,9,0,18,319]
s: 320
i: [1,119,0,11,2,1,0,0] t7_diraf
a: [3,172,0,5,1,119]
s: 172
d: [2,321,1,45,13,0,40,145,140]
d: [2,322,1,31,9,0,18,321]
s: 322
d: [2,381,3,0,0,9,52,307,193,203,198,178,188,183,166,163,322]
e: 381
d: [2,382,2,25,8,1,40,322,163]
d: [2,392,2,21,7,2,21,382,178]
d: [2,393,2,17,6,3,21,392,183]
d: [2,394,2,13,5,4,21,393,188]
d: [2,395,2,11,4,5,21,394,193]
d: [2,396,2,9,3,6,21,395,198]
d: [2,397,2,7,2,7,21,396,203]
d: [2,398,2,2,1,8,21,397,166]
d: [2,399,2,0,0,9,21,398,307]
e: 399
r: 290
r: 307
r: 163
r: 166
a: [3,167,0,5,1,120]
s: 167
s: 307
a: [3,162,0,6,1,120]
s: 162
s: 290
d: [2,405,1,30,8,1,40,162,143]
d: [2,406,1,26,7,2,21,405,178]
d: [2,407,1,22,6,3,21,406,183]
d: [2,408,1,18,5,4,21,407,188]
d: [2,409,1,16,4,5,21,408,193]
d: [2,410,1,14,3,6,21,409,198]
d: [2,411,1,12,2,7,21,410,203]
a: [3,419,1,6,1,411]
s: 419
d: [2,422,3,0,0,7,52,307,193,188,183,178,172,419,141]
e: 422
d: [2,425,2,21,6,1,40,419,141]
d: [2,433,2,16,5,2,21,425,172]
d: [2,434,2,12,4,3,21,433,178]
d: [2,435,2,8,3,4,21,434,183]
d: [2,436,2,4,2,5,21,435,188]
d: [2,437,2,2,1,6,21,436,193]
d: [2,438,2,0,0,7,21,437,307]
e: 438
r: 419
a: [3,415,1,6,1,411]
s: 415
d: [2,440,3,0,0,7,52,307,193,188,178,183,167,415,141]
e: 440
d: [2,443,2,21,6,1,40,415,141]
d: [2,451,2,16,5,2,21,443,167]
d: [2,452,2,12,4,3,21,451,183]
d: [2,453,2,8,3,4,21,452,178]
d: [2,454,2,4,2,5,21,453,188]
d: [2,455,2,2,1,6,21,454,193]
d: [2,456,2,0,0,7,21,455,307]
e: 456
r: 287
r: 290
r: 296
r: 301
r: 307
a: [3,308,2,2,1,304]
s: 308
d: [2,464,3,0,0,3,52,240,203,308,127]
e: 464
d: [2,465,3,4,2,1,40,308,127]
d: [2,468,3,2,1,2,21,465,203]
d: [2,469,3,0,0,3,21,468,240]
e: 469
i: [1,118,0,16,3,1,0,0] t7_diraf
d: [2,364,2,25,8,1,40,320,163]
d: [2,373,2,21,7,2,21,364,178]
d: [2,374,2,17,6,3,21,373,183]
d: [2,375,2,13,5,4,21,374,188]
d: [2,376,2,11,4,5,21,375,193]
d: [2,377,2,9,3,6,21,376,198]
d: [2,378,2,7,2,7,21,377,203]
d: [2,379,2,5,1,8,21,378,307]
f: 469,455,438,411,398,379,235,112,113,114,115,116,117,118,119,120
% Refutation found. Thanks to Tanya!
% SZS status Theorem for diraf__t7_diraf
% SZS output start Proof for diraf__t7_diraf
1. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X1,X3) <=> (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3))))))) [input t7_diraf]
2. ~! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X1,X3) <=> (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3))))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)))))) [input d3_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
37. ! [X0] : ((l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X1,X3) <=> (r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X2,X3))))))) [input t6_diraf]
44. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r2_analoaf(X0,X1,X2,X1,X3) <~> (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
45. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((r2_analoaf(X0,X1,X2,X1,X3) <~> (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 44]
46. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
51. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
52. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 51]
72. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r1_diraf(X0,X1,X2,X3) <=> r2_analoaf(X0,X1,X2,X2,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 72]
76. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_analoaf(X0,X1,X2,X1,X3) <=> (r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0))) [ennf transformation 37]
77. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((r2_analoaf(X0,X1,X2,X1,X3) <=> (r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 76]
92. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((~r1_diraf(X0,X1,X3,X2) & ~r1_diraf(X0,X1,X2,X3)) | ~r2_analoaf(X0,X1,X2,X1,X3)) & ((r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3)) | r2_analoaf(X0,X1,X2,X1,X3))) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [nnf transformation 45]
93. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_diraf(X0,X1,X3,X2) & ~r1_diraf(X0,X1,X2,X3)) | ~r2_analoaf(X0,X1,X2,X1,X3)) & (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3) | r2_analoaf(X0,X1,X2,X1,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) [flattening 92]
94. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r1_diraf(X0,X1,X3,X2) & ~r1_diraf(X0,X1,X2,X3)) | ~r2_analoaf(X0,X1,X2,X1,X3)) & (r1_diraf(X0,X1,X3,X2) | r1_diraf(X0,X1,X2,X3) | r2_analoaf(X0,X1,X2,X1,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v2_analoaf(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (((~r1_diraf(sK0,X1,X3,X2) & ~r1_diraf(sK0,X1,X2,X3)) | ~r2_analoaf(sK0,X1,X2,X1,X3)) & (r1_diraf(sK0,X1,X3,X2) | r1_diraf(sK0,X1,X2,X3) | r2_analoaf(sK0,X1,X2,X1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0)) [choice axiom]
95. ? [X1] : (? [X2] : (? [X3] : (((~r1_diraf(sK0,X1,X3,X2) & ~r1_diraf(sK0,X1,X2,X3)) | ~r2_analoaf(sK0,X1,X2,X1,X3)) & (r1_diraf(sK0,X1,X3,X2) | r1_diraf(sK0,X1,X2,X3) | r2_analoaf(sK0,X1,X2,X1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (((~r1_diraf(sK0,sK1,X3,X2) & ~r1_diraf(sK0,sK1,X2,X3)) | ~r2_analoaf(sK0,sK1,X2,sK1,X3)) & (r1_diraf(sK0,sK1,X3,X2) | r1_diraf(sK0,sK1,X2,X3) | r2_analoaf(sK0,sK1,X2,sK1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
96. ? [X2] : (? [X3] : (((~r1_diraf(sK0,sK1,X3,X2) & ~r1_diraf(sK0,sK1,X2,X3)) | ~r2_analoaf(sK0,sK1,X2,sK1,X3)) & (r1_diraf(sK0,sK1,X3,X2) | r1_diraf(sK0,sK1,X2,X3) | r2_analoaf(sK0,sK1,X2,sK1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (((~r1_diraf(sK0,sK1,X3,sK2) & ~r1_diraf(sK0,sK1,sK2,X3)) | ~r2_analoaf(sK0,sK1,sK2,sK1,X3)) & (r1_diraf(sK0,sK1,X3,sK2) | r1_diraf(sK0,sK1,sK2,X3) | r2_analoaf(sK0,sK1,sK2,sK1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
97. ? [X3] : (((~r1_diraf(sK0,sK1,X3,sK2) & ~r1_diraf(sK0,sK1,sK2,X3)) | ~r2_analoaf(sK0,sK1,sK2,sK1,X3)) & (r1_diraf(sK0,sK1,X3,sK2) | r1_diraf(sK0,sK1,sK2,X3) | r2_analoaf(sK0,sK1,sK2,sK1,X3)) & m1_subset_1(X3,u1_struct_0(sK0))) => (((~r1_diraf(sK0,sK1,sK3,sK2) & ~r1_diraf(sK0,sK1,sK2,sK3)) | ~r2_analoaf(sK0,sK1,sK2,sK1,sK3)) & (r1_diraf(sK0,sK1,sK3,sK2) | r1_diraf(sK0,sK1,sK2,sK3) | r2_analoaf(sK0,sK1,sK2,sK1,sK3)) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
98. (((((~r1_diraf(sK0,sK1,sK3,sK2) & ~r1_diraf(sK0,sK1,sK2,sK3)) | ~r2_analoaf(sK0,sK1,sK2,sK1,sK3)) & (r1_diraf(sK0,sK1,sK3,sK2) | r1_diraf(sK0,sK1,sK2,sK3) | r2_analoaf(sK0,sK1,sK2,sK1,sK3)) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v2_analoaf(sK0) & ~v7_struct_0(sK0) [skolemisation 93,97,96,95,94]
99. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r1_diraf(X0,X1,X2,X3) | ~r2_analoaf(X0,X1,X2,X2,X3)) & (r2_analoaf(X0,X1,X2,X2,X3) | ~r1_diraf(X0,X1,X2,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 73]
100. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r2_analoaf(X0,X1,X2,X1,X3) | (~r2_analoaf(X0,X1,X3,X3,X2) & ~r2_analoaf(X0,X1,X2,X2,X3))) & ((r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X2,X3)) | ~r2_analoaf(X0,X1,X2,X1,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [nnf transformation 77]
101. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((r2_analoaf(X0,X1,X2,X1,X3) | (~r2_analoaf(X0,X1,X3,X3,X2) & ~r2_analoaf(X0,X1,X2,X2,X3))) & (r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X2,X3) | ~r2_analoaf(X0,X1,X2,X1,X3))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0)) [flattening 100]
112. ~v7_struct_0(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
113. v2_analoaf(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
114. l1_analoaf(sK0) [cnf transformation 98] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
115. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 98] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
116. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 98] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
117. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 98] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
118. r1_diraf(sK0,sK1,sK3,sK2) | r1_diraf(sK0,sK1,sK2,sK3) | r2_analoaf(sK0,sK1,sK2,sK1,sK3) [cnf transformation 98] {a:0,w:16,wCS:1600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
119. ~r1_diraf(sK0,sK1,sK2,sK3) | ~r2_analoaf(sK0,sK1,sK2,sK1,sK3) [cnf transformation 98] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
120. ~r1_diraf(sK0,sK1,sK3,sK2) | ~r2_analoaf(sK0,sK1,sK2,sK1,sK3) [cnf transformation 98] {a:0,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
123. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 46] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 52] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. r2_analoaf(X0,X1,X2,X2,X3) | ~r1_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
141. ~r2_analoaf(X0,X1,X2,X2,X3) | r1_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 99] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
143. ~r2_analoaf(X0,X1,X2,X1,X3) | r2_analoaf(X0,X1,X2,X2,X3) | r2_analoaf(X0,X1,X3,X3,X2) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 101] {a:0,w:36,wCS:14400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
144. ~r2_analoaf(X0,X1,X2,X2,X3) | r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 101] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. ~r2_analoaf(X0,X1,X3,X3,X2) | r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) [cnf transformation 101] {a:0,w:30,wCS:12000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
161. 1 <=> r2_analoaf(sK0,sK1,sK2,sK1,sK3) [avatar definition]
162. r2_analoaf(sK0,sK1,sK2,sK1,sK3) <- (1) [avatar component clause 161] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
163. ~r2_analoaf(sK0,sK1,sK2,sK1,sK3) <- (~1) [avatar component clause 161] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
165. 2 <=> r1_diraf(sK0,sK1,sK3,sK2) [avatar definition]
166. r1_diraf(sK0,sK1,sK3,sK2) <- (2) [avatar component clause 165] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
167. ~r1_diraf(sK0,sK1,sK3,sK2) <- (~2) [avatar component clause 165] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
168. ~1 | ~2 [avatar split clause 120,165,161]
170. 3 <=> r1_diraf(sK0,sK1,sK2,sK3) [avatar definition]
172. ~r1_diraf(sK0,sK1,sK2,sK3) <- (~3) [avatar component clause 170] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
173. ~1 | ~3 [avatar split clause 119,170,161]
174. 1 | 3 | 2 [avatar split clause 118,165,170,161]
176. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
178. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 176] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
179. 4 [avatar split clause 117,176]
181. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
183. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 181] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
184. 5 [avatar split clause 116,181]
186. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
188. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 186] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
189. 6 [avatar split clause 115,186]
191. 7 <=> l1_analoaf(sK0) [avatar definition]
193. l1_analoaf(sK0) <- (7) [avatar component clause 191] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
194. 7 [avatar split clause 114,191]
196. 8 <=> v2_analoaf(sK0) [avatar definition]
198. v2_analoaf(sK0) <- (8) [avatar component clause 196] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
199. 8 [avatar split clause 113,196]
201. 9 <=> v7_struct_0(sK0) [avatar definition]
203. ~v7_struct_0(sK0) <- (~9) [avatar component clause 201] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
204. ~9 [avatar split clause 112,201]
235. l1_struct_0(sK0) <- (7) [resolution 123,193] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
238. 16 <=> l1_struct_0(sK0) [avatar definition]
240. l1_struct_0(sK0) <- (16) [avatar component clause 238] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
241. 16 | ~7 [avatar split clause 235,191,238]
306. 26 <=> v2_struct_0(sK0) [avatar definition]
307. ~v2_struct_0(sK0) <- (~26) [avatar component clause 306] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
308. v2_struct_0(sK0) <- (26) [avatar component clause 306] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
319. r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) | ~r1_diraf(X0,X1,X2,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [resolution 144,140] {a:1,w:45,wCS:18000,sine:1,thAx:0,allAx:2,thDist:-2}
320. r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) | ~r1_diraf(X0,X1,X2,X3) | v2_struct_0(X0) [duplicate literal removal 319] {a:1,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
321. r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) | ~r1_diraf(X0,X1,X3,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [resolution 145,140] {a:1,w:45,wCS:18000,sine:1,thAx:0,allAx:2,thDist:-2}
322. r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v2_analoaf(X0) | v7_struct_0(X0) | ~r1_diraf(X0,X1,X3,X2) | v2_struct_0(X0) [duplicate literal removal 321] {a:1,w:31,wCS:12400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
364. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1) [resolution 320,163] {a:2,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
373. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4) [subsumption resolution 364,178] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
374. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4, 5) [subsumption resolution 373,183] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
375. ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4, 5, 6) [subsumption resolution 374,188] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
376. ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7) [subsumption resolution 375,193] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
377. v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7, 8) [subsumption resolution 376,198] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
378. ~r1_diraf(sK0,sK1,sK2,sK3) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7, 8, ~9) [subsumption resolution 377,203] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
379. ~r1_diraf(sK0,sK1,sK2,sK3) <- (~1, 4, 5, 6, 7, 8, ~9, ~26) [subsumption resolution 378,307] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
380. ~3 | 1 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | 26 [avatar split clause 379,306,201,196,191,186,181,176,161,170]
382. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1) [resolution 322,163] {a:2,w:25,wCS:2500,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
392. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4) [subsumption resolution 382,178] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
393. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4, 5) [subsumption resolution 392,183] {a:2,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
394. ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4, 5, 6) [subsumption resolution 393,188] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
395. ~v2_analoaf(sK0) | v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7) [subsumption resolution 394,193] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
396. v7_struct_0(sK0) | ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7, 8) [subsumption resolution 395,198] {a:2,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
397. ~r1_diraf(sK0,sK1,sK3,sK2) | v2_struct_0(sK0) <- (~1, 4, 5, 6, 7, 8, ~9) [subsumption resolution 396,203] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
398. v2_struct_0(sK0) <- (~1, 2, 4, 5, 6, 7, 8, ~9) [subsumption resolution 397,166] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
404. 26 | 1 | ~2 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 [avatar split clause 398,201,196,191,186,181,176,165,161,306]
405. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1) [resolution 162,143] {a:1,w:30,wCS:3000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
406. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, 4) [subsumption resolution 405,178] {a:1,w:26,wCS:2600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
407. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, 4, 5) [subsumption resolution 406,183] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
408. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | ~l1_analoaf(sK0) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, 4, 5, 6) [subsumption resolution 407,188] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
409. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | ~v2_analoaf(sK0) | v7_struct_0(sK0) <- (1, 4, 5, 6, 7) [subsumption resolution 408,193] {a:1,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
410. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) | v7_struct_0(sK0) <- (1, 4, 5, 6, 7, 8) [subsumption resolution 409,198] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
411. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK3,sK3,sK2) <- (1, 4, 5, 6, 7, 8, ~9) [subsumption resolution 410,203] {a:1,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
413. 30 <=> r2_analoaf(sK0,sK1,sK3,sK3,sK2) [avatar definition]
415. r2_analoaf(sK0,sK1,sK3,sK3,sK2) <- (30) [avatar component clause 413] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
417. 31 <=> r2_analoaf(sK0,sK1,sK2,sK2,sK3) [avatar definition]
419. r2_analoaf(sK0,sK1,sK2,sK2,sK3) <- (31) [avatar component clause 417] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
420. 30 | 31 | ~1 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 [avatar split clause 411,201,196,191,186,181,176,161,417,413]
425. r1_diraf(sK0,sK1,sK2,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (31) [resolution 419,141] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
433. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~3, 31) [subsumption resolution 425,172] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
434. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~3, 4, 31) [subsumption resolution 433,178] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
435. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~3, 4, 5, 31) [subsumption resolution 434,183] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
436. ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~3, 4, 5, 6, 31) [subsumption resolution 435,188] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
437. v2_struct_0(sK0) <- (~3, 4, 5, 6, 7, 31) [subsumption resolution 436,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
438. $false <- (~3, 4, 5, 6, 7, ~26, 31) [subsumption resolution 437,307] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
439. 3 | ~4 | ~5 | ~6 | ~7 | 26 | ~31 [avatar contradiction clause 438]
443. r1_diraf(sK0,sK1,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (30) [resolution 415,141] {a:2,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
451. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~2, 30) [subsumption resolution 443,167] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
452. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~2, 5, 30) [subsumption resolution 451,183] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
453. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~2, 4, 5, 30) [subsumption resolution 452,178] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
454. ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~2, 4, 5, 6, 30) [subsumption resolution 453,188] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
455. v2_struct_0(sK0) <- (~2, 4, 5, 6, 7, 30) [subsumption resolution 454,193] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
458. 26 | 2 | ~4 | ~5 | ~6 | ~7 | ~30 [avatar split clause 455,413,191,186,181,176,165,306]
465. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (26) [resolution 308,127] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
468. ~l1_struct_0(sK0) <- (~9, 26) [subsumption resolution 465,203] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
469. $false <- (~9, 16, 26) [subsumption resolution 468,240] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
470. 9 | ~16 | ~26 [avatar contradiction clause 469]
472. $false [avatar sat refutation 168,173,174,179,184,189,194,199,204,241,380,404,420,439,458,470] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:452,thDist:-452}
% SZS output end Proof for diraf__t7_diraf
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 38
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 48
% Generated clauses: 241
% Active clauses: 67
% Passive clauses: 86
% Final active clauses: 56
% Final passive clauses: 12
% Main loop iterations started: 72
% 
% >>> Simplifying Inferences
% Duplicate literals: 85
% Fw subsumption resolutions: 50
% 
% >>> Deletion Inferences
% Simple tautologies: 15
% Forward subsumptions: 22
% 
% >>> Generating Inferences
% Binary resolution: 59
% Unit resulting resolution: 4
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 9
% Split components: 19
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 45
% SAT solver unit clauses: 12
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 9594
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.003 s

