i: [1,41,0,2,1,1,0,0] t33_ordinal1
a: [3,82,0,2,1,41]
s: 82
i: [1,40,0,2,1,1,0,0] t33_ordinal1
a: [3,87,0,2,1,40]
s: 87
i: [1,39,0,2,1,1,0,0] t33_ordinal1
a: [3,92,0,2,1,39]
s: 92
i: [1,37,0,2,1,1,0,0] t33_ordinal1
a: [3,102,0,2,1,37]
s: 102
i: [1,38,0,3,1,1,0,0] t33_ordinal1
a: [3,97,0,3,1,38]
s: 97
i: [1,42,0,17,4,1,0,0] t33_ordinal1
a: [3,65,0,4,1,42]
s: 65
i: [1,59,0,2,1,0,0,1] rc1_ordinal1
a: [3,107,0,2,1,59]
s: 107
i: [1,61,0,2,1,0,0,1] rc1_funct_1
a: [3,112,0,2,1,61]
s: 112
i: [1,60,0,2,1,0,0,1] rc1_funct_1
a: [3,117,0,2,1,60]
s: 117
i: [1,46,0,3,1,0,0,1] rc3_ordinal1
s: 46
i: [1,48,0,3,1,0,0,1] rc3_ordinal1
s: 48
i: [1,49,0,3,1,0,0,1] rc3_ordinal1
s: 49
i: [1,43,0,4,2,0,0,255] cc1_ordinal1
s: 43
d: [2,119,1,2,1,1,40,43,82]
a: [3,124,1,2,1,119]
s: 124
i: [1,44,0,4,2,0,0,255] cc1_ordinal1
s: 44
i: [1,47,0,4,1,0,0,1] rc3_ordinal1
s: 47
d: [2,131,1,2,1,1,40,44,82]
a: [3,136,1,2,1,131]
s: 136
i: [1,45,0,6,3,0,0,255] cc2_ordinal1
s: 45
d: [2,120,1,2,1,1,40,43,107]
a: [3,129,1,2,1,120]
s: 129
i: [1,50,0,6,2,0,0,1] dt_k5_relat_1
s: 50
d: [2,144,1,0,0,2,52,102,65,50]
e: 144
d: [2,145,1,2,1,1,40,50,65]
d: [2,147,1,0,0,2,21,145,102]
e: 147
r: 65
a: [3,64,0,4,1,42]
s: 64
a: [3,69,0,5,1,42]
s: 69
i: [1,55,0,8,3,0,0,1] fc8_funct_1
s: 55
d: [2,132,1,2,1,1,40,44,107]
a: [3,141,1,2,1,132]
s: 141
i: [1,56,0,8,3,0,0,1] fc8_funct_1
s: 56
i: [1,57,0,9,3,0,0,1] fc26_relat_1
s: 57
i: [1,58,0,10,3,0,0,1] fc26_relat_1
s: 58
d: [2,150,1,0,0,3,52,102,97,69,58]
e: 150
d: [2,151,1,5,2,1,40,58,69]
d: [2,153,1,2,1,2,21,151,97]
d: [2,154,1,0,0,3,21,153,102]
e: 154
r: 69
a: [3,73,0,4,1,42]
s: 73
d: [2,156,1,0,0,3,52,102,92,73,56]
e: 156
d: [2,157,1,4,2,1,40,73,56]
d: [2,159,1,2,1,2,21,157,92]
d: [2,160,1,0,0,3,21,159,102]
e: 160
r: 73
a: [3,72,0,4,1,42]
s: 72
a: [3,77,0,4,1,42]
s: 77
a: [3,68,0,5,1,42]
s: 68
i: [1,51,0,11,2,0,0,255] rd5_relat_1
s: 51
i: [1,52,0,12,5,0,0,255] fc5_ordinal1
s: 52
i: [1,54,0,12,5,0,0,255] fc5_ordinal1
s: 54
d: [2,166,1,0,0,5,52,102,92,87,82,77,54]
e: 166
d: [2,167,1,8,4,1,40,54,77]
d: [2,170,1,6,3,2,21,167,82]
d: [2,171,1,4,2,3,21,170,87]
d: [2,172,1,2,1,4,21,171,92]
d: [2,173,1,0,0,5,21,172,102]
e: 173
f: 173,160,154,147,37,38,39,40,41,42
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal1__t33_ordinal1
% SZS output start Proof for ordinal1__t33_ordinal1
1. ! [X0] : ! [X1] : ((v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (v3_ordinal1(X2) => (v5_ordinal1(k5_relat_1(X1,X2)) & v1_funct_1(k5_relat_1(X1,X2)) & v5_relat_1(k5_relat_1(X1,X2),X0) & v1_relat_1(k5_relat_1(X1,X2))))) [input t33_ordinal1]
2. ~! [X0] : ! [X1] : ((v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (v3_ordinal1(X2) => (v5_ordinal1(k5_relat_1(X1,X2)) & v1_funct_1(k5_relat_1(X1,X2)) & v5_relat_1(k5_relat_1(X1,X2),X0) & v1_relat_1(k5_relat_1(X1,X2))))) [negated conjecture 1]
5. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k5_relat_1(X0,X1))) [input dt_k5_relat_1]
7. ! [X0,X1,X2] : ((v5_relat_1(X2,X1) & v1_relat_1(X2)) => (v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0)))) [input fc26_relat_1]
8. ! [X0,X1] : ((v3_ordinal1(X1) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v5_ordinal1(k5_relat_1(X0,X1)) & v5_relat_1(k5_relat_1(X0,X1),k10_xtuple_0(X0)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc5_ordinal1]
9. ! [X0,X1] : ((v1_funct_1(X0) & v1_relat_1(X0)) => (v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1)))) [input fc8_funct_1]
14. ~! [X0,X1] : ((v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => ! [X2] : (v3_ordinal1(X2) => (v5_ordinal1(k5_relat_1(X1,X2)) & v1_funct_1(k5_relat_1(X1,X2)) & v5_relat_1(k5_relat_1(X1,X2),X0) & v1_relat_1(k5_relat_1(X1,X2))))) [flattening 2]
15. ? [X0,X1] : (? [X2] : ((~v5_ordinal1(k5_relat_1(X1,X2)) | ~v1_funct_1(k5_relat_1(X1,X2)) | ~v5_relat_1(k5_relat_1(X1,X2),X0) | ~v1_relat_1(k5_relat_1(X1,X2))) & v3_ordinal1(X2)) & (v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1))) [ennf transformation 14]
16. ? [X0,X1] : (? [X2] : ((~v5_ordinal1(k5_relat_1(X1,X2)) | ~v1_funct_1(k5_relat_1(X1,X2)) | ~v5_relat_1(k5_relat_1(X1,X2),X0) | ~v1_relat_1(k5_relat_1(X1,X2))) & v3_ordinal1(X2)) & v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) [flattening 15]
20. ! [X0,X1] : (v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 5]
22. ! [X0,X1] : ((v5_ordinal1(k5_relat_1(X0,X1)) & v5_relat_1(k5_relat_1(X0,X1),k10_xtuple_0(X0)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v3_ordinal1(X1) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 8]
23. ! [X0,X1] : ((v5_ordinal1(k5_relat_1(X0,X1)) & v5_relat_1(k5_relat_1(X0,X1),k10_xtuple_0(X0)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v3_ordinal1(X1) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 22]
24. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 9]
25. ! [X0,X1] : ((v1_funct_1(k5_relat_1(X0,X1)) & v1_relat_1(k5_relat_1(X0,X1))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 24]
26. ! [X0,X1,X2] : ((v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0))) | (~v5_relat_1(X2,X1) | ~v1_relat_1(X2))) [ennf transformation 7]
27. ! [X0,X1,X2] : ((v5_relat_1(k5_relat_1(X2,X0),X1) & v1_relat_1(k5_relat_1(X2,X0))) | ~v5_relat_1(X2,X1) | ~v1_relat_1(X2)) [flattening 26]
28. ? [X0,X1] : (? [X2] : ((~v5_ordinal1(k5_relat_1(X1,X2)) | ~v1_funct_1(k5_relat_1(X1,X2)) | ~v5_relat_1(k5_relat_1(X1,X2),X0) | ~v1_relat_1(k5_relat_1(X1,X2))) & v3_ordinal1(X2)) & v5_ordinal1(X1) & v1_funct_1(X1) & v5_relat_1(X1,X0) & v1_relat_1(X1)) => (? [X2] : ((~v5_ordinal1(k5_relat_1(sK1,X2)) | ~v1_funct_1(k5_relat_1(sK1,X2)) | ~v5_relat_1(k5_relat_1(sK1,X2),sK0) | ~v1_relat_1(k5_relat_1(sK1,X2))) & v3_ordinal1(X2)) & v5_ordinal1(sK1) & v1_funct_1(sK1) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1)) [choice axiom]
29. ? [X2] : ((~v5_ordinal1(k5_relat_1(sK1,X2)) | ~v1_funct_1(k5_relat_1(sK1,X2)) | ~v5_relat_1(k5_relat_1(sK1,X2),sK0) | ~v1_relat_1(k5_relat_1(sK1,X2))) & v3_ordinal1(X2)) => ((~v5_ordinal1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v5_relat_1(k5_relat_1(sK1,sK2),sK0) | ~v1_relat_1(k5_relat_1(sK1,sK2))) & v3_ordinal1(sK2)) [choice axiom]
30. ((~v5_ordinal1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v5_relat_1(k5_relat_1(sK1,sK2),sK0) | ~v1_relat_1(k5_relat_1(sK1,sK2))) & v3_ordinal1(sK2)) & v5_ordinal1(sK1) & v1_funct_1(sK1) & v5_relat_1(sK1,sK0) & v1_relat_1(sK1) [skolemisation 16,29,28]
37. v1_relat_1(sK1) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. v5_relat_1(sK1,sK0) [cnf transformation 30] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. v1_funct_1(sK1) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. v5_ordinal1(sK1) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
41. v3_ordinal1(sK2) [cnf transformation 30] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~v5_ordinal1(k5_relat_1(sK1,sK2)) | ~v1_funct_1(k5_relat_1(sK1,sK2)) | ~v5_relat_1(k5_relat_1(sK1,sK2),sK0) | ~v1_relat_1(k5_relat_1(sK1,sK2)) [cnf transformation 30] {a:0,w:17,wCS:1700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
50. v1_relat_1(k5_relat_1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 20] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. v5_ordinal1(k5_relat_1(X0,X1)) | ~v3_ordinal1(X1) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 23] {a:0,w:12,wCS:4800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
56. v1_funct_1(k5_relat_1(X0,X1)) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. v5_relat_1(k5_relat_1(X2,X0),X1) | ~v5_relat_1(X2,X1) | ~v1_relat_1(X2) [cnf transformation 27] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
63. 1 <=> v1_relat_1(k5_relat_1(sK1,sK2)) [avatar definition]
65. ~v1_relat_1(k5_relat_1(sK1,sK2)) <- (~1) [avatar component clause 63] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 2 <=> v5_relat_1(k5_relat_1(sK1,sK2),sK0) [avatar definition]
69. ~v5_relat_1(k5_relat_1(sK1,sK2),sK0) <- (~2) [avatar component clause 67] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
71. 3 <=> v1_funct_1(k5_relat_1(sK1,sK2)) [avatar definition]
73. ~v1_funct_1(k5_relat_1(sK1,sK2)) <- (~3) [avatar component clause 71] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
75. 4 <=> v5_ordinal1(k5_relat_1(sK1,sK2)) [avatar definition]
77. ~v5_ordinal1(k5_relat_1(sK1,sK2)) <- (~4) [avatar component clause 75] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
78. ~1 | ~2 | ~3 | ~4 [avatar split clause 42,75,71,67,63]
80. 5 <=> v3_ordinal1(sK2) [avatar definition]
82. v3_ordinal1(sK2) <- (5) [avatar component clause 80] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
83. 5 [avatar split clause 41,80]
85. 6 <=> v5_ordinal1(sK1) [avatar definition]
87. v5_ordinal1(sK1) <- (6) [avatar component clause 85] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
88. 6 [avatar split clause 40,85]
90. 7 <=> v1_funct_1(sK1) [avatar definition]
92. v1_funct_1(sK1) <- (7) [avatar component clause 90] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
93. 7 [avatar split clause 39,90]
95. 8 <=> v5_relat_1(sK1,sK0) [avatar definition]
97. v5_relat_1(sK1,sK0) <- (8) [avatar component clause 95] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
98. 8 [avatar split clause 38,95]
100. 9 <=> v1_relat_1(sK1) [avatar definition]
102. v1_relat_1(sK1) <- (9) [avatar component clause 100] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
103. 9 [avatar split clause 37,100]
145. ~v1_relat_1(sK1) <- (~1) [resolution 50,65] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
147. $false <- (~1, 9) [subsumption resolution 145,102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
148. 1 | ~9 [avatar contradiction clause 147]
151. ~v5_relat_1(sK1,sK0) | ~v1_relat_1(sK1) <- (~2) [resolution 58,69] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
153. ~v1_relat_1(sK1) <- (~2, 8) [subsumption resolution 151,97] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
154. $false <- (~2, 8, 9) [subsumption resolution 153,102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
155. 2 | ~8 | ~9 [avatar contradiction clause 154]
157. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~3) [resolution 73,56] {a:1,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
159. ~v1_relat_1(sK1) <- (~3, 7) [subsumption resolution 157,92] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
160. $false <- (~3, 7, 9) [subsumption resolution 159,102] {a:1,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
161. 3 | ~7 | ~9 [avatar contradiction clause 160]
167. ~v3_ordinal1(sK2) | ~v5_ordinal1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~4) [resolution 54,77] {a:1,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
170. ~v5_ordinal1(sK1) | ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~4, 5) [subsumption resolution 167,82] {a:1,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
171. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) <- (~4, 5, 6) [subsumption resolution 170,87] {a:1,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
172. ~v1_relat_1(sK1) <- (~4, 5, 6, 7) [subsumption resolution 171,92] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
173. $false <- (~4, 5, 6, 7, 9) [subsumption resolution 172,102] {a:1,w:0,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
174. 4 | ~5 | ~6 | ~7 | ~9 [avatar contradiction clause 173]
175. $false [avatar sat refutation 78,83,88,93,98,103,148,155,161,174] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:117,thDist:-117}
% SZS output end Proof for ordinal1__t33_ordinal1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 12
% Introduced skolems: 6
% 
% >>> Saturation
% Initial clauses: 25
% Generated clauses: 72
% Active clauses: 31
% Passive clauses: 35
% Final active clauses: 28
% Final passive clauses: 1
% Main loop iterations started: 34
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 9
% 
% >>> Deletion Inferences
% Simple tautologies: 4
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 10
% Unit resulting resolution: 4
% Forward superposition: 1
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 4
% Sat splitting refutations: 8
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 23
% SAT solver unit clauses: 8
% SAT solver binary clauses: 6
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

