i: [1,42,0,3,1,1,0,0] t58_zfmisc_1
a: [3,70,0,3,1,42]
s: 70
i: [1,43,0,4,1,1,0,0] t58_zfmisc_1
a: [3,65,0,4,1,43]
s: 65
i: [1,41,0,6,1,1,0,0] t58_zfmisc_1
a: [3,75,0,6,1,41]
s: 75
i: [1,44,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,80,0,2,1,44]
s: 80
i: [1,45,0,2,1,0,0,1] fc1_xboole_0
a: [3,85,0,2,1,45]
s: 85
i: [1,58,0,2,1,0,0,2] rc2_xboole_0
a: [3,95,0,2,1,58]
s: 95
i: [1,59,0,2,1,0,0,2] rc1_xboole_0
a: [3,100,0,2,1,59]
s: 100
i: [1,47,0,3,1,0,0,1] fc2_xboole_0
s: 47
i: [1,57,0,3,1,0,0,2] reflexivity_r1_tarski
s: 57
i: [1,46,0,3,1,0,0,255] d2_xboole_0
a: [3,90,0,3,1,46]
s: 90
r: 80
i: [1,61,0,4,1,0,0,1] l3_zfmisc_1
s: 61
i: [1,48,0,5,1,0,0,1] t4_boole
s: 48
i: [1,49,0,5,1,0,0,1] t3_boole
s: 49
i: [1,50,0,5,2,0,0,1] t6_boole
s: 50
i: [1,60,0,5,1,0,0,1] l3_zfmisc_1
s: 60
i: [1,56,0,7,3,0,0,1] t8_boole
s: 56
d: [2,106,1,6,2,0,44,61,50]
s: 106
i: [1,54,0,8,2,0,0,2] t37_xboole_1
s: 54
d: [2,108,1,7,2,1,44,54,75]
d: [2,112,1,4,1,1,19,108]
a: [3,116,1,4,1,112]
s: 116
i: [1,55,0,8,2,0,0,2] t37_xboole_1
s: 55
i: [1,51,0,11,3,0,0,1] l3_zfmisc_1
s: 51
d: [2,121,2,0,0,3,52,70,65,116,51]
e: 121
d: [2,125,2,7,2,1,40,51,116]
d: [2,127,2,4,1,2,21,125,70]
d: [2,128,2,0,0,3,21,127,65]
e: 128
f: 128,112,41,42,43
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t58_zfmisc_1
% SZS output start Proof for zfmisc_1__t58_zfmisc_1
1. ! [X0] : ! [X1] : ~(k1_tarski(X1) != X0 & k1_xboole_0 != X0 & k4_xboole_0(X0,k1_tarski(X1)) = k1_xboole_0) [input t58_zfmisc_1]
2. ~! [X0] : ! [X1] : ~(k1_tarski(X1) != X0 & k1_xboole_0 != X0 & k4_xboole_0(X0,k1_tarski(X1)) = k1_xboole_0) [negated conjecture 1]
11. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input l3_zfmisc_1]
15. ! [X0] : ! [X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [input t37_xboole_1]
21. ~! [X0,X1] : ~(k1_tarski(X1) != X0 & k1_xboole_0 != X0 & k4_xboole_0(X0,k1_tarski(X1)) = k1_xboole_0) [flattening 2]
22. ! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 11]
23. ! [X0,X1] : (k1_xboole_0 = k4_xboole_0(X0,X1) <=> r1_tarski(X0,X1)) [flattening 15]
29. ? [X0,X1] : (k1_tarski(X1) != X0 & k1_xboole_0 != X0 & k4_xboole_0(X0,k1_tarski(X1)) = k1_xboole_0) [ennf transformation 21]
32. ? [X0,X1] : (k1_tarski(X1) != X0 & k1_xboole_0 != X0 & k4_xboole_0(X0,k1_tarski(X1)) = k1_xboole_0) => (sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0 & k1_xboole_0 = k4_xboole_0(sK0,k1_tarski(sK1))) [choice axiom]
33. sK0 != k1_tarski(sK1) & k1_xboole_0 != sK0 & k1_xboole_0 = k4_xboole_0(sK0,k1_tarski(sK1)) [skolemisation 29,32]
34. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | ~r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 22]
35. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | ~r1_tarski(X0,k1_tarski(X1)))) [flattening 34]
36. ! [X0,X1] : ((k1_xboole_0 = k4_xboole_0(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | k1_xboole_0 != k4_xboole_0(X0,X1))) [nnf transformation 23]
41. k1_xboole_0 = k4_xboole_0(sK0,k1_tarski(sK1)) [cnf transformation 33] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
42. k1_xboole_0 != sK0 [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
43. sK0 != k1_tarski(sK1) [cnf transformation 33] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
51. ~r1_tarski(X0,k1_tarski(X1)) | k1_xboole_0 = X0 | k1_tarski(X1) = X0 [cnf transformation 35] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
54. k1_xboole_0 != k4_xboole_0(X0,X1) | r1_tarski(X0,X1) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
63. 1 <=> sK0 = k1_tarski(sK1) [avatar definition]
65. sK0 != k1_tarski(sK1) <- (~1) [avatar component clause 63] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
66. ~1 [avatar split clause 43,63]
68. 2 <=> k1_xboole_0 = sK0 [avatar definition]
70. k1_xboole_0 != sK0 <- (~2) [avatar component clause 68] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
71. ~2 [avatar split clause 42,68]
73. 3 <=> k1_xboole_0 = k4_xboole_0(sK0,k1_tarski(sK1)) [avatar definition]
75. k1_xboole_0 = k4_xboole_0(sK0,k1_tarski(sK1)) <- (3) [avatar component clause 73] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
76. 3 [avatar split clause 41,73]
108. k1_xboole_0 != k1_xboole_0 | r1_tarski(sK0,k1_tarski(sK1)) <- (3) [superposition 54,75] {a:1,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
112. r1_tarski(sK0,k1_tarski(sK1)) <- (3) [trivial inequality removal 108] {a:1,w:4,wCS:400,goal:1,sine:2,thAx:0,allAx:3,thDist:-3}
114. 9 <=> r1_tarski(sK0,k1_tarski(sK1)) [avatar definition]
116. r1_tarski(sK0,k1_tarski(sK1)) <- (9) [avatar component clause 114] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
117. 9 | ~3 [avatar split clause 112,73,114]
125. k1_xboole_0 = sK0 | sK0 = k1_tarski(sK1) <- (9) [resolution 51,116] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
127. sK0 = k1_tarski(sK1) <- (~2, 9) [subsumption resolution 125,70] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
128. $false <- (~1, ~2, 9) [subsumption resolution 127,65] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
129. 1 | 2 | ~9 [avatar contradiction clause 128]
130. $false [avatar sat refutation 66,71,76,117,129] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:38,thDist:-38}
% SZS output end Proof for zfmisc_1__t58_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 19
% >>> Preprocessing
% Introduced skolems: 4
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 49
% Active clauses: 20
% Passive clauses: 28
% Final active clauses: 19
% Final passive clauses: 7
% Main loop iterations started: 21
% 
% >>> Simplifying Inferences
% Trivial inequalities: 2
% Fw subsumption resolutions: 2
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Equational tautologies: 2
% Forward subsumptions: 3
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 4
% Backward superposition: 6
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 8
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% unit resulting resolution: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

