i: [1,30,0,4,2,1,0,0] t1_grfunc_1
a: [3,41,0,2,1,30]
s: 41
i: [1,28,0,2,1,1,0,0] t1_grfunc_1
a: [3,55,0,2,1,28]
s: 55
i: [1,27,0,2,1,1,0,0] t1_grfunc_1
a: [3,60,0,2,1,27]
s: 60
i: [1,29,0,3,1,1,0,0] t1_grfunc_1
a: [3,50,0,3,1,29]
s: 50
i: [1,37,0,2,1,0,0,1] rc1_funct_1
a: [3,65,0,2,1,37]
s: 65
i: [1,36,0,2,1,0,0,1] rc1_funct_1
a: [3,70,0,2,1,36]
s: 70
i: [1,35,0,3,1,0,0,1] reflexivity_r1_tarski
s: 35
i: [1,34,0,4,1,0,0,2] existence_m1_subset_1
s: 34
i: [1,33,0,7,2,0,0,1] t3_subset
s: 33
i: [1,31,0,8,3,0,0,2] cc2_relat_1
s: 31
i: [1,32,0,10,4,0,0,1] cc3_funct_1
s: 32
d: [2,72,1,6,2,0,40,31,34]
s: 72
d: [2,73,1,7,3,0,40,31,33]
s: 73
d: [2,76,2,0,0,3,52,41,60,50,73]
e: 76
d: [2,77,2,4,2,1,40,73,50]
d: [2,80,2,2,1,2,21,77,60]
d: [2,81,2,0,0,3,21,80,41]
e: 81
r: 41
a: [3,40,0,2,1,30]
s: 40
a: [3,45,0,2,1,30]
s: 45
d: [2,74,1,8,3,0,40,32,34]
s: 74
d: [2,75,1,9,4,0,40,32,33]
s: 75
d: [2,84,2,0,0,4,52,45,60,55,50,75]
e: 84
d: [2,85,2,6,3,1,40,75,50]
d: [2,88,2,4,2,2,21,85,55]
d: [2,89,2,2,1,3,21,88,60]
d: [2,90,2,0,0,4,21,89,45]
e: 90
f: 90,80,27,28,29,30
% Refutation found. Thanks to Tanya!
% SZS status Theorem for grfunc_1__t1_grfunc_1
% SZS output start Proof for grfunc_1__t1_grfunc_1
1. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (r1_tarski(X1,X0) => (v1_funct_1(X1) & v1_relat_1(X1)))) [input t1_grfunc_1]
2. ~! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (r1_tarski(X1,X0) => (v1_funct_1(X1) & v1_relat_1(X1)))) [negated conjecture 1]
6. ! [X0] : (v1_relat_1(X0) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_relat_1(X1))) [input cc2_relat_1]
7. ! [X0] : ((v1_funct_1(X0) & v1_relat_1(X0)) => ! [X1] : (m1_subset_1(X1,k1_zfmisc_1(X0)) => v1_funct_1(X1))) [input cc3_funct_1]
10. ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [input t3_subset]
11. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1)) [flattening 10]
13. ! [X0,X1] : (r1_tarski(X0,X1) => m1_subset_1(X0,k1_zfmisc_1(X1))) [unused predicate definition removal 11]
14. ? [X0] : (? [X1] : ((~v1_funct_1(X1) | ~v1_relat_1(X1)) & r1_tarski(X1,X0)) & (v1_funct_1(X0) & v1_relat_1(X0))) [ennf transformation 2]
15. ? [X0] : (? [X1] : ((~v1_funct_1(X1) | ~v1_relat_1(X1)) & r1_tarski(X1,X0)) & v1_funct_1(X0) & v1_relat_1(X0)) [flattening 14]
16. ! [X0] : (! [X1] : (v1_relat_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_relat_1(X0)) [ennf transformation 6]
17. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | (~v1_funct_1(X0) | ~v1_relat_1(X0))) [ennf transformation 7]
18. ! [X0] : (! [X1] : (v1_funct_1(X1) | ~m1_subset_1(X1,k1_zfmisc_1(X0))) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) [flattening 17]
19. ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 13]
20. ? [X0] : (? [X1] : ((~v1_funct_1(X1) | ~v1_relat_1(X1)) & r1_tarski(X1,X0)) & v1_funct_1(X0) & v1_relat_1(X0)) => (? [X1] : ((~v1_funct_1(X1) | ~v1_relat_1(X1)) & r1_tarski(X1,sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0)) [choice axiom]
21. ? [X1] : ((~v1_funct_1(X1) | ~v1_relat_1(X1)) & r1_tarski(X1,sK0)) => ((~v1_funct_1(sK1) | ~v1_relat_1(sK1)) & r1_tarski(sK1,sK0)) [choice axiom]
22. ((~v1_funct_1(sK1) | ~v1_relat_1(sK1)) & r1_tarski(sK1,sK0)) & v1_funct_1(sK0) & v1_relat_1(sK0) [skolemisation 15,21,20]
27. v1_relat_1(sK0) [cnf transformation 22] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
28. v1_funct_1(sK0) [cnf transformation 22] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. r1_tarski(sK1,sK0) [cnf transformation 22] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. ~v1_funct_1(sK1) | ~v1_relat_1(sK1) [cnf transformation 22] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 16] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
32. ~m1_subset_1(X1,k1_zfmisc_1(X0)) | v1_funct_1(X1) | ~v1_funct_1(X0) | ~v1_relat_1(X0) [cnf transformation 18] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
33. m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1) [cnf transformation 19] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
39. 1 <=> v1_relat_1(sK1) [avatar definition]
43. 2 <=> v1_funct_1(sK1) [avatar definition]
45. ~v1_funct_1(sK1) <- (~2) [avatar component clause 43] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
46. ~1 | ~2 [avatar split clause 30,43,39]
48. 3 <=> r1_tarski(sK1,sK0) [avatar definition]
50. r1_tarski(sK1,sK0) <- (3) [avatar component clause 48] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
51. 3 [avatar split clause 29,48]
53. 4 <=> v1_funct_1(sK0) [avatar definition]
55. v1_funct_1(sK0) <- (4) [avatar component clause 53] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
56. 4 [avatar split clause 28,53]
58. 5 <=> v1_relat_1(sK0) [avatar definition]
60. v1_relat_1(sK0) <- (5) [avatar component clause 58] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
61. 5 [avatar split clause 27,58]
73. ~r1_tarski(X1,X2) | ~v1_relat_1(X2) | v1_relat_1(X1) [resolution 31,33] {a:1,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
75. ~r1_tarski(X1,X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v1_funct_1(X1) [resolution 32,33] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
77. ~v1_relat_1(sK0) | v1_relat_1(sK1) <- (3) [resolution 73,50] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
80. v1_relat_1(sK1) <- (3, 5) [subsumption resolution 77,60] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
83. 1 | ~3 | ~5 [avatar split clause 80,58,48,39]
85. ~v1_funct_1(sK0) | ~v1_relat_1(sK0) | v1_funct_1(sK1) <- (3) [resolution 75,50] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
88. ~v1_relat_1(sK0) | v1_funct_1(sK1) <- (3, 4) [subsumption resolution 85,55] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
89. v1_funct_1(sK1) <- (3, 4, 5) [subsumption resolution 88,60] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
90. $false <- (~2, 3, 4, 5) [subsumption resolution 89,45] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
91. 2 | ~3 | ~4 | ~5 [avatar contradiction clause 90]
92. $false [avatar sat refutation 46,51,56,61,83,91] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:66,thDist:-66}
% SZS output end Proof for grfunc_1__t1_grfunc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% >>> Preprocessing
% Introduced skolems: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 11
% Generated clauses: 35
% Active clauses: 17
% Passive clauses: 17
% Final active clauses: 16
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 5
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% 
% >>> Generating Inferences
% Binary resolution: 8
% Unit resulting resolution: 2
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 2
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 13
% SAT solver unit clauses: 5
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.001 s

