i: [1,69,0,3,1,1,0,0] t124_zfmisc_1
a: [3,116,0,3,1,69]
s: 116
i: [1,67,0,3,1,1,0,0] t124_zfmisc_1
a: [3,126,0,3,1,67]
s: 126
i: [1,70,0,6,1,1,0,0] t124_zfmisc_1
a: [3,111,0,6,1,70]
s: 111
i: [1,68,0,6,1,1,0,0] t124_zfmisc_1
a: [3,121,0,6,1,68]
s: 121
i: [1,71,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,131,0,2,1,71]
s: 131
i: [1,72,0,2,1,0,0,2] fc1_xboole_0
a: [3,136,0,2,1,72]
s: 136
i: [1,97,0,2,1,0,0,2] rc2_xboole_0
a: [3,146,0,2,1,97]
s: 146
i: [1,98,0,2,1,0,0,2] rc1_xboole_0
a: [3,151,0,2,1,98]
s: 151
i: [1,99,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,156,0,2,1,99]
s: 156
i: [1,74,0,3,1,0,0,1] fc2_xboole_0
s: 74
i: [1,93,0,3,1,0,0,1] reflexivity_r1_tarski
s: 93
i: [1,73,0,3,1,0,0,255] d2_xboole_0
a: [3,141,0,3,1,73]
s: 141
r: 131
i: [1,101,0,4,1,0,0,1] d1_tarski
s: 101
i: [1,104,0,4,1,0,0,1] l3_zfmisc_1
s: 104
i: [1,75,0,5,1,0,0,2] t2_boole
s: 75
i: [1,76,0,5,2,0,0,2] t6_boole
s: 76
i: [1,85,0,5,2,0,0,1] t7_boole
s: 85
i: [1,94,0,5,1,0,0,1] idempotence_k3_xboole_0
s: 94
d: [2,163,1,2,1,1,40,85,116]
a: [3,168,1,2,1,163]
s: 168
i: [1,103,0,5,1,0,0,1] l3_zfmisc_1
s: 103
i: [1,96,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 96
d: [2,170,1,3,1,1,40,96,116]
a: [3,175,1,3,1,170]
s: 175
i: [1,84,0,7,3,0,0,1] t8_boole
s: 84
d: [2,171,1,4,1,0,40,96,101]
s: 171
i: [1,95,0,7,1,0,0,1] commutativity_k3_xboole_0
s: 95
i: [1,102,0,7,2,0,0,1] d1_tarski
s: 102
d: [2,177,1,5,1,0,44,95,75]
s: 177
i: [1,106,0,8,2,0,0,1] d4_xboole_0
s: 106
d: [2,187,1,7,2,1,44,106,121]
s: 187
i: [1,107,0,8,2,0,0,1] d4_xboole_0
s: 107
i: [1,86,0,10,2,0,0,1] t26_xboole_1
s: 86
d: [2,193,2,3,1,1,40,187,101]
a: [3,197,2,3,1,193]
s: 197
i: [1,81,0,11,3,0,0,1] l3_zfmisc_1
s: 81
d: [2,220,3,2,1,1,40,197,85]
a: [3,229,3,2,1,220]
s: 229
i: [1,105,0,11,3,0,0,1] d4_xboole_0
s: 105
i: [1,79,0,14,3,0,0,1] d1_tarski
s: 79
d: [2,219,3,3,1,1,40,197,96]
a: [3,224,3,3,1,219]
s: 224
i: [1,80,0,14,3,0,0,1] d1_tarski
s: 80
d: [2,199,1,7,2,1,44,107,121]
s: 199
i: [1,90,0,17,3,0,0,1] d4_xboole_0
s: 90
i: [1,91,0,17,3,0,0,1] d4_xboole_0
s: 91
d: [2,251,2,3,1,1,40,199,101]
a: [3,256,2,3,1,251]
s: 256
i: [1,92,0,23,4,0,0,1] d4_xboole_0
s: 92
d: [2,275,3,2,1,1,40,256,85]
a: [3,284,3,2,1,275]
s: 284
d: [2,205,1,9,2,1,44,86,121]
s: 205
d: [2,211,1,9,2,1,44,86,121]
s: 211
d: [2,292,2,7,2,1,44,205,94]
a: [3,299,2,3,1,292]
s: 299
d: [2,238,1,10,3,1,44,105,121]
s: 238
d: [2,295,2,7,2,1,44,205,177]
a: [3,308,2,3,1,295]
s: 308
d: [2,178,1,5,1,0,44,95,75]
s: 178
d: [2,179,1,5,1,0,44,75,95]
s: 179
d: [2,316,2,7,2,1,44,211,94]
a: [3,323,2,3,1,316]
s: 323
d: [2,180,1,5,1,0,44,75,95]
s: 180
d: [2,274,3,3,1,1,40,256,96]
a: [3,279,3,3,1,274]
s: 279
d: [2,161,1,6,2,0,44,104,76]
s: 161
d: [2,200,1,6,2,0,44,107,75]
s: 200
d: [2,334,2,10,2,1,40,238,171]
a: [3,339,2,5,1,334]
s: 339
d: [2,159,1,7,3,0,44,76,76]
s: 159
d: [2,345,3,5,2,1,44,308,76]
s: 345
d: [2,160,1,7,2,0,44,75,76]
s: 160
d: [2,190,1,8,2,0,44,106,95]
s: 190
d: [2,293,2,9,2,1,44,205,95]
s: 293
d: [2,191,1,8,2,0,44,106,95]
s: 191
d: [2,294,2,9,2,1,44,205,95]
s: 294
d: [2,207,1,8,2,0,44,86,94]
s: 207
d: [2,213,1,8,2,0,44,86,94]
s: 213
d: [2,374,4,7,2,1,40,207,345]
s: 374
d: [2,208,1,10,2,0,44,86,95]
s: 208
d: [2,392,5,7,2,1,44,374,95]
s: 392
d: [2,209,1,10,2,0,44,86,95]
s: 209
d: [2,214,1,10,2,0,44,86,95]
s: 214
d: [2,393,5,7,2,1,44,374,95]
s: 393
d: [2,215,1,10,2,0,44,86,95]
s: 215
d: [2,317,2,9,2,1,44,211,95]
s: 317
d: [2,236,1,10,3,0,40,105,85]
s: 236
d: [2,235,1,11,3,0,40,105,96]
s: 235
d: [2,318,2,9,2,1,44,211,95]
s: 318
d: [2,246,1,11,3,0,40,79,85]
s: 246
d: [2,331,2,9,3,1,40,238,102]
s: 331
d: [2,265,1,11,2,0,42,90]
s: 265
d: [2,273,1,11,2,0,42,91]
s: 273
d: [2,396,2,9,2,1,44,208,121]
s: 396
d: [2,259,1,13,3,0,40,90,85]
s: 259
d: [2,497,3,7,2,1,44,396,177]
a: [3,501,3,3,1,497]
s: 501
d: [2,267,1,13,3,0,40,91,85]
s: 267
d: [2,245,1,14,3,0,40,79,96]
s: 245
d: [2,515,4,5,2,1,44,501,76]
s: 515
d: [2,247,1,17,3,0,40,79,102]
s: 247
d: [2,531,5,7,2,1,40,515,207]
s: 531
d: [2,258,1,17,3,0,40,90,96]
s: 258
d: [2,266,1,17,3,0,40,91,96]
s: 266
d: [2,546,6,7,2,1,44,531,95]
s: 546
d: [2,249,1,20,3,0,40,79,107]
s: 249
d: [2,547,6,7,2,1,44,531,95]
s: 547
d: [2,250,1,20,3,0,40,79,106]
s: 250
d: [2,260,1,20,3,0,40,90,102]
s: 260
d: [2,403,2,9,2,1,44,208,121]
s: 403
d: [2,268,1,20,3,0,40,91,102]
s: 268
d: [2,426,2,9,2,1,44,214,121]
s: 426
d: [2,263,1,23,3,0,40,90,107]
s: 263
d: [2,264,1,23,3,0,40,90,106]
s: 264
d: [2,495,3,9,2,1,44,396,95]
s: 495
d: [2,271,1,23,3,0,40,91,107]
s: 271
d: [2,496,3,9,2,1,44,396,95]
s: 496
d: [2,272,1,23,3,0,40,91,106]
s: 272
d: [2,288,1,39,5,0,40,92,105]
s: 288
d: [2,623,3,9,2,1,44,426,95]
s: 623
d: [2,332,2,10,3,1,40,238,96]
s: 332
d: [2,742,3,16,4,1,40,332,238]
d: [2,744,3,8,2,1,18,742]
a: [3,748,3,4,1,744]
s: 748
d: [2,453,2,10,3,1,44,235,121]
s: 453
d: [2,314,2,14,3,1,40,211,81]
s: 314
a: [3,303,2,4,1,292]
s: 303
a: [3,327,2,4,1,316]
s: 327
r: 327
a: [3,326,2,4,1,316]
s: 326
d: [2,475,2,15,2,1,40,265,199]
s: 475
d: [2,476,2,15,2,1,40,265,187]
s: 476
d: [2,759,3,13,3,2,44,111,314]
d: [2,797,3,8,2,2,19,759]
d: [2,798,3,5,1,3,21,797,126]
a: [3,802,3,5,1,798]
s: 802
r: 111
d: [2,487,2,15,2,1,40,273,199]
s: 487
d: [2,856,4,7,2,2,44,205,802]
a: [3,888,4,3,1,856]
s: 888
d: [2,488,2,15,2,1,40,273,187]
s: 488
d: [2,509,2,15,3,1,40,259,199]
s: 509
d: [2,866,4,6,2,1,44,213,802]
a: [3,895,4,3,1,866]
s: 895
d: [2,510,2,15,3,1,40,259,187]
s: 510
d: [2,854,6,5,2,2,44,531,802]
d: [2,879,6,3,1,3,21,854,136]
a: [3,883,6,3,1,879]
s: 883
d: [2,522,2,15,3,1,40,267,199]
s: 522
d: [2,523,2,15,3,1,40,267,187]
s: 523
d: [2,794,3,12,3,1,46,314]
a: [3,830,3,4,1,794]
s: 830
d: [2,248,2,17,3,1,40,79,187]
s: 248
d: [2,870,4,8,3,1,44,236,802]
d: [2,901,4,6,2,2,21,870,136]
s: 901
d: [2,945,5,0,0,4,52,116,197,901]
e: 945
d: [2,946,5,3,1,3,40,901,197]
d: [2,967,5,0,0,4,21,946,116]
e: 967
f: 967,798,193,72,67,68,69,70
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t124_zfmisc_1
% SZS output start Proof for zfmisc_1__t124_zfmisc_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1)) => k1_tarski(X3) = k3_xboole_0(X0,X2)) [input t124_zfmisc_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1)) => k1_tarski(X3) = k3_xboole_0(X0,X2)) [negated conjecture 1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
9. ! [X0] : ! [X1] : ! [X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [input d4_xboole_0]
14. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
18. ! [X0] : ! [X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [input l3_zfmisc_1]
23. ! [X0] : ! [X1] : ! [X2] : (r1_tarski(X0,X1) => r1_tarski(k3_xboole_0(X0,X2),k3_xboole_0(X1,X2))) [input t26_xboole_1]
26. ! [X0] : ! [X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [input t7_boole]
28. ~! [X0,X1,X2,X3] : ((r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1)) => k1_tarski(X3) = k3_xboole_0(X0,X2)) [flattening 2]
29. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
30. ! [X0,X1] : (r1_tarski(X0,k1_tarski(X1)) <=> (k1_tarski(X1) = X0 | k1_xboole_0 = X0)) [flattening 18]
32. ! [X0,X1] : ~(v1_xboole_0(X1) & r2_hidden(X0,X1)) [flattening 26]
33. ! [X0,X1,X2] : (r1_tarski(X0,X1) => r1_tarski(k3_xboole_0(X0,X2),k3_xboole_0(X1,X2))) [flattening 23]
34. ! [X0,X1,X2] : (k3_xboole_0(X0,X1) = X2 <=> ! [X3] : (r2_hidden(X3,X2) <=> (r2_hidden(X3,X1) & r2_hidden(X3,X0)))) [flattening 9]
41. ? [X0,X1,X2,X3] : (k1_tarski(X3) != k3_xboole_0(X0,X2) & (r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1))) [ennf transformation 28]
42. ? [X0,X1,X2,X3] : (k1_tarski(X3) != k3_xboole_0(X0,X2) & r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1)) [flattening 41]
45. ! [X0,X1] : (~v1_xboole_0(X1) | ~r2_hidden(X0,X1)) [ennf transformation 32]
46. ! [X0,X1,X2] : (r1_tarski(k3_xboole_0(X0,X2),k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X1)) [ennf transformation 33]
48. ? [X0,X1,X2,X3] : (k1_tarski(X3) != k3_xboole_0(X0,X2) & r2_hidden(X3,X0) & k3_xboole_0(X1,X2) = k1_tarski(X3) & r1_tarski(X0,X1)) => (k1_tarski(sK3) != k3_xboole_0(sK0,sK2) & r2_hidden(sK3,sK0) & k1_tarski(sK3) = k3_xboole_0(sK1,sK2) & r1_tarski(sK0,sK1)) [choice axiom]
49. k1_tarski(sK3) != k3_xboole_0(sK0,sK2) & r2_hidden(sK3,sK0) & k1_tarski(sK3) = k3_xboole_0(sK1,sK2) & r1_tarski(sK0,sK1) [skolemisation 42,48]
50. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 29]
51. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 50]
52. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
53. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 51,52]
54. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & ((k1_tarski(X1) = X0 | k1_xboole_0 = X0) | ~r1_tarski(X0,k1_tarski(X1)))) [nnf transformation 30]
55. ! [X0,X1] : ((r1_tarski(X0,k1_tarski(X1)) | (k1_tarski(X1) != X0 & k1_xboole_0 != X0)) & (k1_tarski(X1) = X0 | k1_xboole_0 = X0 | ~r1_tarski(X0,k1_tarski(X1)))) [flattening 54]
56. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : (((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | (~r2_hidden(X3,X1) | ~r2_hidden(X3,X0))) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [nnf transformation 34]
57. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X3] : ((r2_hidden(X3,X2) | ~r2_hidden(X3,X1) | ~r2_hidden(X3,X0)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | ~r2_hidden(X3,X2))) | k3_xboole_0(X0,X1) != X2)) [flattening 56]
58. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [rectify 57]
59. ! [X2,X1,X0] : (? [X3] : ((~r2_hidden(X3,X1) | ~r2_hidden(X3,X0) | ~r2_hidden(X3,X2)) & ((r2_hidden(X3,X1) & r2_hidden(X3,X0)) | r2_hidden(X3,X2))) => ((~r2_hidden(sK5(X0,X1,X2),X1) | ~r2_hidden(sK5(X0,X1,X2),X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & ((r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK5(X0,X1,X2),X0)) | r2_hidden(sK5(X0,X1,X2),X2)))) [choice axiom]
60. ! [X0,X1,X2] : ((k3_xboole_0(X0,X1) = X2 | ((~r2_hidden(sK5(X0,X1,X2),X1) | ~r2_hidden(sK5(X0,X1,X2),X0) | ~r2_hidden(sK5(X0,X1,X2),X2)) & ((r2_hidden(sK5(X0,X1,X2),X1) & r2_hidden(sK5(X0,X1,X2),X0)) | r2_hidden(sK5(X0,X1,X2),X2)))) & (! [X4] : ((r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0)) & ((r2_hidden(X4,X1) & r2_hidden(X4,X0)) | ~r2_hidden(X4,X2))) | k3_xboole_0(X0,X1) != X2)) [skolemisation 58,59]
67. r1_tarski(sK0,sK1) [cnf transformation 49] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
68. k1_tarski(sK3) = k3_xboole_0(sK1,sK2) [cnf transformation 49] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
69. r2_hidden(sK3,sK0) [cnf transformation 49] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. k1_tarski(sK3) != k3_xboole_0(sK0,sK2) [cnf transformation 49] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
72. v1_xboole_0(k1_xboole_0) [cnf transformation 14] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
78. r2_hidden(X3,X1) | X0 != X3 | k1_tarski(X0) != X1 [cnf transformation 53] {a:0,w:10,wCS:4000,sine:1,thAx:0,allAx:2,thDist:-2}
81. ~r1_tarski(X0,k1_tarski(X1)) | k1_xboole_0 = X0 | k1_tarski(X1) = X0 [cnf transformation 55] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
85. ~r2_hidden(X0,X1) | ~v1_xboole_0(X1) [cnf transformation 45] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
86. r1_tarski(k3_xboole_0(X0,X2),k3_xboole_0(X1,X2)) | ~r1_tarski(X0,X1) [cnf transformation 46] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
88. r2_hidden(X4,X1) | ~r2_hidden(X4,X2) | k3_xboole_0(X0,X1) != X2 [cnf transformation 60] {a:0,w:11,wCS:4400,sine:1,thAx:0,allAx:2,thDist:-2}
89. r2_hidden(X4,X2) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) | k3_xboole_0(X0,X1) != X2 [cnf transformation 60] {a:0,w:14,wCS:5600,sine:1,thAx:0,allAx:2,thDist:-2}
100. r2_hidden(X3,X1) | k1_tarski(X3) != X1 [equality resolution 78] {a:0,w:7,wCS:2800,sine:1,thAx:0,allAx:2,thDist:-2}
101. r2_hidden(X3,k1_tarski(X3)) [equality resolution 100] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
105. r2_hidden(X4,k3_xboole_0(X0,X1)) | ~r2_hidden(X4,X1) | ~r2_hidden(X4,X0) [equality resolution 89] {a:0,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
106. ~r2_hidden(X4,k3_xboole_0(X0,X1)) | r2_hidden(X4,X1) [equality resolution 88] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
109. 1 <=> k1_tarski(sK3) = k3_xboole_0(sK0,sK2) [avatar definition]
111. k1_tarski(sK3) != k3_xboole_0(sK0,sK2) <- (~1) [avatar component clause 109] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
112. ~1 [avatar split clause 70,109]
114. 2 <=> r2_hidden(sK3,sK0) [avatar definition]
116. r2_hidden(sK3,sK0) <- (2) [avatar component clause 114] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
117. 2 [avatar split clause 69,114]
119. 3 <=> k1_tarski(sK3) = k3_xboole_0(sK1,sK2) [avatar definition]
121. k1_tarski(sK3) = k3_xboole_0(sK1,sK2) <- (3) [avatar component clause 119] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
122. 3 [avatar split clause 68,119]
124. 4 <=> r1_tarski(sK0,sK1) [avatar definition]
126. r1_tarski(sK0,sK1) <- (4) [avatar component clause 124] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
127. 4 [avatar split clause 67,124]
134. 6 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
136. v1_xboole_0(k1_xboole_0) <- (6) [avatar component clause 134] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
137. 6 [avatar split clause 72,134]
187. ~r2_hidden(X0,k1_tarski(sK3)) | r2_hidden(X0,sK2) <- (3) [superposition 106,121] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
193. r2_hidden(sK3,sK2) <- (3) [resolution 187,101] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. 13 <=> r2_hidden(sK3,sK2) [avatar definition]
197. r2_hidden(sK3,sK2) <- (13) [avatar component clause 195] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
198. 13 | ~3 [avatar split clause 193,119,195]
211. r1_tarski(k3_xboole_0(X0,sK2),k1_tarski(sK3)) | ~r1_tarski(X0,sK1) <- (3) [superposition 86,121] {a:1,w:9,wCS:900,nSel:1,goal:1,sine:1,thAx:0,allAx:3,thDist:-3}
236. ~v1_xboole_0(k3_xboole_0(X11,X10)) | ~r2_hidden(X9,X11) | ~r2_hidden(X9,X10) [resolution 105,85] {a:1,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
314. k1_tarski(sK3) = k3_xboole_0(X0,sK2) | k1_xboole_0 = k3_xboole_0(X0,sK2) | ~r1_tarski(X0,sK1) <- (3) [resolution 211,81] {a:2,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
759. k1_tarski(sK3) != k1_tarski(sK3) | k1_xboole_0 = k3_xboole_0(sK0,sK2) | ~r1_tarski(sK0,sK1) <- (~1, 3) [superposition 111,314] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
797. k1_xboole_0 = k3_xboole_0(sK0,sK2) | ~r1_tarski(sK0,sK1) <- (~1, 3) [trivial inequality removal 759] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
798. k1_xboole_0 = k3_xboole_0(sK0,sK2) <- (~1, 3, 4) [subsumption resolution 797,126] {a:3,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
800. 30 <=> k1_xboole_0 = k3_xboole_0(sK0,sK2) [avatar definition]
802. k1_xboole_0 = k3_xboole_0(sK0,sK2) <- (30) [avatar component clause 800] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
803. 30 | 1 | ~3 | ~4 [avatar split clause 798,124,119,109,800]
870. ~v1_xboole_0(k1_xboole_0) | ~r2_hidden(X10,sK0) | ~r2_hidden(X10,sK2) <- (30) [superposition 236,802] {a:4,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
901. ~r2_hidden(X10,sK2) | ~r2_hidden(X10,sK0) <- (6, 30) [subsumption resolution 870,136] {a:4,w:6,wCS:600,nSel:1,goal:1,sine:1,thAx:0,allAx:9,thDist:-9}
946. ~r2_hidden(sK3,sK0) <- (6, 13, 30) [resolution 901,197] {a:5,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
967. $false <- (2, 6, 13, 30) [subsumption resolution 946,116] {a:5,w:0,goal:1,sine:1,thAx:0,allAx:15,thDist:-15}
968. ~2 | ~6 | ~13 | ~30 [avatar contradiction clause 967]
969. $false [avatar sat refutation 112,117,122,127,137,198,803,968] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:317,thDist:-317}
% SZS output end Proof for zfmisc_1__t124_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 26
% >>> Preprocessing
% Introduced skolems: 9
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 33
% Generated clauses: 716
% Active clauses: 112
% Passive clauses: 354
% Final active clauses: 109
% Final passive clauses: 225
% Main loop iterations started: 129
% 
% >>> Simplifying Inferences
% Duplicate literals: 40
% Trivial inequalities: 2
% Fw subsumption resolutions: 20
% Fw demodulations: 28
% Bw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 24
% Equational tautologies: 4
% Forward subsumptions: 203
% 
% >>> Generating Inferences
% Binary resolution: 247
% Unit resulting resolution: 1
% Factoring: 6
% Forward superposition: 217
% Backward superposition: 83
% Self superposition: 1
% Equality factoring: 2
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 30
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 42
% SAT solver unit clauses: 10
% SAT solver binary clauses: 9
% 
% Memory used [KB]: 9850
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% preprocessing: 0.001 s
% forward superposition index maintenance: 0.001 s
% splitting component index maintenance: 0.001 s
% forward subsumption: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.012 s ( own 0.005 s ) 

