i: [1,40,0,2,1,1,0,0] t56_relat_1
a: [3,80,0,2,1,40]
s: 80
i: [1,39,0,2,1,1,0,0] t56_relat_1
a: [3,85,0,2,1,39]
s: 85
i: [1,38,0,2,1,1,0,0] t56_relat_1
a: [3,90,0,2,1,38]
s: 90
i: [1,44,0,3,1,1,0,0] t56_relat_1
a: [3,60,0,3,1,44]
s: 60
i: [1,41,0,4,1,1,0,0] t56_relat_1
a: [3,75,0,4,1,41]
s: 75
i: [1,43,0,6,1,1,0,0] t56_relat_1
a: [3,65,0,6,1,43]
s: 65
i: [1,42,0,7,1,1,0,0] t56_relat_1
a: [3,70,0,7,1,42]
s: 70
i: [1,45,0,3,1,0,0,1] dt_k4_relat_1
s: 45
i: [1,55,0,3,1,0,0,1] reflexivity_r1_tarski
s: 55
i: [1,54,0,4,1,0,0,255] existence_m1_subset_1
s: 54
i: [1,56,0,4,1,0,0,1] dt_k3_relat_1
s: 56
i: [1,46,0,5,1,0,0,1] rd1_relat_1
s: 46
i: [1,47,0,5,1,0,0,1] rd2_relat_1
s: 47
i: [1,52,0,7,2,0,0,255] t3_subset
s: 52
i: [1,53,0,7,2,0,0,255] t3_subset
s: 53
i: [1,49,0,8,3,0,0,255] cc2_relat_1
s: 49
d: [2,94,1,5,1,0,40,52,54]
s: 94
i: [1,50,0,12,3,0,0,1] t53_relat_1
s: 50
d: [2,96,1,6,2,0,40,49,54]
s: 96
i: [1,51,0,12,3,0,0,1] t51_relat_1
s: 51
i: [1,48,0,17,4,0,0,1] t36_relat_1
s: 48
d: [2,114,1,16,4,1,44,48,65]
d: [2,126,1,14,3,2,21,114,80]
d: [2,127,1,12,2,3,21,126,85]
s: 127
d: [2,113,1,17,4,1,44,48,70]
d: [2,124,1,15,3,2,21,113,85]
d: [2,125,1,13,2,3,21,124,90]
s: 125
d: [2,132,2,18,4,3,44,127,50]
d: [2,135,2,15,3,3,21,132,45]
d: [2,136,2,13,2,3,21,135,80]
d: [2,137,2,12,2,3,22,136,65]
s: 137
d: [2,97,1,7,3,0,40,49,53]
s: 97
d: [2,99,1,21,5,0,44,51,50]
d: [2,102,1,18,4,0,21,99,45]
d: [2,103,1,15,3,0,21,102,45]
d: [2,104,1,13,3,0,22,103,46]
d: [2,105,1,11,3,0,22,104,47]
s: 105
d: [2,166,2,5,2,1,40,97,75]
a: [3,176,2,2,1,166]
s: 176
d: [2,101,1,21,5,0,44,50,51]
d: [2,106,1,18,4,0,21,101,45]
d: [2,107,1,15,3,0,21,106,45]
d: [2,108,1,13,3,0,22,107,47]
d: [2,109,1,11,3,0,22,108,46]
s: 109
d: [2,156,3,17,4,3,44,137,51]
d: [2,160,3,14,3,3,21,156,45]
d: [2,161,3,12,3,3,22,160,46]
s: 161
d: [2,110,1,23,6,0,44,48,50]
d: [2,120,1,21,5,0,18,110]
d: [2,121,1,18,4,0,21,120,45]
s: 121
d: [2,112,1,23,6,0,44,48,51]
d: [2,118,1,21,5,0,18,112]
d: [2,123,1,18,4,0,21,118,45]
s: 123
d: [2,215,4,19,4,3,44,137,161]
a: [3,231,4,4,1,215]
s: 231
d: [2,115,1,27,6,0,44,48,50]
d: [2,128,1,24,5,0,21,115,45]
d: [2,129,1,20,4,0,21,128,56]
s: 129
d: [2,224,4,21,4,5,44,125,161]
a: [3,253,4,4,1,224]
s: 253
d: [2,117,1,27,6,0,44,50,48]
d: [2,130,1,23,5,0,21,117,56]
d: [2,131,1,20,4,0,21,130,45]
s: 131
d: [2,111,1,29,7,0,44,48,48]
d: [2,119,1,27,6,0,18,111]
d: [2,122,1,23,5,0,21,119,56]
s: 122
d: [2,223,4,19,4,3,44,137,161]
d: [2,225,4,15,3,3,18,223]
a: [3,240,4,7,2,225]
s: 240
r: 161
d: [2,368,2,22,5,3,44,129,127]
d: [2,403,2,20,4,3,21,368,80]
d: [2,404,2,18,3,3,21,403,85]
d: [2,405,2,15,2,3,21,404,45]
d: [2,406,2,14,2,3,22,405,65]
d: [2,407,2,13,2,3,22,406,65]
d: [2,408,2,11,2,3,22,407,47]
s: 408
d: [2,209,4,12,3,3,44,46,161]
d: [2,226,4,10,3,3,22,209,46]
s: 226
d: [2,378,2,22,5,3,44,127,129]
d: [2,422,2,19,4,3,21,378,45]
d: [2,423,2,17,3,3,21,422,80]
d: [2,424,2,15,2,3,21,423,85]
d: [2,425,2,14,2,3,22,424,65]
d: [2,426,2,13,2,3,22,425,65]
d: [2,427,2,11,2,3,22,426,47]
s: 427
d: [2,315,2,16,4,1,44,123,65]
d: [2,340,2,14,3,2,21,315,80]
d: [2,341,2,12,2,3,21,340,85]
s: 341
d: [2,551,3,17,3,5,44,341,125]
d: [2,560,3,14,2,5,21,551,45]
a: [3,564,3,5,1,560]
s: 564
d: [2,142,2,16,4,3,44,51,125]
d: [2,147,2,14,3,3,18,142]
s: 147
d: [2,584,3,13,3,4,44,147,65]
d: [2,597,3,8,2,4,21,584,55]
d: [2,598,3,6,1,5,21,597,80]
a: [3,602,3,6,1,598]
s: 602
d: [2,638,4,9,3,1,44,602,50]
d: [2,645,4,6,2,2,21,638,60]
d: [2,646,4,2,1,3,21,645,75]
d: [2,647,4,0,0,4,21,646,90]
e: 647
d: [2,640,4,9,3,1,44,50,602]
d: [2,650,4,6,2,2,21,640,60]
d: [2,651,4,2,1,3,21,650,75]
d: [2,652,4,0,0,4,21,651,90]
e: 652
f: 652,598,38,39,40,41,42,43,44
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t56_relat_1
% SZS output start Proof for relat_1__t56_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0)) => X1 = X3)))) [input t56_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0)) => X1 = X3)))) [negated conjecture 1]
13. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
14. ! [X0] : (v1_relat_1(X0) => ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2))))) [input t36_relat_1]
16. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X1),X0) => k3_relat_1(k4_relat_1(X0),X1) = X1)) [input t51_relat_1]
17. ! [X0] : ! [X1] : (v1_relat_1(X1) => (r1_tarski(k10_xtuple_0(X1),X0) => k3_relat_1(X1,k4_relat_1(X0)) = X1)) [input t53_relat_1]
18. ~! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => ! [X3] : (v1_relat_1(X3) => ((k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0)) => X1 = X3)))) [flattening 2]
19. ! [X0,X1] : (v1_relat_1(X1) => (r1_tarski(k10_xtuple_0(X1),X0) => k3_relat_1(X1,k4_relat_1(X0)) = X1)) [flattening 17]
20. ! [X0,X1] : (v1_relat_1(X1) => (r1_tarski(k9_xtuple_0(X1),X0) => k3_relat_1(k4_relat_1(X0),X1) = X1)) [flattening 16]
22. ! [X0] : r1_tarski(X0,X0) [rectify 13]
23. ? [X0,X1] : (? [X2] : (? [X3] : ((X1 != X3 & (k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0))) & v1_relat_1(X3)) & v1_relat_1(X2)) & v1_relat_1(X1)) [ennf transformation 18]
24. ? [X0,X1] : (? [X2] : (? [X3] : (X1 != X3 & k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0) & v1_relat_1(X3)) & v1_relat_1(X2)) & v1_relat_1(X1)) [flattening 23]
25. ! [X0] : (! [X1] : (! [X2] : (k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2)) | ~v1_relat_1(X2)) | ~v1_relat_1(X1)) | ~v1_relat_1(X0)) [ennf transformation 14]
27. ! [X0,X1] : ((k3_relat_1(X1,k4_relat_1(X0)) = X1 | ~r1_tarski(k10_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 19]
28. ! [X0,X1] : (k3_relat_1(X1,k4_relat_1(X0)) = X1 | ~r1_tarski(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1)) [flattening 27]
29. ! [X0,X1] : ((k3_relat_1(k4_relat_1(X0),X1) = X1 | ~r1_tarski(k9_xtuple_0(X1),X0)) | ~v1_relat_1(X1)) [ennf transformation 20]
30. ! [X0,X1] : (k3_relat_1(k4_relat_1(X0),X1) = X1 | ~r1_tarski(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1)) [flattening 29]
31. ? [X0,X1] : (? [X2] : (? [X3] : (X1 != X3 & k3_relat_1(X2,X3) = k4_relat_1(X0) & k3_relat_1(X1,X2) = k4_relat_1(k9_xtuple_0(X3)) & r1_tarski(k10_xtuple_0(X1),X0) & v1_relat_1(X3)) & v1_relat_1(X2)) & v1_relat_1(X1)) => (? [X2] : (? [X3] : (sK1 != X3 & k3_relat_1(X2,X3) = k4_relat_1(sK0) & k4_relat_1(k9_xtuple_0(X3)) = k3_relat_1(sK1,X2) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(X3)) & v1_relat_1(X2)) & v1_relat_1(sK1)) [choice axiom]
32. ? [X2] : (? [X3] : (sK1 != X3 & k3_relat_1(X2,X3) = k4_relat_1(sK0) & k4_relat_1(k9_xtuple_0(X3)) = k3_relat_1(sK1,X2) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(X3)) & v1_relat_1(X2)) => (? [X3] : (sK1 != X3 & k4_relat_1(sK0) = k3_relat_1(sK2,X3) & k4_relat_1(k9_xtuple_0(X3)) = k3_relat_1(sK1,sK2) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(X3)) & v1_relat_1(sK2)) [choice axiom]
33. ? [X3] : (sK1 != X3 & k4_relat_1(sK0) = k3_relat_1(sK2,X3) & k4_relat_1(k9_xtuple_0(X3)) = k3_relat_1(sK1,sK2) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(X3)) => (sK1 != sK3 & k4_relat_1(sK0) = k3_relat_1(sK2,sK3) & k3_relat_1(sK1,sK2) = k4_relat_1(k9_xtuple_0(sK3)) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(sK3)) [choice axiom]
34. ((sK1 != sK3 & k4_relat_1(sK0) = k3_relat_1(sK2,sK3) & k3_relat_1(sK1,sK2) = k4_relat_1(k9_xtuple_0(sK3)) & r1_tarski(k10_xtuple_0(sK1),sK0) & v1_relat_1(sK3)) & v1_relat_1(sK2)) & v1_relat_1(sK1) [skolemisation 24,33,32,31]
38. v1_relat_1(sK1) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
39. v1_relat_1(sK2) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
40. v1_relat_1(sK3) [cnf transformation 34] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
41. r1_tarski(k10_xtuple_0(sK1),sK0) [cnf transformation 34] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
42. k3_relat_1(sK1,sK2) = k4_relat_1(k9_xtuple_0(sK3)) [cnf transformation 34] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
43. k4_relat_1(sK0) = k3_relat_1(sK2,sK3) [cnf transformation 34] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
44. sK1 != sK3 [cnf transformation 34] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
48. k3_relat_1(k3_relat_1(X0,X1),X2) = k3_relat_1(X0,k3_relat_1(X1,X2)) | ~v1_relat_1(X2) | ~v1_relat_1(X1) | ~v1_relat_1(X0) [cnf transformation 25] {a:0,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
50. k3_relat_1(X1,k4_relat_1(X0)) = X1 | ~r1_tarski(k10_xtuple_0(X1),X0) | ~v1_relat_1(X1) [cnf transformation 28] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
51. k3_relat_1(k4_relat_1(X0),X1) = X1 | ~r1_tarski(k9_xtuple_0(X1),X0) | ~v1_relat_1(X1) [cnf transformation 30] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
55. r1_tarski(X0,X0) [cnf transformation 22] {a:0,w:3,wCS:1200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
58. 1 <=> sK1 = sK3 [avatar definition]
60. sK1 != sK3 <- (~1) [avatar component clause 58] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
61. ~1 [avatar split clause 44,58]
63. 2 <=> k4_relat_1(sK0) = k3_relat_1(sK2,sK3) [avatar definition]
65. k4_relat_1(sK0) = k3_relat_1(sK2,sK3) <- (2) [avatar component clause 63] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
66. 2 [avatar split clause 43,63]
68. 3 <=> k3_relat_1(sK1,sK2) = k4_relat_1(k9_xtuple_0(sK3)) [avatar definition]
70. k3_relat_1(sK1,sK2) = k4_relat_1(k9_xtuple_0(sK3)) <- (3) [avatar component clause 68] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
71. 3 [avatar split clause 42,68]
73. 4 <=> r1_tarski(k10_xtuple_0(sK1),sK0) [avatar definition]
75. r1_tarski(k10_xtuple_0(sK1),sK0) <- (4) [avatar component clause 73] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
76. 4 [avatar split clause 41,73]
78. 5 <=> v1_relat_1(sK3) [avatar definition]
80. v1_relat_1(sK3) <- (5) [avatar component clause 78] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
81. 5 [avatar split clause 40,78]
83. 6 <=> v1_relat_1(sK2) [avatar definition]
85. v1_relat_1(sK2) <- (6) [avatar component clause 83] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
86. 6 [avatar split clause 39,83]
88. 7 <=> v1_relat_1(sK1) [avatar definition]
90. v1_relat_1(sK1) <- (7) [avatar component clause 88] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
91. 7 [avatar split clause 38,88]
113. k3_relat_1(sK1,k3_relat_1(sK2,X10)) = k3_relat_1(k4_relat_1(k9_xtuple_0(sK3)),X10) | ~v1_relat_1(X10) | ~v1_relat_1(sK2) | ~v1_relat_1(sK1) <- (3) [superposition 48,70] {a:1,w:17,wCS:1700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
124. k3_relat_1(sK1,k3_relat_1(sK2,X10)) = k3_relat_1(k4_relat_1(k9_xtuple_0(sK3)),X10) | ~v1_relat_1(X10) | ~v1_relat_1(sK1) <- (3, 6) [subsumption resolution 113,85] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
125. k3_relat_1(sK1,k3_relat_1(sK2,X10)) = k3_relat_1(k4_relat_1(k9_xtuple_0(sK3)),X10) | ~v1_relat_1(X10) <- (3, 6, 7) [subsumption resolution 124,90] {a:1,w:13,wCS:1300,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
142. k3_relat_1(sK1,k3_relat_1(sK2,X0)) = X0 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(sK3)) | ~v1_relat_1(X0) | ~v1_relat_1(X0) <- (3, 6, 7) [superposition 51,125] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
147. k3_relat_1(sK1,k3_relat_1(sK2,X0)) = X0 | ~r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(sK3)) | ~v1_relat_1(X0) <- (3, 6, 7) [duplicate literal removal 142] {a:2,w:14,wCS:1400,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
584. sK3 = k3_relat_1(sK1,k4_relat_1(sK0)) | ~r1_tarski(k9_xtuple_0(sK3),k9_xtuple_0(sK3)) | ~v1_relat_1(sK3) <- (2, 3, 6, 7) [superposition 147,65] {a:3,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
597. sK3 = k3_relat_1(sK1,k4_relat_1(sK0)) | ~v1_relat_1(sK3) <- (2, 3, 6, 7) [subsumption resolution 584,55] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
598. sK3 = k3_relat_1(sK1,k4_relat_1(sK0)) <- (2, 3, 5, 6, 7) [subsumption resolution 597,80] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
600. 19 <=> sK3 = k3_relat_1(sK1,k4_relat_1(sK0)) [avatar definition]
602. sK3 = k3_relat_1(sK1,k4_relat_1(sK0)) <- (19) [avatar component clause 600] {a:3,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
603. 19 | ~2 | ~3 | ~5 | ~6 | ~7 [avatar split clause 598,88,83,78,68,63,600]
640. sK1 = sK3 | ~r1_tarski(k10_xtuple_0(sK1),sK0) | ~v1_relat_1(sK1) <- (19) [superposition 50,602] {a:4,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
650. ~r1_tarski(k10_xtuple_0(sK1),sK0) | ~v1_relat_1(sK1) <- (~1, 19) [subsumption resolution 640,60] {a:4,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
651. ~v1_relat_1(sK1) <- (~1, 4, 19) [subsumption resolution 650,75] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
652. $false <- (~1, 4, 7, 19) [subsumption resolution 651,90] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:11,thDist:-11}
653. 1 | ~4 | ~7 | ~19 [avatar contradiction clause 652]
660. $false [avatar sat refutation 61,66,71,76,81,86,91,603,653] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:174,thDist:-174}
% SZS output end Proof for relat_1__t56_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 5
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 546
% Active clauses: 40
% Passive clauses: 194
% Final active clauses: 39
% Final passive clauses: 150
% Main loop iterations started: 44
% 
% >>> Simplifying Inferences
% Duplicate literals: 43
% Fw subsumption resolutions: 206
% Fw demodulations: 54
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 33
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 8
% Forward superposition: 96
% Backward superposition: 102
% Self superposition: 10
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 13
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 18
% SAT solver unit clauses: 7
% 
% Memory used [KB]: 9722
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward superposition index maintenance: 0.001 s
% forward subsumption resolution: 0.001 s
% superposition: 0.001 s
% other: 0.005 s ( own 0.002 s ) 

