i: [1,37,0,2,1,1,0,0] t33_orders_1
a: [3,69,0,2,1,37]
s: 69
i: [1,38,0,3,1,1,0,0] t33_orders_1
a: [3,64,0,3,1,38]
s: 64
i: [1,39,0,4,1,1,0,0] t33_orders_1
a: [3,59,0,4,1,39]
s: 59
i: [1,49,0,4,1,0,0,255] fc6_relat_1
s: 49
i: [1,50,0,5,1,0,0,255] idempotence_k3_xboole_0
s: 50
i: [1,40,0,6,3,0,0,1] d3_orders_1
s: 40
i: [1,41,0,6,3,0,0,1] d3_orders_1
s: 41
i: [1,52,0,6,2,0,0,1] dt_k2_wellord1
s: 52
i: [1,53,0,6,2,0,0,3] fc1_relat_1
s: 53
i: [1,51,0,7,1,0,0,255] commutativity_k3_xboole_0
s: 51
i: [1,42,0,8,4,0,0,1] d3_orders_1
s: 42
d: [2,72,1,6,2,0,44,53,51]
s: 72
i: [1,46,0,8,3,0,0,1] d6_orders_1
s: 46
d: [2,77,1,5,2,1,40,46,64]
d: [2,78,1,3,1,2,21,77,69]
a: [3,82,1,3,1,78]
s: 82
i: [1,47,0,8,3,0,0,1] d6_orders_1
s: 47
i: [1,44,0,9,3,0,0,2] l47_orders_1
s: 44
d: [2,84,1,5,2,1,40,47,64]
d: [2,85,1,3,1,2,21,84,69]
a: [3,89,1,3,1,85]
s: 89
i: [1,45,0,9,3,0,0,2] l46_orders_1
s: 45
d: [2,73,1,6,2,0,44,53,51]
s: 73
i: [1,43,0,11,2,0,0,255] d6_wellord1
s: 43
i: [1,48,0,11,4,0,0,1] d6_orders_1
s: 48
d: [2,93,1,13,2,0,44,43,50]
d: [2,102,1,9,1,0,21,93,49]
s: 102
i: [1,54,0,12,5,0,0,2] fc22_relat_2
s: 54
d: [2,96,1,13,2,0,44,50,43]
d: [2,103,1,9,1,0,21,96,49]
s: 103
i: [1,55,0,12,5,0,0,2] fc5_relat_2
s: 55
d: [2,94,1,11,2,0,44,43,51]
s: 94
d: [2,105,2,13,3,0,44,45,102]
d: [2,108,2,9,2,0,21,105,49]
s: 108
d: [2,95,1,11,2,0,44,43,51]
s: 95
d: [2,106,2,13,3,0,44,44,102]
d: [2,109,2,9,2,0,21,106,49]
s: 109
d: [2,98,1,11,2,0,44,51,43]
s: 98
d: [2,99,1,11,2,0,44,51,43]
s: 99
d: [2,129,2,19,3,0,44,94,43]
d: [2,141,2,15,2,0,21,129,49]
d: [2,142,2,11,1,0,21,141,49]
s: 142
d: [2,114,1,18,6,0,44,54,43]
d: [2,115,1,16,5,0,18,114]
d: [2,118,1,12,4,0,21,115,49]
s: 118
d: [2,149,3,10,2,0,44,52,142]
d: [2,155,3,6,1,0,21,149,49]
s: 155
d: [2,122,1,18,6,0,44,55,43]
d: [2,123,1,16,5,0,18,122]
d: [2,125,1,12,4,0,21,123,49]
s: 125
d: [2,91,1,17,5,0,40,44,42]
d: [2,92,1,13,4,0,21,91,52]
s: 92
d: [2,152,3,10,2,0,44,52,142]
d: [2,158,3,6,1,0,21,152,49]
s: 158
d: [2,110,1,20,7,0,40,54,42]
d: [2,117,1,16,6,0,21,110,72]
s: 117
d: [2,137,2,19,3,0,44,43,94]
d: [2,145,2,15,2,0,21,137,49]
d: [2,146,2,11,1,0,21,145,49]
s: 146
d: [2,131,2,18,6,0,44,55,94]
d: [2,140,2,16,5,0,18,131]
d: [2,143,2,12,4,0,21,140,49]
s: 143
d: [2,132,2,18,6,0,44,54,94]
d: [2,139,2,16,5,0,18,132]
d: [2,144,2,12,4,0,21,139,49]
s: 144
d: [2,147,3,15,3,0,44,45,142]
d: [2,153,3,11,2,0,21,147,49]
s: 153
d: [2,162,2,14,5,0,40,92,45]
d: [2,166,2,12,4,0,18,162]
s: 166
d: [2,187,3,0,0,4,52,69,82,89,59,166]
e: 187
d: [2,188,3,8,3,1,40,166,59]
d: [2,195,3,6,2,2,21,188,69]
d: [2,196,3,3,1,3,21,195,89]
d: [2,197,3,0,0,4,21,196,82]
e: 197
f: 197,85,78,37,38,39
% Refutation found. Thanks to Tanya!
% SZS status Theorem for orders_1__t33_orders_1
% SZS output start Proof for orders_1__t33_orders_1
1. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) => v1_orders_1(k2_wellord1(X0,X1)))) [input t33_orders_1]
2. ~! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) => v1_orders_1(k2_wellord1(X0,X1)))) [negated conjecture 1]
4. ! [X0] : (v1_relat_1(X0) => (v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0)))) [input d3_orders_1]
5. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1)))) [input d6_orders_1]
7. ! [X0,X1] : (v1_relat_1(X0) => v1_relat_1(k2_wellord1(X0,X1))) [input dt_k2_wellord1]
15. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r1_relat_2(X0,X1) => v1_relat_2(k2_wellord1(X0,X1)))) [input l46_orders_1]
16. ! [X0] : (v1_relat_1(X0) => ! [X1] : (r8_relat_2(X0,X1) => v8_relat_2(k2_wellord1(X0,X1)))) [input l47_orders_1]
18. ? [X0] : (? [X1] : (~v1_orders_1(k2_wellord1(X0,X1)) & r1_orders_1(X0,X1)) & v1_relat_1(X0)) [ennf transformation 2]
19. ! [X0] : ((v1_orders_1(X0) <=> (v8_relat_2(X0) & v1_relat_2(X0))) | ~v1_relat_1(X0)) [ennf transformation 4]
21. ! [X0] : (! [X1] : (v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 16]
22. ! [X0] : (! [X1] : (v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 15]
23. ! [X0] : (! [X1] : (r1_orders_1(X0,X1) <=> (r8_relat_2(X0,X1) & r1_relat_2(X0,X1))) | ~v1_relat_1(X0)) [ennf transformation 5]
24. ! [X0,X1] : (v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0)) [ennf transformation 7]
30. ? [X0] : (? [X1] : (~v1_orders_1(k2_wellord1(X0,X1)) & r1_orders_1(X0,X1)) & v1_relat_1(X0)) => (? [X1] : (~v1_orders_1(k2_wellord1(sK0,X1)) & r1_orders_1(sK0,X1)) & v1_relat_1(sK0)) [choice axiom]
31. ? [X1] : (~v1_orders_1(k2_wellord1(sK0,X1)) & r1_orders_1(sK0,X1)) => (~v1_orders_1(k2_wellord1(sK0,sK1)) & r1_orders_1(sK0,sK1)) [choice axiom]
32. (~v1_orders_1(k2_wellord1(sK0,sK1)) & r1_orders_1(sK0,sK1)) & v1_relat_1(sK0) [skolemisation 18,31,30]
33. ! [X0] : (((v1_orders_1(X0) | (~v8_relat_2(X0) | ~v1_relat_2(X0))) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [nnf transformation 19]
34. ! [X0] : (((v1_orders_1(X0) | ~v8_relat_2(X0) | ~v1_relat_2(X0)) & ((v8_relat_2(X0) & v1_relat_2(X0)) | ~v1_orders_1(X0))) | ~v1_relat_1(X0)) [flattening 33]
35. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | (~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1))) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [nnf transformation 23]
36. ! [X0] : (! [X1] : ((r1_orders_1(X0,X1) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1)) & ((r8_relat_2(X0,X1) & r1_relat_2(X0,X1)) | ~r1_orders_1(X0,X1))) | ~v1_relat_1(X0)) [flattening 35]
37. v1_relat_1(sK0) [cnf transformation 32] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. r1_orders_1(sK0,sK1) [cnf transformation 32] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. ~v1_orders_1(k2_wellord1(sK0,sK1)) [cnf transformation 32] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
42. ~v8_relat_2(X0) | v1_orders_1(X0) | ~v1_relat_2(X0) | ~v1_relat_1(X0) [cnf transformation 34] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
44. v8_relat_2(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 21] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
45. v1_relat_2(k2_wellord1(X0,X1)) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 22] {a:0,w:9,wCS:3600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
46. ~r1_orders_1(X0,X1) | r1_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
47. ~r1_orders_1(X0,X1) | r8_relat_2(X0,X1) | ~v1_relat_1(X0) [cnf transformation 36] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
52. v1_relat_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) [cnf transformation 24] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
57. 1 <=> v1_orders_1(k2_wellord1(sK0,sK1)) [avatar definition]
59. ~v1_orders_1(k2_wellord1(sK0,sK1)) <- (~1) [avatar component clause 57] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. ~1 [avatar split clause 39,57]
62. 2 <=> r1_orders_1(sK0,sK1) [avatar definition]
64. r1_orders_1(sK0,sK1) <- (2) [avatar component clause 62] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 2 [avatar split clause 38,62]
67. 3 <=> v1_relat_1(sK0) [avatar definition]
69. v1_relat_1(sK0) <- (3) [avatar component clause 67] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
70. 3 [avatar split clause 37,67]
77. r1_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 46,64] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
78. r1_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 77,69] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
80. 4 <=> r1_relat_2(sK0,sK1) [avatar definition]
82. r1_relat_2(sK0,sK1) <- (4) [avatar component clause 80] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
83. 4 | ~2 | ~3 [avatar split clause 78,67,62,80]
84. r8_relat_2(sK0,sK1) | ~v1_relat_1(sK0) <- (2) [resolution 47,64] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
85. r8_relat_2(sK0,sK1) <- (2, 3) [subsumption resolution 84,69] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
87. 5 <=> r8_relat_2(sK0,sK1) [avatar definition]
89. r8_relat_2(sK0,sK1) <- (5) [avatar component clause 87] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
90. 5 | ~2 | ~3 [avatar split clause 85,67,62,87]
91. ~r8_relat_2(X0,X1) | ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(k2_wellord1(X0,X1)) [resolution 44,42] {a:1,w:17,wCS:6800,sine:1,thAx:0,allAx:2,thDist:-2}
92. ~v1_relat_2(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1) [subsumption resolution 91,52] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
162. ~v1_relat_1(X0) | v1_orders_1(k2_wellord1(X0,X1)) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) | ~v1_relat_1(X0) [resolution 92,45] {a:2,w:14,wCS:5600,sine:1,thAx:0,allAx:3,thDist:-3}
166. v1_orders_1(k2_wellord1(X0,X1)) | ~v1_relat_1(X0) | ~r8_relat_2(X0,X1) | ~r1_relat_2(X0,X1) [duplicate literal removal 162] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
188. ~v1_relat_1(sK0) | ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1) [resolution 166,59] {a:3,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
195. ~r8_relat_2(sK0,sK1) | ~r1_relat_2(sK0,sK1) <- (~1, 3) [subsumption resolution 188,69] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
196. ~r1_relat_2(sK0,sK1) <- (~1, 3, 5) [subsumption resolution 195,89] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
197. $false <- (~1, 3, 4, 5) [subsumption resolution 196,82] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
198. 1 | ~3 | ~4 | ~5 [avatar contradiction clause 197]
202. $false [avatar sat refutation 60,65,70,83,90,198] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:49,thDist:-49}
% SZS output end Proof for orders_1__t33_orders_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 15
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 19
% Generated clauses: 143
% Active clauses: 34
% Passive clauses: 55
% Final active clauses: 34
% Final passive clauses: 12
% Main loop iterations started: 43
% 
% >>> Simplifying Inferences
% Duplicate literals: 17
% Fw subsumption resolutions: 34
% 
% >>> Deletion Inferences
% Simple tautologies: 5
% Forward subsumptions: 35
% 
% >>> Generating Inferences
% Binary resolution: 11
% Unit resulting resolution: 1
% Forward superposition: 36
% Backward superposition: 24
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 9
% SAT solver unit clauses: 3
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

