i: [1,19,0,9,3,1,0,0] t48_mcart_1
a: [3,31,0,3,1,19]
s: 31
i: [1,18,0,9,1,1,0,0] t48_mcart_1
a: [3,40,0,9,1,18]
s: 40
i: [1,23,0,6,2,0,0,1] symmetry_r1_xboole_0
s: 23
d: [2,42,1,3,1,1,40,23,31]
a: [3,46,1,3,1,42]
s: 46
i: [1,20,0,10,1,0,0,255] d3_zfmisc_1
s: 20
i: [1,21,0,10,2,0,0,1] t104_zfmisc_1
s: 21
i: [1,22,0,10,2,0,0,1] t104_zfmisc_1
s: 22
d: [2,50,1,10,2,0,40,21,23]
s: 50
d: [2,53,1,10,2,0,40,22,23]
s: 53
d: [2,54,1,11,2,0,44,22,20]
s: 54
d: [2,55,1,11,2,0,44,22,20]
s: 55
d: [2,60,2,11,2,0,44,53,20]
s: 60
d: [2,49,1,13,1,0,44,20,20]
s: 49
d: [2,61,2,11,2,0,44,53,20]
s: 61
d: [2,51,1,13,2,0,44,21,20]
s: 51
d: [2,52,1,13,2,0,44,21,20]
s: 52
d: [2,62,2,11,2,0,40,54,23]
s: 62
d: [2,64,2,11,2,0,40,55,23]
s: 64
d: [2,63,2,12,2,0,44,54,20]
s: 63
a: [3,34,0,3,1,19]
s: 34
d: [2,65,2,12,2,0,44,55,20]
s: 65
d: [2,67,3,12,2,0,44,60,20]
s: 67
d: [2,57,2,13,2,0,44,50,20]
s: 57
d: [2,89,4,3,1,1,40,67,40]
a: [3,96,4,3,1,89]
s: 96
d: [2,58,2,13,2,0,44,50,20]
s: 58
d: [2,76,2,13,2,0,40,51,23]
s: 76
d: [2,74,3,12,2,0,44,61,20]
s: 74
d: [2,80,2,13,2,0,40,52,23]
s: 80
d: [2,85,3,12,2,0,40,63,23]
s: 85
d: [2,78,2,16,2,0,44,51,20]
s: 78
d: [2,81,2,16,2,0,44,52,20]
s: 81
d: [2,106,3,7,1,1,40,78,40]
a: [3,113,3,7,1,106]
s: 113
d: [2,117,4,0,0,2,52,46,113,53]
e: 117
d: [2,118,4,0,0,2,52,31,113,22]
e: 118
d: [2,119,4,3,1,1,40,113,53]
d: [2,125,4,0,0,2,21,119,46]
e: 125
d: [2,121,4,3,1,1,40,113,22]
d: [2,127,4,0,0,2,21,121,31]
e: 127
r: 31
r: 46
d: [2,84,3,3,1,1,40,63,40]
d: [2,122,4,3,1,1,40,113,21]
f: 122,121,106,84,18,19
% Refutation found. Thanks to Tanya!
% SZS status Theorem for mcart_1__t48_mcart_1
% SZS output start Proof for mcart_1__t48_mcart_1
1. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (~r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k3_zfmisc_1(X3,X4,X5)) => (~r1_xboole_0(X2,X5) & ~r1_xboole_0(X1,X4) & ~r1_xboole_0(X0,X3))) [input t48_mcart_1]
2. ~! [X0] : ! [X1] : ! [X2] : ! [X3] : ! [X4] : ! [X5] : (~r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k3_zfmisc_1(X3,X4,X5)) => (~r1_xboole_0(X2,X5) & ~r1_xboole_0(X1,X4) & ~r1_xboole_0(X0,X3))) [negated conjecture 1]
3. ! [X0] : ! [X1] : ! [X2] : k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [input d3_zfmisc_1]
8. ! [X0] : ! [X1] : ! [X2] : ! [X3] : ((r1_xboole_0(X2,X3) | r1_xboole_0(X0,X1)) => r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [input t104_zfmisc_1]
9. ~! [X0,X1,X2,X3,X4,X5] : (~r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k3_zfmisc_1(X3,X4,X5)) => (~r1_xboole_0(X2,X5) & ~r1_xboole_0(X1,X4) & ~r1_xboole_0(X0,X3))) [flattening 2]
10. ! [X0,X1,X2] : k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [flattening 3]
11. ! [X0,X1,X2,X3] : ((r1_xboole_0(X2,X3) | r1_xboole_0(X0,X1)) => r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3))) [flattening 8]
13. ? [X0,X1,X2,X3,X4,X5] : ((r1_xboole_0(X2,X5) | r1_xboole_0(X1,X4) | r1_xboole_0(X0,X3)) & ~r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k3_zfmisc_1(X3,X4,X5))) [ennf transformation 9]
14. ! [X0,X1,X2,X3] : (r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | (~r1_xboole_0(X2,X3) & ~r1_xboole_0(X0,X1))) [ennf transformation 11]
16. ? [X0,X1,X2,X3,X4,X5] : ((r1_xboole_0(X2,X5) | r1_xboole_0(X1,X4) | r1_xboole_0(X0,X3)) & ~r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k3_zfmisc_1(X3,X4,X5))) => ((r1_xboole_0(sK2,sK5) | r1_xboole_0(sK1,sK4) | r1_xboole_0(sK0,sK3)) & ~r1_xboole_0(k3_zfmisc_1(sK0,sK1,sK2),k3_zfmisc_1(sK3,sK4,sK5))) [choice axiom]
17. (r1_xboole_0(sK2,sK5) | r1_xboole_0(sK1,sK4) | r1_xboole_0(sK0,sK3)) & ~r1_xboole_0(k3_zfmisc_1(sK0,sK1,sK2),k3_zfmisc_1(sK3,sK4,sK5)) [skolemisation 13,16]
18. ~r1_xboole_0(k3_zfmisc_1(sK0,sK1,sK2),k3_zfmisc_1(sK3,sK4,sK5)) [cnf transformation 17] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
19. r1_xboole_0(sK2,sK5) | r1_xboole_0(sK1,sK4) | r1_xboole_0(sK0,sK3) [cnf transformation 17] {a:0,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
20. k3_zfmisc_1(X0,X1,X2) = k2_zfmisc_1(k2_zfmisc_1(X0,X1),X2) [cnf transformation 10] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
21. r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_xboole_0(X0,X1) [cnf transformation 14] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
22. r1_xboole_0(k2_zfmisc_1(X0,X2),k2_zfmisc_1(X1,X3)) | ~r1_xboole_0(X2,X3) [cnf transformation 14] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
25. 1 <=> r1_xboole_0(sK0,sK3) [avatar definition]
29. 2 <=> r1_xboole_0(sK1,sK4) [avatar definition]
33. 3 <=> r1_xboole_0(sK2,sK5) [avatar definition]
36. 1 | 2 | 3 [avatar split clause 19,33,29,25]
38. 4 <=> r1_xboole_0(k3_zfmisc_1(sK0,sK1,sK2),k3_zfmisc_1(sK3,sK4,sK5)) [avatar definition]
40. ~r1_xboole_0(k3_zfmisc_1(sK0,sK1,sK2),k3_zfmisc_1(sK3,sK4,sK5)) <- (~4) [avatar component clause 38] {a:0,w:9,wCS:900,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
41. ~4 [avatar split clause 18,38]
51. r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k2_zfmisc_1(X3,X4)) | ~r1_xboole_0(k2_zfmisc_1(X0,X1),X3) [superposition 21,20] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. r1_xboole_0(k3_zfmisc_1(X0,X1,X2),k2_zfmisc_1(X3,X4)) | ~r1_xboole_0(X2,X4) [superposition 22,20] {a:1,w:11,wCS:4400,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
63. r1_xboole_0(k3_zfmisc_1(X3,X4,X5),k3_zfmisc_1(X0,X1,X2)) | ~r1_xboole_0(X5,X2) [superposition 54,20] {a:2,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
78. r1_xboole_0(k3_zfmisc_1(X3,X4,X5),k3_zfmisc_1(X0,X1,X2)) | ~r1_xboole_0(k2_zfmisc_1(X3,X4),k2_zfmisc_1(X0,X1)) [superposition 51,20] {a:2,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
84. ~r1_xboole_0(sK2,sK5) <- (~4) [resolution 63,40] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
88. ~3 | 4 [avatar split clause 84,38,33]
106. ~r1_xboole_0(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK3,sK4)) <- (~4) [resolution 78,40] {a:3,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
111. 7 <=> r1_xboole_0(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK3,sK4)) [avatar definition]
113. ~r1_xboole_0(k2_zfmisc_1(sK0,sK1),k2_zfmisc_1(sK3,sK4)) <- (~7) [avatar component clause 111] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
114. ~7 | 4 [avatar split clause 106,38,111]
121. ~r1_xboole_0(sK1,sK4) <- (~7) [resolution 113,22] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
122. ~r1_xboole_0(sK0,sK3) <- (~7) [resolution 113,21] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
129. ~2 | 7 [avatar split clause 121,111,29]
131. ~1 | 7 [avatar split clause 122,111,25]
137. $false [avatar sat refutation 36,41,88,114,129,131] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:100,thDist:-100}
% SZS output end Proof for mcart_1__t48_mcart_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% >>> Preprocessing
% Introduced skolems: 6
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 6
% Generated clauses: 86
% Active clauses: 24
% Passive clauses: 38
% Final active clauses: 22
% Final passive clauses: 6
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 2
% Fw demodulations: 7
% 
% >>> Deletion Inferences
% Forward subsumptions: 31
% 
% >>> Generating Inferences
% Binary resolution: 24
% Unit resulting resolution: 2
% Forward superposition: 28
% Backward superposition: 3
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1
% Split components: 3
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 1
% SAT solver binary clauses: 12
% 
% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% binary resolution index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

