i: [1,39,0,3,1,1,0,0] t19_arytm_1
a: [3,61,0,3,1,39]
s: 61
i: [1,38,0,3,1,1,0,0] t19_arytm_1
a: [3,66,0,3,1,38]
s: 66
i: [1,37,0,3,1,1,0,0] t19_arytm_1
a: [3,71,0,3,1,37]
s: 71
i: [1,36,0,3,1,1,0,0] t19_arytm_1
a: [3,76,0,3,1,36]
s: 76
i: [1,40,0,7,1,1,0,0] t19_arytm_1
a: [3,56,0,7,1,40]
s: 56
d: [2,103,0,7,1,2,22,56,66]
a: [3,107,0,7,1,103]
s: 107
i: [1,41,0,3,1,0,0,255] redefinition_k11_arytm_3
a: [3,81,0,3,1,41]
s: 81
i: [1,42,0,3,1,0,0,255] dt_k11_arytm_3
a: [3,86,0,3,1,42]
s: 86
i: [1,51,0,9,3,0,0,2] t6_arytm_1
a: [3,90,0,3,1,51]
s: 90
i: [1,48,0,4,1,0,0,1] existence_m1_subset_1
s: 48
a: [3,94,0,6,2,51]
s: 94
d: [2,109,1,3,1,2,40,94,71]
a: [3,115,1,3,1,109]
s: 115
i: [1,52,0,11,3,0,0,2] l17_arytm_1
a: [3,98,0,8,2,52]
s: 98
d: [2,110,1,3,1,2,40,94,90]
a: [3,120,1,3,1,110]
s: 120
i: [1,49,0,11,3,0,0,2] dt_k1_arytm_1
s: 49
i: [1,50,0,12,4,0,0,2] connectedness_r1_arytm_2
s: 50
d: [2,132,1,9,3,1,40,50,71]
s: 132
i: [1,47,0,15,5,0,0,2] t4_arytm_1
s: 47
d: [2,136,2,6,2,1,40,132,71]
d: [2,138,2,3,1,1,18,136]
a: [3,142,2,3,1,138]
s: 142
i: [1,43,0,16,4,0,0,1] d2_arytm_1
s: 43
i: [1,44,0,18,4,0,0,1] d2_arytm_1
s: 44
d: [2,153,2,12,4,1,40,47,115]
d: [2,156,2,9,3,2,21,153,61]
d: [2,157,2,6,2,3,21,156,90]
d: [2,158,2,3,1,4,21,157,71]
a: [3,162,2,3,1,158]
s: 162
d: [2,111,1,4,1,1,40,94,48]
a: [3,125,1,4,1,111]
s: 125
d: [2,137,2,8,2,1,40,132,48]
a: [3,151,2,4,1,137]
s: 151
d: [2,166,1,17,5,1,44,43,98]
d: [2,171,1,14,4,1,18,166]
d: [2,172,1,11,3,2,21,171,90]
d: [2,173,1,8,2,3,21,172,94]
s: 173
d: [2,167,1,17,5,1,44,98,43]
d: [2,170,1,14,4,1,18,167]
d: [2,174,1,11,3,2,21,170,90]
d: [2,175,1,8,2,3,21,174,94]
s: 175
a: [3,146,2,4,1,137]
s: 146
d: [2,127,1,11,3,1,40,49,94]
s: 127
d: [2,135,2,16,4,1,40,132,49]
s: 135
d: [2,133,1,11,3,0,40,50,48]
s: 133
d: [2,176,1,19,5,1,44,44,98]
d: [2,179,1,16,4,1,18,176]
d: [2,180,1,13,3,2,21,179,90]
s: 180
d: [2,224,3,0,0,5,52,71,162,107,180]
e: 224
d: [2,226,2,13,3,3,44,107,180]
d: [2,228,2,6,2,3,19,226]
d: [2,231,2,3,1,4,21,228,162]
d: [2,232,2,0,0,5,21,231,71]
e: 232
d: [2,101,0,3,1,2,22,76,66]
f: 232,158,109,103,101,52,51,36,37,38,39,40
% Refutation found. Thanks to Tanya!
% SZS status Theorem for arytm_1__t19_arytm_1
% SZS output start Proof for arytm_1__t19_arytm_1
1. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => (k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,X1) | k11_arytm_3 = X1)))) [input t19_arytm_1]
2. ~! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => (k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,X1) | k11_arytm_3 = X1)))) [negated conjecture 1]
4. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((~r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0))) & (r1_arytm_2(X1,X0) => k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1))))) [input d2_arytm_1]
14. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => k1_arytm_1(X1,X0) = X1))) [input l17_arytm_1]
16. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => ((r1_arytm_2(X1,X0) & r1_arytm_2(X0,X1)) => X0 = X1))) [input t4_arytm_1]
17. ! [X0] : (m1_subset_1(X0,k2_arytm_2) => ! [X1] : (m1_subset_1(X1,k2_arytm_2) => (k11_arytm_3 = X0 => r1_arytm_2(X0,X1)))) [input t6_arytm_1]
18. ? [X0] : (? [X1] : (((k2_arytm_1(X0,X1) != k4_tarski(k11_arytm_3,X1) & k11_arytm_3 != X1) & k11_arytm_3 = X0) & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [ennf transformation 2]
19. ? [X0] : (? [X1] : (k2_arytm_1(X0,X1) != k4_tarski(k11_arytm_3,X1) & k11_arytm_3 != X1 & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) [flattening 18]
20. ! [X0] : (! [X1] : (((k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0)) | r1_arytm_2(X1,X0)) & (k2_arytm_1(X0,X1) = k1_arytm_1(X0,X1) | ~r1_arytm_2(X1,X0))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 4]
21. ! [X0] : (! [X1] : ((r1_arytm_2(X0,X1) | k11_arytm_3 != X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 17]
22. ! [X0] : (! [X1] : (r1_arytm_2(X0,X1) | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 21]
23. ! [X0] : (! [X1] : ((k1_arytm_1(X1,X0) = X1 | k11_arytm_3 != X0) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 14]
24. ! [X0] : (! [X1] : (k1_arytm_1(X1,X0) = X1 | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 23]
25. ! [X0] : (! [X1] : ((X0 = X1 | (~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1))) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [ennf transformation 16]
26. ! [X0] : (! [X1] : (X0 = X1 | ~r1_arytm_2(X1,X0) | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2)) | ~m1_subset_1(X0,k2_arytm_2)) [flattening 25]
31. ? [X0] : (? [X1] : (k2_arytm_1(X0,X1) != k4_tarski(k11_arytm_3,X1) & k11_arytm_3 != X1 & k11_arytm_3 = X0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(X0,k2_arytm_2)) => (? [X1] : (k4_tarski(k11_arytm_3,X1) != k2_arytm_1(sK0,X1) & k11_arytm_3 != X1 & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2)) [choice axiom]
32. ? [X1] : (k4_tarski(k11_arytm_3,X1) != k2_arytm_1(sK0,X1) & k11_arytm_3 != X1 & k11_arytm_3 = sK0 & m1_subset_1(X1,k2_arytm_2)) => (k2_arytm_1(sK0,sK1) != k4_tarski(k11_arytm_3,sK1) & k11_arytm_3 != sK1 & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) [choice axiom]
33. (k2_arytm_1(sK0,sK1) != k4_tarski(k11_arytm_3,sK1) & k11_arytm_3 != sK1 & k11_arytm_3 = sK0 & m1_subset_1(sK1,k2_arytm_2)) & m1_subset_1(sK0,k2_arytm_2) [skolemisation 19,32,31]
36. m1_subset_1(sK0,k2_arytm_2) [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
37. m1_subset_1(sK1,k2_arytm_2) [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
38. k11_arytm_3 = sK0 [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
39. k11_arytm_3 != sK1 [cnf transformation 33] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
40. k2_arytm_1(sK0,sK1) != k4_tarski(k11_arytm_3,sK1) [cnf transformation 33] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
44. k2_arytm_1(X0,X1) = k4_tarski(k11_arytm_3,k1_arytm_1(X1,X0)) | r1_arytm_2(X1,X0) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 20] {a:0,w:18,wCS:7200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
45. r1_arytm_2(X0,X1) | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 22] {a:0,w:12,wCS:4800,sine:2,thAx:0,allAx:1,thDist:-1}
46. k1_arytm_1(X1,X0) = X1 | k11_arytm_3 != X0 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 24] {a:0,w:14,wCS:5600,sine:2,thAx:0,allAx:1,thDist:-1}
47. ~r1_arytm_2(X1,X0) | X0 = X1 | ~r1_arytm_2(X0,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) [cnf transformation 26] {a:0,w:15,wCS:6000,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
51. r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 45] {a:0,w:9,wCS:3600,sine:2,thAx:0,allAx:1,thDist:-1}
52. k1_arytm_1(X1,k11_arytm_3) = X1 | ~m1_subset_1(X1,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) [equality resolution 46] {a:0,w:11,wCS:4400,sine:2,thAx:0,allAx:1,thDist:-1}
54. 1 <=> k2_arytm_1(sK0,sK1) = k4_tarski(k11_arytm_3,sK1) [avatar definition]
56. k2_arytm_1(sK0,sK1) != k4_tarski(k11_arytm_3,sK1) <- (~1) [avatar component clause 54] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
57. ~1 [avatar split clause 40,54]
59. 2 <=> k11_arytm_3 = sK1 [avatar definition]
61. k11_arytm_3 != sK1 <- (~2) [avatar component clause 59] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
62. ~2 [avatar split clause 39,59]
64. 3 <=> k11_arytm_3 = sK0 [avatar definition]
66. k11_arytm_3 = sK0 <- (3) [avatar component clause 64] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
67. 3 [avatar split clause 38,64]
69. 4 <=> m1_subset_1(sK1,k2_arytm_2) [avatar definition]
71. m1_subset_1(sK1,k2_arytm_2) <- (4) [avatar component clause 69] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
72. 4 [avatar split clause 37,69]
74. 5 <=> m1_subset_1(sK0,k2_arytm_2) [avatar definition]
76. m1_subset_1(sK0,k2_arytm_2) <- (5) [avatar component clause 74] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
77. 5 [avatar split clause 36,74]
89. 8 <=> m1_subset_1(k11_arytm_3,k2_arytm_2) [avatar definition]
90. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (8) [avatar component clause 89] {a:0,w:3,wCS:1200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
93. 9 <=> ! [X1] : (r1_arytm_2(k11_arytm_3,X1) | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
94. ~m1_subset_1(X1,k2_arytm_2) | r1_arytm_2(k11_arytm_3,X1) <- (9) [avatar component clause 93] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
95. ~8 | 9 [avatar split clause 51,93,89]
97. 10 <=> ! [X1] : (k1_arytm_1(X1,k11_arytm_3) = X1 | ~m1_subset_1(X1,k2_arytm_2)) [avatar definition]
98. k1_arytm_1(X1,k11_arytm_3) = X1 | ~m1_subset_1(X1,k2_arytm_2) <- (10) [avatar component clause 97] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
99. ~8 | 10 [avatar split clause 52,97,89]
101. m1_subset_1(k11_arytm_3,k2_arytm_2) <- (3, 5) [forward demodulation 76,66] {a:0,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
102. 8 | ~3 | ~5 [avatar split clause 101,74,64,89]
103. k4_tarski(k11_arytm_3,sK1) != k2_arytm_1(k11_arytm_3,sK1) <- (~1, 3) [forward demodulation 56,66] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
105. 11 <=> k4_tarski(k11_arytm_3,sK1) = k2_arytm_1(k11_arytm_3,sK1) [avatar definition]
107. k4_tarski(k11_arytm_3,sK1) != k2_arytm_1(k11_arytm_3,sK1) <- (~11) [avatar component clause 105] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
108. ~11 | 1 | ~3 [avatar split clause 103,64,54,105]
109. r1_arytm_2(k11_arytm_3,sK1) <- (4, 9) [resolution 94,71] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
113. 12 <=> r1_arytm_2(k11_arytm_3,sK1) [avatar definition]
115. r1_arytm_2(k11_arytm_3,sK1) <- (12) [avatar component clause 113] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
116. 12 | ~4 | ~9 [avatar split clause 109,93,69,113]
153. k11_arytm_3 = sK1 | ~r1_arytm_2(sK1,k11_arytm_3) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (12) [resolution 47,115] {a:2,w:12,wCS:1200,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
156. ~r1_arytm_2(sK1,k11_arytm_3) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(sK1,k2_arytm_2) <- (~2, 12) [subsumption resolution 153,61] {a:2,w:9,wCS:900,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
157. ~r1_arytm_2(sK1,k11_arytm_3) | ~m1_subset_1(sK1,k2_arytm_2) <- (~2, 8, 12) [subsumption resolution 156,90] {a:2,w:6,wCS:600,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
158. ~r1_arytm_2(sK1,k11_arytm_3) <- (~2, 4, 8, 12) [subsumption resolution 157,71] {a:2,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:5,thDist:-5}
160. 18 <=> r1_arytm_2(sK1,k11_arytm_3) [avatar definition]
162. ~r1_arytm_2(sK1,k11_arytm_3) <- (~18) [avatar component clause 160] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
163. ~18 | 2 | ~4 | ~8 | ~12 [avatar split clause 158,113,89,69,59,160]
176. k2_arytm_1(k11_arytm_3,X0) = k4_tarski(k11_arytm_3,X0) | r1_arytm_2(X0,k11_arytm_3) | ~m1_subset_1(X0,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) | ~m1_subset_1(X0,k2_arytm_2) <- (10) [superposition 44,98] {a:1,w:19,wCS:7600,sine:1,thAx:0,allAx:2,thDist:-2}
179. k2_arytm_1(k11_arytm_3,X0) = k4_tarski(k11_arytm_3,X0) | r1_arytm_2(X0,k11_arytm_3) | ~m1_subset_1(X0,k2_arytm_2) | ~m1_subset_1(k11_arytm_3,k2_arytm_2) <- (10) [duplicate literal removal 176] {a:1,w:16,wCS:6400,sine:1,thAx:0,allAx:2,thDist:-2}
180. k2_arytm_1(k11_arytm_3,X0) = k4_tarski(k11_arytm_3,X0) | r1_arytm_2(X0,k11_arytm_3) | ~m1_subset_1(X0,k2_arytm_2) <- (8, 10) [subsumption resolution 179,90] {a:1,w:13,wCS:5200,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
226. k2_arytm_1(k11_arytm_3,sK1) != k2_arytm_1(k11_arytm_3,sK1) | r1_arytm_2(sK1,k11_arytm_3) | ~m1_subset_1(sK1,k2_arytm_2) <- (8, 10, ~11) [superposition 107,180] {a:2,w:13,wCS:1300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
228. r1_arytm_2(sK1,k11_arytm_3) | ~m1_subset_1(sK1,k2_arytm_2) <- (8, 10, ~11) [trivial inequality removal 226] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
231. ~m1_subset_1(sK1,k2_arytm_2) <- (8, 10, ~11, ~18) [subsumption resolution 228,162] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
232. $false <- (4, 8, 10, ~11, ~18) [subsumption resolution 231,71] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
233. ~4 | ~8 | ~10 | 11 | 18 [avatar contradiction clause 232]
235. $false [avatar sat refutation 57,62,67,72,77,95,99,102,108,116,163,233] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:129,thDist:-129}
% SZS output end Proof for arytm_1__t19_arytm_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 16
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 15
% Generated clauses: 108
% Active clauses: 28
% Passive clauses: 39
% Final active clauses: 28
% Final passive clauses: 8
% Main loop iterations started: 31
% 
% >>> Simplifying Inferences
% Duplicate literals: 18
% Trivial inequalities: 1
% Fw subsumption resolutions: 20
% Fw demodulations: 2
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Equational tautologies: 2
% Forward subsumptions: 6
% 
% >>> Generating Inferences
% Binary resolution: 21
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 9
% Backward superposition: 4
% 
% >>> AVATAR
% Split clauses: 5
% Split components: 10
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 25
% SAT solver unit clauses: 8
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward superposition index maintenance: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

