i: [1,59,0,2,1,1,0,0] t131_zfmisc_1
a: [3,93,0,2,1,59]
s: 93
i: [1,58,0,2,1,1,0,0] t131_zfmisc_1
a: [3,98,0,2,1,58]
s: 98
i: [1,60,0,4,1,1,0,0] t131_zfmisc_1
s: 60
i: [1,61,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,103,0,2,1,61]
s: 103
i: [1,62,0,2,1,0,0,255] fc1_xboole_0
a: [3,108,0,2,1,62]
s: 108
i: [1,82,0,2,1,0,0,1] rc2_xboole_0
a: [3,118,0,2,1,82]
s: 118
i: [1,83,0,2,1,0,0,1] rc2_zfmisc_1
a: [3,123,0,2,1,83]
s: 123
i: [1,84,0,2,1,0,0,1] rc1_xboole_0
a: [3,128,0,2,1,84]
s: 128
i: [1,86,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,133,0,2,1,86]
s: 133
i: [1,85,0,2,1,0,0,1] rc1_zfmisc_1
a: [3,138,0,2,1,85]
s: 138
i: [1,64,0,3,1,0,0,1] fc2_xboole_0
s: 64
i: [1,65,0,3,1,0,0,1] fc2_zfmisc_1
s: 65
i: [1,63,0,3,1,0,0,255] d2_xboole_0
a: [3,113,0,3,1,63]
s: 113
r: 103
i: [1,66,0,4,2,0,0,1] cc2_zfmisc_1
s: 66
i: [1,67,0,4,2,0,0,1] cc1_zfmisc_1
s: 67
i: [1,88,0,4,1,0,0,1] d1_tarski
s: 88
d: [2,141,1,2,1,1,40,66,128]
a: [3,146,1,2,1,141]
s: 146
i: [1,68,0,5,2,0,0,255] t6_boole
s: 68
d: [2,142,1,2,1,1,40,66,108]
a: [3,151,1,2,1,142]
s: 151
i: [1,69,0,5,2,0,0,2] d1_xboole_0
s: 69
i: [1,80,0,5,2,0,0,2] t7_boole
s: 80
i: [1,70,0,6,2,0,0,2] d1_xboole_0
s: 70
i: [1,72,0,6,2,0,0,2] d10_zfmisc_1
s: 72
i: [1,73,0,6,2,0,0,2] d10_zfmisc_1
s: 73
i: [1,81,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 81
i: [1,74,0,7,2,0,0,2] d10_zfmisc_1
s: 74
d: [2,160,1,4,1,0,40,81,88]
s: 160
i: [1,79,0,7,3,0,0,1] t8_boole
s: 79
d: [2,161,1,6,2,0,40,81,70]
s: 161
i: [1,89,0,7,2,0,0,1] d1_tarski
s: 89
i: [1,71,0,11,4,0,0,2] d10_zfmisc_1
s: 71
d: [2,165,1,8,2,0,40,89,70]
d: [2,168,1,5,1,0,21,165,64]
s: 168
i: [1,77,0,14,3,0,0,1] d1_tarski
s: 77
d: [2,162,1,6,2,0,40,81,72]
s: 162
i: [1,78,0,14,3,0,0,1] d1_tarski
s: 78
d: [2,163,1,6,2,0,40,81,73]
s: 163
d: [2,153,1,7,3,0,44,68,68]
s: 153
d: [2,170,1,11,4,0,40,71,70]
d: [2,173,1,9,3,0,21,170,69]
s: 173
d: [2,178,1,11,3,0,40,77,69]
s: 178
d: [2,177,1,14,3,0,40,77,81]
s: 177
d: [2,179,1,17,3,0,40,77,89]
s: 179
d: [2,196,2,14,3,0,46,179]
s: 196
d: [2,176,1,19,5,0,40,77,71]
s: 176
d: [2,194,2,23,5,0,44,78,179]
d: [2,198,2,18,4,0,18,194]
d: [2,201,2,14,3,0,21,198,88]
s: 201
d: [2,211,2,17,5,0,46,176]
s: 211
d: [2,208,2,24,7,0,44,78,176]
d: [2,214,2,17,5,0,18,208]
s: 214
d: [2,219,3,14,4,0,47,211]
s: 219
d: [2,224,4,19,6,0,44,78,219]
d: [2,227,4,12,4,0,18,224]
d: [2,228,4,9,3,0,19,227]
s: 228
r: 219
d: [2,232,5,8,3,0,44,60,228]
s: 232
d: [2,231,5,13,5,0,44,228,228]
s: 231
d: [2,240,6,5,2,0,47,232]
d: [2,241,6,3,1,1,21,240,93]
s: 241
d: [2,247,7,0,0,2,52,98,241,70]
e: 247
d: [2,248,7,2,1,1,40,241,70]
d: [2,253,7,0,0,2,21,248,98]
e: 253
f: 253,58,59
% Refutation found. Thanks to Tanya!
% SZS status Theorem for zfmisc_1__t131_zfmisc_1
% SZS output start Proof for zfmisc_1__t131_zfmisc_1
1. ! [X0] : ~(! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [input t131_zfmisc_1]
2. ~! [X0] : ~(! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [negated conjecture 1]
6. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1] : ! [X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => X1 = X2)) [input d10_zfmisc_1]
7. ! [X0] : ! [X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [input d1_tarski]
8. ! [X0] : (v1_xboole_0(X0) <=> ! [X1] : ~r2_hidden(X1,X0)) [input d1_xboole_0]
23. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : ((r2_hidden(X2,X0) & r2_hidden(X1,X0)) => X1 = X2)) [flattening 6]
24. ! [X0,X1] : (k1_tarski(X0) = X1 <=> ! [X2] : (r2_hidden(X2,X1) <=> X0 = X2)) [flattening 7]
27. ? [X0] : (! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) [ennf transformation 2]
31. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : (X1 = X2 | (~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)))) [ennf transformation 23]
32. ! [X0] : (v1_zfmisc_1(X0) <=> ! [X1,X2] : (X1 = X2 | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0))) [flattening 31]
36. ? [X0] : (! [X1] : k1_tarski(X1) != X0 & v1_zfmisc_1(X0) & ~v1_xboole_0(X0)) => (! [X1] : k1_tarski(X1) != sK0 & v1_zfmisc_1(sK0) & ~v1_xboole_0(sK0)) [choice axiom]
37. ! [X1] : k1_tarski(X1) != sK0 & v1_zfmisc_1(sK0) & ~v1_xboole_0(sK0) [skolemisation 27,36]
38. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X1] : ~r2_hidden(X1,X0) | ~v1_xboole_0(X0))) [nnf transformation 8]
39. ! [X0] : ((v1_xboole_0(X0) | ? [X1] : r2_hidden(X1,X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [rectify 38]
40. ! [X0] : (? [X1] : r2_hidden(X1,X0) => r2_hidden(sK1(X0),X0)) [choice axiom]
41. ! [X0] : ((v1_xboole_0(X0) | r2_hidden(sK1(X0),X0)) & (! [X2] : ~r2_hidden(X2,X0) | ~v1_xboole_0(X0))) [skolemisation 39,40]
42. ! [X0] : ((v1_zfmisc_1(X0) | ? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X1,X2] : (X1 = X2 | ~r2_hidden(X2,X0) | ~r2_hidden(X1,X0)) | ~v1_zfmisc_1(X0))) [nnf transformation 32]
43. ! [X0] : ((v1_zfmisc_1(X0) | ? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0))) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_zfmisc_1(X0))) [rectify 42]
44. ! [X0] : (? [X1,X2] : (X1 != X2 & r2_hidden(X2,X0) & r2_hidden(X1,X0)) => (sK2(X0) != sK3(X0) & r2_hidden(sK3(X0),X0) & r2_hidden(sK2(X0),X0))) [choice axiom]
45. ! [X0] : ((v1_zfmisc_1(X0) | (sK2(X0) != sK3(X0) & r2_hidden(sK3(X0),X0) & r2_hidden(sK2(X0),X0))) & (! [X3,X4] : (X3 = X4 | ~r2_hidden(X4,X0) | ~r2_hidden(X3,X0)) | ~v1_zfmisc_1(X0))) [skolemisation 43,44]
46. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X2] : ((r2_hidden(X2,X1) | X0 != X2) & (X0 = X2 | ~r2_hidden(X2,X1))) | k1_tarski(X0) != X1)) [nnf transformation 24]
47. ! [X0,X1] : ((k1_tarski(X0) = X1 | ? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [rectify 46]
48. ! [X1,X0] : (? [X2] : ((X0 != X2 | ~r2_hidden(X2,X1)) & (X0 = X2 | r2_hidden(X2,X1))) => ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) [choice axiom]
49. ! [X0,X1] : ((k1_tarski(X0) = X1 | ((sK4(X0,X1) != X0 | ~r2_hidden(sK4(X0,X1),X1)) & (sK4(X0,X1) = X0 | r2_hidden(sK4(X0,X1),X1)))) & (! [X3] : ((r2_hidden(X3,X1) | X0 != X3) & (X0 = X3 | ~r2_hidden(X3,X1))) | k1_tarski(X0) != X1)) [skolemisation 47,48]
58. ~v1_xboole_0(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
59. v1_zfmisc_1(sK0) [cnf transformation 37] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
60. k1_tarski(X1) != sK0 [cnf transformation 37] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
70. r2_hidden(sK1(X0),X0) | v1_xboole_0(X0) [cnf transformation 41] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
71. ~r2_hidden(X4,X0) | X3 = X4 | ~r2_hidden(X3,X0) | ~v1_zfmisc_1(X0) [cnf transformation 45] {a:0,w:11,wCS:4400,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
77. r2_hidden(sK4(X0,X1),X1) | sK4(X0,X1) = X0 | k1_tarski(X0) = X1 [cnf transformation 49] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
78. sK4(X0,X1) != X0 | k1_tarski(X0) = X1 | ~r2_hidden(sK4(X0,X1),X1) [cnf transformation 49] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
91. 1 <=> v1_zfmisc_1(sK0) [avatar definition]
93. v1_zfmisc_1(sK0) <- (1) [avatar component clause 91] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
94. 1 [avatar split clause 59,91]
96. 2 <=> v1_xboole_0(sK0) [avatar definition]
98. ~v1_xboole_0(sK0) <- (~2) [avatar component clause 96] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
99. ~2 [avatar split clause 58,96]
176. sK4(X0,X1) = X2 | k1_tarski(X0) = X1 | sK4(X0,X1) = X0 | ~r2_hidden(X2,X1) | ~v1_zfmisc_1(X1) [resolution 77,71] {a:1,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
211. X0 != X2 | k1_tarski(X0) = X1 | sK4(X0,X1) = X0 | ~r2_hidden(X2,X1) | ~v1_zfmisc_1(X1) [equality factoring 176] {a:2,w:17,wCS:6800,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
219. sK4(X0,X1) = X0 | k1_tarski(X0) = X1 | ~r2_hidden(X0,X1) | ~v1_zfmisc_1(X1) [equality resolution 211] {a:3,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
224. X7 != X7 | k1_tarski(X7) = X8 | ~r2_hidden(X7,X8) | k1_tarski(X7) = X8 | ~r2_hidden(X7,X8) | ~v1_zfmisc_1(X8) [superposition 78,219] {a:4,w:19,wCS:7600,sine:1,thAx:0,allAx:6,thDist:-6}
227. X7 != X7 | k1_tarski(X7) = X8 | ~r2_hidden(X7,X8) | ~v1_zfmisc_1(X8) [duplicate literal removal 224] {a:4,w:12,wCS:4800,sine:1,thAx:0,allAx:6,thDist:-6}
228. k1_tarski(X7) = X8 | ~r2_hidden(X7,X8) | ~v1_zfmisc_1(X8) [trivial inequality removal 227] {a:4,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:6,thDist:-6}
232. sK0 != X1 | ~r2_hidden(X0,X1) | ~v1_zfmisc_1(X1) [superposition 60,228] {a:5,w:8,wCS:800,nSel:1,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
240. ~r2_hidden(X0,sK0) | ~v1_zfmisc_1(sK0) [equality resolution 232] {a:6,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
241. ~r2_hidden(X0,sK0) <- (1) [subsumption resolution 240,93] {a:6,w:3,wCS:300,nSel:1,goal:1,sine:0,thAx:0,allAx:8,thDist:-8}
248. v1_xboole_0(sK0) <- (1) [resolution 241,70] {a:7,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
253. $false <- (1, ~2) [subsumption resolution 248,98] {a:7,w:0,goal:1,sine:0,thAx:0,allAx:10,thDist:-10}
254. ~1 | 2 [avatar contradiction clause 253]
256. $false [avatar sat refutation 94,99,254] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:58,thDist:-58}
% SZS output end Proof for zfmisc_1__t131_zfmisc_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 21
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 29
% Generated clauses: 137
% Active clauses: 46
% Passive clauses: 53
% Final active clauses: 44
% Final passive clauses: 2
% Main loop iterations started: 51
% 
% >>> Simplifying Inferences
% Duplicate literals: 20
% Trivial inequalities: 2
% Fw subsumption resolutions: 6
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 10
% Equational tautologies: 2
% Forward subsumptions: 38
% Backward subsumptions: 1
% 
% >>> Generating Inferences
% Binary resolution: 31
% Unit resulting resolution: 1
% Forward superposition: 5
% Backward superposition: 25
% Self superposition: 4
% Equality factoring: 2
% Equality resolution: 3
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 16
% SAT solver unit clauses: 10
% SAT solver binary clauses: 4
% 
% Memory used [KB]: 9466
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% Time measurement results:
% backward subsumption index maintenance: 0.001 s
% other: 0.003 s ( own 0.002 s ) 

