i: [1,207,0,2,1,1,0,0] t33_ordinal2
a: [3,328,0,2,1,207]
s: 328
i: [1,206,0,2,1,1,0,0] t33_ordinal2
a: [3,333,0,2,1,206]
s: 333
i: [1,205,0,2,1,1,0,0] t33_ordinal2
a: [3,338,0,2,1,205]
s: 338
i: [1,208,0,3,1,1,0,0] t33_ordinal2
a: [3,323,0,3,1,208]
s: 323
i: [1,209,0,7,1,1,0,0] t33_ordinal2
a: [3,318,0,7,1,209]
s: 318
i: [1,210,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,343,0,2,1,210]
s: 343
i: [1,211,0,2,1,0,0,255] fc1_xboole_0
a: [3,348,0,2,1,211]
s: 348
i: [1,289,0,2,1,0,0,2] rc2_xboole_0
a: [3,365,0,2,1,289]
s: 365
i: [1,290,0,2,1,0,0,255] rc4_ordinal1
a: [3,370,0,2,1,290]
s: 370
i: [1,291,0,2,1,0,0,1] rc1_ordinal1
a: [3,375,0,2,1,291]
s: 375
i: [1,292,0,2,1,0,0,2] rc1_xboole_0
a: [3,380,0,2,1,292]
s: 380
i: [1,296,0,2,1,0,0,255] rc2_ordinal1
a: [3,385,0,2,1,296]
s: 385
i: [1,295,0,2,1,0,0,255] rc2_ordinal1
a: [3,390,0,2,1,295]
s: 390
i: [1,294,0,2,1,0,0,255] rc2_ordinal1
a: [3,395,0,2,1,294]
s: 395
i: [1,293,0,2,1,0,0,255] rc2_ordinal1
a: [3,400,0,2,1,293]
s: 400
i: [1,298,0,2,1,0,0,255] rc5_ordinal1
a: [3,405,0,2,1,298]
s: 405
i: [1,297,0,2,1,0,0,255] rc5_ordinal1
a: [3,410,0,2,1,297]
s: 410
i: [1,300,0,2,1,0,0,255] rc7_funct_1
a: [3,415,0,2,1,300]
s: 415
i: [1,299,0,2,1,0,0,255] rc7_funct_1
a: [3,420,0,2,1,299]
s: 420
i: [1,302,0,2,1,0,0,255] rc1_funct_1
a: [3,425,0,2,1,302]
s: 425
i: [1,301,0,2,1,0,0,255] rc1_funct_1
a: [3,430,0,2,1,301]
s: 430
i: [1,305,0,2,1,0,0,255] rc5_funct_1
a: [3,435,0,2,1,305]
s: 435
i: [1,304,0,2,1,0,0,255] rc5_funct_1
a: [3,440,0,2,1,304]
s: 440
i: [1,303,0,2,1,0,0,255] rc5_funct_1
a: [3,445,0,2,1,303]
s: 445
i: [1,307,0,2,1,0,0,255] rc2_funct_1
a: [3,450,0,2,1,307]
s: 450
i: [1,306,0,2,1,0,0,255] rc2_funct_1
a: [3,455,0,2,1,306]
s: 455
i: [1,310,0,2,1,0,0,255] rc1_funcop_1
a: [3,460,0,2,1,310]
s: 460
i: [1,309,0,2,1,0,0,255] rc1_funcop_1
a: [3,465,0,2,1,309]
s: 465
i: [1,308,0,2,1,0,0,255] rc1_funcop_1
a: [3,470,0,2,1,308]
s: 470
i: [1,314,0,2,1,0,0,255] rc2_funcop_1
a: [3,475,0,2,1,314]
s: 475
i: [1,313,0,2,1,0,0,255] rc2_funcop_1
a: [3,480,0,2,1,313]
s: 480
i: [1,312,0,2,1,0,0,255] rc2_funcop_1
a: [3,485,0,2,1,312]
s: 485
i: [1,311,0,2,1,0,0,255] rc2_funcop_1
a: [3,490,0,2,1,311]
s: 490
i: [1,213,0,3,1,0,0,3] fc1_subset_1
s: 213
i: [1,277,0,3,1,0,0,3] rc2_subset_1
s: 277
i: [1,280,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 280
i: [1,281,0,3,1,0,0,2] reflexivity_r1_tarski
s: 281
i: [1,212,0,3,1,0,0,255] d2_xboole_0
a: [3,353,0,3,1,212]
s: 353
r: 343
i: [1,227,0,4,2,0,0,255] cc6_ordinal1
s: 227
i: [1,232,0,4,2,0,0,255] cc1_ordinal1
s: 232
i: [1,233,0,4,2,0,0,255] cc1_ordinal1
s: 233
d: [2,505,1,2,1,1,40,232,338]
a: [3,513,1,2,1,505]
s: 513
i: [1,237,0,4,2,0,0,255] cc9_ordinal1
s: 237
d: [2,506,1,2,1,1,40,232,333]
a: [3,518,1,2,1,506]
s: 518
i: [1,238,0,4,2,0,0,255] cc7_ordinal1
s: 238
i: [1,239,0,4,2,0,0,255] cc7_funct_1
s: 239
d: [2,507,1,2,1,1,40,232,328]
a: [3,523,1,2,1,507]
s: 523
i: [1,240,0,4,2,0,0,1] cc3_ordinal1
s: 240
d: [2,530,1,2,1,1,40,233,338]
a: [3,538,1,2,1,530]
s: 538
i: [1,241,0,4,2,0,0,255] cc1_funct_1
s: 241
i: [1,275,0,4,1,0,0,2] existence_m1_subset_1
s: 275
d: [2,531,1,2,1,1,40,233,333]
a: [3,543,1,2,1,531]
s: 543
i: [1,279,0,4,1,0,0,255] rc3_subset_1
s: 279
d: [2,532,1,2,1,1,40,233,328]
a: [3,548,1,2,1,532]
s: 548
i: [1,215,0,5,2,0,0,255] rc6_subset_1
s: 215
i: [1,219,0,5,2,0,0,3] rc1_subset_1
s: 219
d: [2,493,1,2,1,1,40,227,370]
a: [3,498,1,2,1,493]
s: 498
i: [1,221,0,5,2,0,0,255] rc5_subset_1
s: 221
d: [2,494,1,2,1,1,40,227,405]
a: [3,503,1,2,1,494]
s: 503
i: [1,222,0,5,2,0,0,255] rc5_subset_1
s: 222
i: [1,242,0,5,2,0,0,255] t6_boole
s: 242
d: [2,508,1,2,1,1,40,232,375]
a: [3,528,1,2,1,508]
s: 528
i: [1,272,0,5,2,0,0,2] t7_boole
s: 272
d: [2,533,1,2,1,1,40,233,375]
a: [3,553,1,2,1,533]
s: 553
i: [1,276,0,5,1,0,0,3] rc2_subset_1
s: 276
i: [1,278,0,5,1,0,0,255] rc3_subset_1
s: 278
d: [2,555,1,2,1,1,40,237,348]
a: [3,561,1,2,1,555]
s: 561
i: [1,284,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,360,0,5,2,284]
s: 360
d: [2,557,1,2,1,1,40,237,380]
a: [3,566,1,2,1,557]
s: 566
i: [1,224,0,6,2,0,0,255] rc4_subset_1
s: 224
i: [1,245,0,6,3,0,0,255] cc2_ordinal1
s: 245
d: [2,568,1,2,1,1,40,238,348]
a: [3,574,1,2,1,568]
s: 574
i: [1,264,0,6,2,0,0,2] d2_ordinal1
s: 264
d: [2,570,1,2,1,1,40,238,380]
a: [3,579,1,2,1,570]
s: 579
i: [1,265,0,6,2,0,0,2] d2_ordinal1
s: 265
i: [1,266,0,6,2,0,0,2] t1_subset
s: 266
d: [2,581,1,2,1,1,40,239,348]
a: [3,587,1,2,1,581]
s: 587
i: [1,282,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 282
d: [2,583,1,2,1,1,40,239,380]
a: [3,592,1,2,1,583]
s: 592
i: [1,283,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 283
i: [1,214,0,7,2,0,0,255] rc6_subset_1
s: 214
d: [2,594,1,2,1,1,40,240,348]
a: [3,600,1,2,1,594]
s: 600
i: [1,218,0,7,2,0,0,3] rc1_subset_1
s: 218
d: [2,596,1,2,1,1,40,240,380]
a: [3,605,1,2,1,596]
s: 605
i: [1,220,0,7,2,0,0,255] rc5_subset_1
s: 220
i: [1,223,0,7,2,0,0,255] rc4_subset_1
s: 223
d: [2,607,1,2,1,1,40,241,348]
a: [3,613,1,2,1,607]
s: 613
i: [1,228,0,7,3,0,0,255] cc8_funct_1
s: 228
d: [2,609,1,2,1,1,40,241,380]
a: [3,618,1,2,1,609]
s: 618
i: [1,229,0,7,3,0,0,255] cc8_funct_1
s: 229
i: [1,235,0,7,3,0,0,1] cc2_ordinal2
s: 235
d: [2,620,2,2,1,1,40,498,233]
a: [3,625,2,2,1,620]
s: 625
i: [1,236,0,7,3,0,0,1] cc5_ordinal1
s: 236
d: [2,621,2,2,1,1,40,498,232]
a: [3,630,2,2,1,621]
s: 630
i: [1,269,0,7,2,0,0,2] t3_subset
s: 269
i: [1,270,0,7,2,0,0,2] t3_subset
s: 270
d: [2,632,2,2,1,1,40,503,233]
a: [3,637,2,2,1,632]
s: 637
i: [1,271,0,7,3,0,0,255] t8_boole
s: 271
d: [2,633,2,2,1,1,40,503,232]
a: [3,642,2,2,1,633]
s: 642
i: [1,225,0,8,3,0,0,255] cc10_ordinal1
s: 225
i: [1,230,0,8,3,0,0,255] cc9_funct_1
s: 230
d: [2,661,2,2,1,1,40,600,233]
a: [3,667,2,2,1,661]
s: 667
i: [1,231,0,8,3,0,0,255] cc5_subset_1
s: 231
d: [2,662,2,2,1,1,40,600,232]
a: [3,672,2,2,1,662]
s: 672
i: [1,243,0,8,3,0,0,3] cc1_subset_1
s: 243
i: [1,248,0,8,4,0,0,255] cc6_funct_1
s: 248
d: [2,674,2,2,1,1,40,605,233]
a: [3,679,2,2,1,674]
s: 679
i: [1,253,0,8,4,0,0,255] cc4_funct_1
s: 253
d: [2,675,2,2,1,1,40,605,232]
a: [3,684,2,2,1,675]
s: 684
d: [2,758,0,6,3,0,21,253,241]
s: 758
i: [1,254,0,8,4,0,0,255] cc5_funct_1
s: 254
d: [2,759,1,4,2,1,40,758,430]
a: [3,767,1,2,1,759]
s: 767
i: [1,261,0,8,4,0,0,255] cc3_funcop_1
s: 261
d: [2,760,1,4,2,1,40,758,445]
d: [2,773,1,2,1,2,21,760,435]
a: [3,777,1,2,1,773]
s: 777
d: [2,797,0,6,3,0,21,261,241]
s: 797
i: [1,263,0,8,3,0,0,2] d2_ordinal1
s: 263
d: [2,761,1,4,2,1,40,758,455]
a: [3,782,1,2,1,761]
s: 782
i: [1,267,0,8,3,0,0,2] t2_subset
s: 267
d: [2,762,1,4,2,1,40,758,470]
a: [3,791,1,2,1,762]
s: 791
i: [1,285,0,8,3,0,0,1] dt_k10_ordinal2
s: 285
i: [1,244,0,9,3,0,0,255] cc4_subset_1
s: 244
d: [2,556,1,3,1,0,40,237,277]
s: 556
i: [1,268,0,9,3,0,0,255] d8_xboole_0
s: 268
d: [2,569,1,3,1,0,40,238,277]
s: 569
i: [1,274,0,9,3,0,0,2] t5_subset
s: 274
i: [1,226,0,10,4,0,0,255] t11_ordinal1
s: 226
d: [2,582,1,3,1,0,40,239,277]
s: 582
i: [1,262,0,10,4,0,0,255] cc3_funct_1
s: 262
d: [2,595,1,3,1,0,40,240,277]
s: 595
i: [1,273,0,10,3,0,0,2] t4_subset
s: 273
i: [1,286,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 286
d: [2,862,1,15,3,1,40,286,318]
a: [3,868,1,4,1,862]
s: 868
d: [2,878,2,0,0,3,52,328,333,868,285]
e: 878
d: [2,879,2,4,2,1,40,868,285]
d: [2,881,2,2,1,2,21,879,333]
d: [2,882,2,0,0,3,21,881,328]
e: 882
r: 868
i: [1,287,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 287
d: [2,884,1,7,3,1,40,287,323]
d: [2,889,1,5,2,2,21,884,333]
d: [2,890,1,3,1,3,21,889,338]
a: [3,894,1,3,1,890]
s: 894
i: [1,288,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 288
i: [1,216,0,11,4,0,0,255] cc3_subset_1
s: 216
d: [2,896,2,6,2,1,40,894,268]
a: [3,900,2,3,1,896]
s: 900
i: [1,217,0,11,4,0,0,255] cc2_subset_1
s: 217
a: [3,903,2,3,1,896]
s: 903
i: [1,234,0,16,5,0,0,1] t32_ordinal2
s: 234
a: [3,867,1,4,1,862]
s: 867
d: [2,922,3,7,3,1,40,900,226]
d: [2,924,3,5,2,2,21,922,333]
d: [2,925,3,3,1,3,21,924,513]
a: [3,929,3,3,1,925]
s: 929
a: [3,876,1,7,1,862]
s: 876
d: [2,957,4,2,1,1,40,929,272]
a: [3,971,4,2,1,957]
s: 971
d: [2,608,1,3,1,0,40,241,277]
s: 608
d: [2,655,1,4,2,0,40,264,272]
s: 655
d: [2,923,3,3,1,1,40,900,282]
a: [3,934,3,3,1,923]
s: 934
d: [2,714,1,4,1,0,40,269,276]
s: 714
d: [2,955,4,3,1,1,40,929,283]
a: [3,961,4,3,1,955]
s: 961
d: [2,715,1,4,1,0,40,269,278]
s: 715
d: [2,691,1,5,2,0,40,228,275]
s: 691
d: [2,956,4,3,1,1,40,929,266]
a: [3,966,4,3,1,956]
s: 966
d: [2,698,1,5,2,0,40,229,275]
s: 698
d: [2,992,5,4,2,1,40,966,229]
a: [3,997,5,2,1,992]
s: 997
d: [2,705,1,5,2,0,40,235,275]
s: 705
d: [2,713,1,5,1,0,40,269,275]
s: 713
d: [2,942,2,4,1,1,40,867,233]
a: [3,947,2,4,1,942]
s: 947
d: [2,727,1,5,2,0,40,225,276]
s: 727
d: [2,943,2,4,1,1,40,867,232]
a: [3,952,2,4,1,943]
s: 952
d: [2,728,1,5,2,0,40,225,278]
s: 728
d: [2,735,1,5,2,0,40,230,276]
s: 735
d: [2,973,2,15,3,1,40,876,287]
d: [2,974,2,11,2,2,21,973,867]
a: [3,978,2,7,1,974]
s: 978
d: [2,736,1,5,2,0,40,230,278]
s: 736
d: [2,1014,3,14,2,1,40,978,268]
a: [3,1018,3,7,1,1014]
s: 1018
d: [2,745,1,5,2,0,40,231,276]
s: 745
d: [2,746,1,5,2,0,40,231,278]
s: 746
d: [2,1025,4,15,3,1,40,1018,226]
d: [2,1027,4,11,2,2,21,1025,952]
a: [3,1031,4,7,1,1027]
s: 1031
d: [2,755,1,5,2,0,40,243,278]
s: 755
d: [2,1044,5,4,1,1,40,1031,272]
a: [3,1058,5,4,1,1044]
s: 1058
d: [2,815,1,8,2,0,40,267,276]
d: [2,822,1,5,1,0,21,815,213]
s: 822
d: [2,816,1,8,2,0,40,267,278]
d: [2,823,1,5,1,0,21,816,213]
s: 823
d: [2,1026,4,7,1,1,40,1018,282]
a: [3,1036,4,7,1,1026]
s: 1036
d: [2,916,1,9,3,0,40,216,278]
d: [2,921,1,5,2,0,21,916,279]
s: 921
d: [2,1042,5,7,1,1,40,1031,283]
a: [3,1048,5,7,1,1042]
s: 1048
d: [2,1075,6,0,0,5,52,338,333,328,929,1048,234]
e: 1075
d: [2,1076,6,9,4,1,40,1048,234]
d: [2,1078,6,6,3,2,21,1076,929]
d: [2,1079,6,4,2,3,21,1078,328]
d: [2,1080,6,2,1,4,21,1079,333]
d: [2,1081,6,0,0,5,21,1080,338]
e: 1081
r: 978
r: 1018
r: 1031
r: 1048
a: [3,872,1,4,1,862]
s: 872
d: [2,1083,2,0,0,3,52,328,338,872,285]
e: 1083
d: [2,1084,2,4,2,1,40,872,285]
d: [2,1086,2,2,1,2,21,1084,338]
d: [2,1087,2,0,0,3,21,1086,328]
e: 1087
f: 1087,1081,1042,1027,1014,974,943,925,896,890,882,862,505,205,206,207,208,209
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal2__t33_ordinal2
% SZS output start Proof for ordinal2__t33_ordinal2
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(X0,X1) => r1_ordinal1(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)))))) [input t33_ordinal2]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(X0,X1) => r1_ordinal1(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)))))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
7. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
31. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
34. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
35. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k10_ordinal2(X0,X1))) [input dt_k10_ordinal2]
62. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
65. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
68. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(X0,X1) => r2_hidden(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)))))) [input t32_ordinal2]
77. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 34]
85. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 77]
90. ? [X0] : (? [X1] : (? [X2] : ((~r1_ordinal1(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) & r1_ordinal1(X0,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
91. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) & r1_ordinal1(X0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 90]
101. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 65]
102. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 101]
107. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 7]
108. ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) | ~r2_hidden(X0,X1)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 68]
109. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) | ~r2_hidden(X0,X1) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 108]
140. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 85]
141. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 140]
148. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
151. ! [X0,X1] : (v3_ordinal1(k10_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 35]
152. ! [X0,X1] : (v3_ordinal1(k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 151]
153. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 31]
154. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 153]
155. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 62]
156. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 155]
157. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) & r1_ordinal1(X0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (? [X2] : (~r1_ordinal1(k10_ordinal2(X2,sK0),k10_ordinal2(X2,X1)) & r1_ordinal1(sK0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
158. ? [X1] : (? [X2] : (~r1_ordinal1(k10_ordinal2(X2,sK0),k10_ordinal2(X2,X1)) & r1_ordinal1(sK0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (~r1_ordinal1(k10_ordinal2(X2,sK0),k10_ordinal2(X2,sK1)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(X2)) & v3_ordinal1(sK1)) [choice axiom]
159. ? [X2] : (~r1_ordinal1(k10_ordinal2(X2,sK0),k10_ordinal2(X2,sK1)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(X2)) => (~r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(sK2)) [choice axiom]
160. ((~r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(sK2)) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 91,159,158,157]
180. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 156]
205. v3_ordinal1(sK0) [cnf transformation 160] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
206. v3_ordinal1(sK1) [cnf transformation 160] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
207. v3_ordinal1(sK2) [cnf transformation 160] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
208. r1_ordinal1(sK0,sK1) [cnf transformation 160] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
209. ~r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) [cnf transformation 160] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
226. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 102] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
232. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 107] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
234. r2_hidden(k10_ordinal2(X2,X0),k10_ordinal2(X2,X1)) | ~r2_hidden(X0,X1) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 109] {a:0,w:16,wCS:6400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
268. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 141] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
283. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 148] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
285. v3_ordinal1(k10_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 152] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
286. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 154] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
287. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 180] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
316. 1 <=> r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) [avatar definition]
318. ~r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) <- (~1) [avatar component clause 316] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
319. ~1 [avatar split clause 209,316]
321. 2 <=> r1_ordinal1(sK0,sK1) [avatar definition]
323. r1_ordinal1(sK0,sK1) <- (2) [avatar component clause 321] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
324. 2 [avatar split clause 208,321]
326. 3 <=> v3_ordinal1(sK2) [avatar definition]
328. v3_ordinal1(sK2) <- (3) [avatar component clause 326] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
329. 3 [avatar split clause 207,326]
331. 4 <=> v3_ordinal1(sK1) [avatar definition]
333. v3_ordinal1(sK1) <- (4) [avatar component clause 331] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
334. 4 [avatar split clause 206,331]
336. 5 <=> v3_ordinal1(sK0) [avatar definition]
338. v3_ordinal1(sK0) <- (5) [avatar component clause 336] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
339. 5 [avatar split clause 205,336]
505. v1_ordinal1(sK0) <- (5) [resolution 232,338] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
511. 39 <=> v1_ordinal1(sK0) [avatar definition]
513. v1_ordinal1(sK0) <- (39) [avatar component clause 511] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
514. 39 | ~5 [avatar split clause 505,336,511]
862. r1_ordinal1(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK1)) <- (~1) [resolution 286,318] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
866. 73 <=> v3_ordinal1(k10_ordinal2(sK2,sK1)) [avatar definition]
867. v3_ordinal1(k10_ordinal2(sK2,sK1)) <- (73) [avatar component clause 866] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
868. ~v3_ordinal1(k10_ordinal2(sK2,sK1)) <- (~73) [avatar component clause 866] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
870. 74 <=> v3_ordinal1(k10_ordinal2(sK2,sK0)) [avatar definition]
872. ~v3_ordinal1(k10_ordinal2(sK2,sK0)) <- (~74) [avatar component clause 870] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
874. 75 <=> r1_ordinal1(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [avatar definition]
876. r1_ordinal1(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) <- (75) [avatar component clause 874] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
877. ~73 | ~74 | 75 | 1 [avatar split clause 862,316,874,870,866]
879. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK2) <- (~73) [resolution 868,285] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
881. ~v3_ordinal1(sK2) <- (4, ~73) [subsumption resolution 879,333] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
882. $false <- (3, 4, ~73) [subsumption resolution 881,328] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
883. ~3 | ~4 | 73 [avatar contradiction clause 882]
884. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (2) [resolution 287,323] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
889. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK0) <- (2, 4) [subsumption resolution 884,333] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
890. r1_tarski(sK0,sK1) <- (2, 4, 5) [subsumption resolution 889,338] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
892. 76 <=> r1_tarski(sK0,sK1) [avatar definition]
894. r1_tarski(sK0,sK1) <- (76) [avatar component clause 892] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
895. 76 | ~2 | ~4 | ~5 [avatar split clause 890,336,331,321,892]
896. sK0 = sK1 | r2_xboole_0(sK0,sK1) <- (76) [resolution 894,268] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
898. 77 <=> r2_xboole_0(sK0,sK1) [avatar definition]
900. r2_xboole_0(sK0,sK1) <- (77) [avatar component clause 898] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
902. 78 <=> sK0 = sK1 [avatar definition]
905. 77 | 78 | ~76 [avatar split clause 896,892,902,898]
922. r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v1_ordinal1(sK0) <- (77) [resolution 900,226] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
924. r2_hidden(sK0,sK1) | ~v1_ordinal1(sK0) <- (4, 77) [subsumption resolution 922,333] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
925. r2_hidden(sK0,sK1) <- (4, 39, 77) [subsumption resolution 924,513] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
927. 79 <=> r2_hidden(sK0,sK1) [avatar definition]
929. r2_hidden(sK0,sK1) <- (79) [avatar component clause 927] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
930. 79 | ~4 | ~39 | ~77 [avatar split clause 925,898,511,331,927]
937. sK0 != sK1 | r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) | ~r1_ordinal1(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [theory tautology sat conflict] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:1,thDist:-1}
943. v1_ordinal1(k10_ordinal2(sK2,sK1)) <- (73) [resolution 867,232] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
950. 82 <=> v1_ordinal1(k10_ordinal2(sK2,sK1)) [avatar definition]
952. v1_ordinal1(k10_ordinal2(sK2,sK1)) <- (82) [avatar component clause 950] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
953. 82 | ~73 [avatar split clause 943,866,950]
973. r1_tarski(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK1)) <- (75) [resolution 876,287] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
974. r1_tarski(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK0)) <- (73, 75) [subsumption resolution 973,867] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
976. 86 <=> r1_tarski(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [avatar definition]
978. r1_tarski(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) <- (86) [avatar component clause 976] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
979. ~74 | 86 | ~73 | ~75 [avatar split clause 974,874,866,976,870]
1014. k10_ordinal2(sK2,sK0) = k10_ordinal2(sK2,sK1) | r2_xboole_0(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) <- (86) [resolution 978,268] {a:3,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1016. 90 <=> r2_xboole_0(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [avatar definition]
1018. r2_xboole_0(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) <- (90) [avatar component clause 1016] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1020. 91 <=> k10_ordinal2(sK2,sK0) = k10_ordinal2(sK2,sK1) [avatar definition]
1023. 90 | 91 | ~86 [avatar split clause 1014,976,1020,1016]
1024. k10_ordinal2(sK2,sK0) != k10_ordinal2(sK2,sK1) | r1_ordinal1(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) | ~r1_ordinal1(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [theory tautology sat conflict] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
1025. r2_hidden(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK0)) | ~v1_ordinal1(k10_ordinal2(sK2,sK1)) <- (90) [resolution 1018,226] {a:4,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1027. r2_hidden(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) | ~v3_ordinal1(k10_ordinal2(sK2,sK0)) <- (82, 90) [subsumption resolution 1025,952] {a:4,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1029. 92 <=> r2_hidden(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) [avatar definition]
1031. r2_hidden(k10_ordinal2(sK2,sK1),k10_ordinal2(sK2,sK0)) <- (92) [avatar component clause 1029] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1032. ~74 | 92 | ~82 | ~90 [avatar split clause 1027,1016,950,1029,870]
1042. ~r2_hidden(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) <- (92) [resolution 1031,283] {a:5,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
1046. 94 <=> r2_hidden(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) [avatar definition]
1048. ~r2_hidden(k10_ordinal2(sK2,sK0),k10_ordinal2(sK2,sK1)) <- (~94) [avatar component clause 1046] {a:5,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1049. ~94 | ~92 [avatar split clause 1042,1029,1046]
1076. ~r2_hidden(sK0,sK1) | ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~94) [resolution 1048,234] {a:6,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1078. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (79, ~94) [subsumption resolution 1076,929] {a:6,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1079. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (3, 79, ~94) [subsumption resolution 1078,328] {a:6,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1080. ~v3_ordinal1(sK0) <- (3, 4, 79, ~94) [subsumption resolution 1079,333] {a:6,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1081. $false <- (3, 4, 5, 79, ~94) [subsumption resolution 1080,338] {a:6,w:0,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1082. ~3 | ~4 | ~5 | ~79 | 94 [avatar contradiction clause 1081]
1084. ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (~74) [resolution 872,285] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1086. ~v3_ordinal1(sK2) <- (5, ~74) [subsumption resolution 1084,338] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1087. $false <- (3, 5, ~74) [subsumption resolution 1086,328] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1088. ~3 | ~5 | 74 [avatar contradiction clause 1087]
1089. $false [avatar sat refutation 319,324,329,334,339,514,877,883,895,905,930,937,953,979,1023,1024,1032,1049,1082,1088] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:580,thDist:-580}
% SZS output end Proof for ordinal2__t33_ordinal2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 73
% >>> Preprocessing
% Introduced skolems: 23
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 110
% Generated clauses: 497
% Active clauses: 168
% Passive clauses: 298
% Final active clauses: 162
% Final passive clauses: 122
% Main loop iterations started: 175
% 
% >>> Simplifying Inferences
% Duplicate literals: 9
% Fw subsumption resolutions: 28
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Simple tautologies: 26
% Equational tautologies: 1
% Forward subsumptions: 49
% 
% >>> Generating Inferences
% Binary resolution: 245
% Unit resulting resolution: 3
% Factoring: 1
% Forward superposition: 7
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 12
% Split components: 25
% Sat splitting refutations: 6
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 100
% SAT solver unit clauses: 34
% SAT solver binary clauses: 41
% 
% Memory used [KB]: 5628
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.002 s
% congruence closure: 0.001 s
% unit resulting resolution: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

