i: [1,115,0,2,1,1,0,0] t28_ordinal3
a: [3,174,0,2,1,115]
s: 174
i: [1,114,0,2,1,1,0,0] t28_ordinal3
a: [3,179,0,2,1,114]
s: 179
i: [1,116,0,3,1,1,0,0] t28_ordinal3
a: [3,169,0,3,1,116]
s: 169
i: [1,118,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,184,0,2,1,118]
s: 184
i: [1,119,0,2,1,0,0,1] fc1_xboole_0
a: [3,189,0,2,1,119]
s: 189
i: [1,156,0,2,1,0,0,2] rc2_xboole_0
a: [3,206,0,2,1,156]
s: 206
i: [1,157,0,2,1,0,0,2] rc4_ordinal1
a: [3,211,0,2,1,157]
s: 211
i: [1,158,0,2,1,0,0,2] rc1_xboole_0
a: [3,216,0,2,1,158]
s: 216
i: [1,159,0,2,1,0,0,1] rc1_ordinal1
a: [3,221,0,2,1,159]
s: 221
i: [1,163,0,2,1,0,0,255] rc2_ordinal1
a: [3,226,0,2,1,163]
s: 226
i: [1,162,0,2,1,0,0,255] rc2_ordinal1
a: [3,231,0,2,1,162]
s: 231
i: [1,161,0,2,1,0,0,255] rc2_ordinal1
a: [3,236,0,2,1,161]
s: 236
i: [1,160,0,2,1,0,0,255] rc2_ordinal1
a: [3,241,0,2,1,160]
s: 241
i: [1,165,0,2,1,0,0,2] rc5_ordinal1
a: [3,246,0,2,1,165]
s: 246
i: [1,164,0,2,1,0,0,2] rc5_ordinal1
a: [3,251,0,2,1,164]
s: 251
i: [1,117,0,10,3,1,0,0] t28_ordinal3
s: 117
i: [1,147,0,3,1,0,0,255] reflexivity_r1_tarski
s: 147
i: [1,120,0,3,1,0,0,255] d2_xboole_0
a: [3,194,0,3,1,120]
s: 194
r: 184
i: [1,122,0,4,2,0,0,2] cc6_ordinal1
s: 122
i: [1,123,0,4,2,0,0,255] cc9_ordinal1
s: 123
i: [1,124,0,4,2,0,0,2] cc7_ordinal1
s: 124
d: [2,254,1,2,1,1,40,122,211]
a: [3,259,1,2,1,254]
s: 259
i: [1,125,0,4,2,0,0,2] cc3_ordinal1
s: 125
d: [2,255,1,2,1,1,40,122,246]
a: [3,264,1,2,1,255]
s: 264
i: [1,128,0,4,2,0,0,255] cc1_ordinal1
s: 128
i: [1,129,0,4,2,0,0,255] cc1_ordinal1
s: 129
d: [2,302,1,2,1,1,40,128,179]
a: [3,311,1,2,1,302]
s: 311
i: [1,146,0,4,1,0,0,2] existence_m1_subset_1
s: 146
d: [2,303,1,2,1,1,40,128,174]
a: [3,316,1,2,1,303]
s: 316
i: [1,126,0,5,2,0,0,1] t6_boole
s: 126
i: [1,143,0,5,2,0,0,1] t7_boole
s: 143
d: [2,333,1,2,1,1,40,129,179]
a: [3,342,1,2,1,333]
s: 342
i: [1,151,0,7,3,0,0,255] reflexivity_r1_ordinal1
a: [3,201,0,5,2,151]
s: 201
d: [2,334,1,2,1,1,40,129,174]
a: [3,347,1,2,1,334]
s: 347
i: [1,134,0,6,3,0,0,255] cc2_ordinal1
s: 134
i: [1,136,0,6,2,0,0,255] d2_ordinal1
s: 136
d: [2,366,1,2,1,1,40,143,169]
a: [3,370,1,2,1,366]
s: 370
i: [1,137,0,6,2,0,0,255] d2_ordinal1
s: 137
d: [2,266,1,2,1,1,40,123,216]
a: [3,271,1,2,1,266]
s: 271
i: [1,138,0,6,2,0,0,1] t1_subset
s: 138
i: [1,148,0,6,2,0,0,1] antisymmetry_r2_hidden
s: 148
d: [2,376,1,3,1,1,40,138,169]
a: [3,381,1,3,1,376]
s: 381
i: [1,127,0,7,2,0,0,1] t27_ordinal2
s: 127
d: [2,383,1,3,1,1,40,148,169]
a: [3,388,1,3,1,383]
s: 388
i: [1,132,0,7,3,0,0,2] cc2_ordinal2
s: 132
i: [1,133,0,7,3,0,0,2] cc5_ordinal1
s: 133
d: [2,267,1,2,1,1,40,123,189]
a: [3,276,1,2,1,267]
s: 276
i: [1,140,0,7,2,0,0,255] t3_subset
s: 140
d: [2,278,1,2,1,1,40,124,216]
a: [3,283,1,2,1,278]
s: 283
i: [1,141,0,7,2,0,0,255] t3_subset
s: 141
i: [1,142,0,7,3,0,0,1] t8_boole
s: 142
d: [2,279,1,2,1,1,40,124,189]
a: [3,288,1,2,1,279]
s: 288
i: [1,121,0,8,3,0,0,255] cc10_ordinal1
s: 121
d: [2,290,1,2,1,1,40,125,216]
a: [3,295,1,2,1,290]
s: 295
i: [1,135,0,8,3,0,0,255] d2_ordinal1
s: 135
i: [1,139,0,8,3,0,0,1] t2_subset
s: 139
d: [2,417,1,5,2,1,40,135,169]
d: [2,419,1,3,1,2,21,417,316]
a: [3,423,1,3,1,419]
s: 423
i: [1,149,0,8,3,0,0,1] fc5_ordinal2
s: 149
d: [2,291,1,2,1,1,40,125,189]
a: [3,300,1,2,1,291]
s: 300
i: [1,150,0,8,3,0,0,1] fc5_ordinal2
s: 150
i: [1,152,0,8,3,0,0,1] dt_k10_ordinal2
s: 152
d: [2,305,1,2,1,1,40,128,221]
a: [3,326,1,2,1,305]
s: 326
i: [1,145,0,9,3,0,0,255] t5_subset
s: 145
d: [2,336,1,2,1,1,40,129,221]
a: [3,357,1,2,1,336]
s: 357
i: [1,144,0,10,3,0,0,255] t4_subset
s: 144
i: [1,153,0,10,4,0,0,255] connectedness_r1_ordinal1
s: 153
d: [2,304,2,2,1,1,40,128,259]
a: [3,321,2,2,1,304]
s: 321
i: [1,154,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 154
d: [2,307,2,2,1,1,40,128,264]
a: [3,331,2,2,1,307]
s: 331
i: [1,155,0,10,4,0,0,255] redefinition_r1_ordinal1
s: 155
i: [1,130,0,11,4,0,0,255] t27_ordinal3
s: 130
d: [2,462,2,7,3,1,40,155,423]
d: [2,464,2,5,2,2,21,462,174]
d: [2,465,2,3,1,3,21,464,179]
a: [3,469,2,3,1,465]
s: 469
i: [1,131,0,14,4,0,0,255] t27_ordinal3
s: 131
d: [2,335,2,2,1,1,40,129,259]
a: [3,352,2,2,1,335]
s: 352
d: [2,375,1,4,2,0,40,136,143]
s: 375
d: [2,474,1,19,6,0,44,117,131]
d: [2,479,1,17,5,1,21,474,179]
s: 479
d: [2,486,2,14,4,1,47,479]
d: [2,487,2,11,3,2,21,486,469]
d: [2,488,2,9,2,3,21,487,174]
a: [3,492,2,4,1,488]
s: 492
d: [2,498,3,0,0,4,52,179,174,469,492,130]
e: 498
d: [2,499,3,7,3,1,40,492,130]
d: [2,501,3,4,2,2,21,499,469]
d: [2,502,3,2,1,3,21,501,174]
d: [2,503,3,0,0,4,21,502,179]
e: 503
r: 492
d: [2,392,1,5,2,0,40,132,146]
s: 392
a: [3,491,2,4,1,488]
s: 491
d: [2,395,1,5,1,0,40,140,146]
s: 395
d: [2,377,1,6,2,0,40,138,136]
s: 377
d: [2,507,3,4,1,1,40,491,129]
a: [3,512,3,4,1,507]
s: 512
d: [2,384,1,6,2,0,40,148,136]
s: 384
d: [2,508,3,4,1,1,40,491,128]
a: [3,517,3,4,1,508]
s: 517
d: [2,403,1,6,2,0,40,121,146]
s: 403
d: [2,425,1,6,2,0,40,139,146]
s: 425
a: [3,496,2,5,1,488]
s: 496
r: 517
r: 512
r: 491
d: [2,364,1,7,3,0,44,126,126]
s: 364
d: [2,534,3,12,4,1,44,131,496]
d: [2,538,3,9,3,2,21,534,469]
d: [2,539,3,7,2,3,21,538,174]
d: [2,540,3,5,1,4,21,539,179]
a: [3,544,3,5,1,540]
s: 544
d: [2,549,4,5,2,1,44,127,544]
d: [2,559,4,3,1,2,21,549,179]
f: 559,540,503,488,465,419,383,303,114,115,116
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal3__t28_ordinal3
% SZS output start Proof for ordinal3__t28_ordinal3
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(! [X2] : (v3_ordinal1(X2) => ~(k1_xboole_0 != X2 & k10_ordinal2(X0,X2) = X1)) & r2_hidden(X0,X1)))) [input t28_ordinal3]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(! [X2] : (v3_ordinal1(X2) => ~(k1_xboole_0 != X2 & k10_ordinal2(X0,X2) = X1)) & r2_hidden(X0,X1)))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
5. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
15. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r2_hidden(X1,X0) => r1_tarski(X1,X0))) [input d2_ordinal1]
31. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
35. ! [X0] : (v3_ordinal1(X0) => k10_ordinal2(X0,k1_xboole_0) = X0) [input t27_ordinal2]
36. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ~(! [X2] : (v3_ordinal1(X2) => k10_ordinal2(X0,X2) != X1) & r1_ordinal1(X0,X1)))) [input t27_ordinal3]
53. ? [X0] : (? [X1] : ((! [X2] : ((k1_xboole_0 = X2 | k10_ordinal2(X0,X2) != X1) | ~v3_ordinal1(X2)) & r2_hidden(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
54. ? [X0] : (? [X1] : (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2)) & r2_hidden(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 53]
61. ! [X0] : (k10_ordinal2(X0,k1_xboole_0) = X0 | ~v3_ordinal1(X0)) [ennf transformation 35]
62. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 5]
63. ! [X0] : (! [X1] : ((? [X2] : (k10_ordinal2(X0,X2) = X1 & v3_ordinal1(X2)) | ~r1_ordinal1(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 36]
64. ! [X0] : (! [X1] : (? [X2] : (k10_ordinal2(X0,X2) = X1 & v3_ordinal1(X2)) | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 63]
69. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0))) [ennf transformation 15]
78. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
87. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 31]
88. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 87]
89. ? [X0] : (? [X1] : (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(X0,X2) != X1 | ~v3_ordinal1(X2)) & r2_hidden(X0,X1) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(sK0,X2) != X1 | ~v3_ordinal1(X2)) & r2_hidden(sK0,X1) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
90. ? [X1] : (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(sK0,X2) != X1 | ~v3_ordinal1(X2)) & r2_hidden(sK0,X1) & v3_ordinal1(X1)) => (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(sK0,X2) != sK1 | ~v3_ordinal1(X2)) & r2_hidden(sK0,sK1) & v3_ordinal1(sK1)) [choice axiom]
91. (! [X2] : (k1_xboole_0 = X2 | k10_ordinal2(sK0,X2) != sK1 | ~v3_ordinal1(X2)) & r2_hidden(sK0,sK1) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 54,90,89]
92. ! [X1,X0] : (? [X2] : (k10_ordinal2(X0,X2) = X1 & v3_ordinal1(X2)) => (k10_ordinal2(X0,sK2(X0,X1)) = X1 & v3_ordinal1(sK2(X0,X1)))) [choice axiom]
93. ! [X0] : (! [X1] : ((k10_ordinal2(X0,sK2(X0,X1)) = X1 & v3_ordinal1(sK2(X0,X1))) | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [skolemisation 64,92]
94. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)) | ~v1_ordinal1(X0))) [nnf transformation 69]
95. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [rectify 94]
96. ! [X0] : (? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0)) => (~r1_tarski(sK3(X0),X0) & r2_hidden(sK3(X0),X0))) [choice axiom]
97. ! [X0] : ((v1_ordinal1(X0) | (~r1_tarski(sK3(X0),X0) & r2_hidden(sK3(X0),X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [skolemisation 95,96]
101. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 88]
114. v3_ordinal1(sK0) [cnf transformation 91] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
115. v3_ordinal1(sK1) [cnf transformation 91] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
116. r2_hidden(sK0,sK1) [cnf transformation 91] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
117. k10_ordinal2(sK0,X2) != sK1 | k1_xboole_0 = X2 | ~v3_ordinal1(X2) [cnf transformation 91] {a:0,w:10,wCS:1000,nSel:1,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
127. k10_ordinal2(X0,k1_xboole_0) = X0 | ~v3_ordinal1(X0) [cnf transformation 61] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
128. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 62] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
130. v3_ordinal1(sK2(X0,X1)) | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 93] {a:0,w:11,wCS:4400,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
131. k10_ordinal2(X0,sK2(X0,X1)) = X1 | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 93] {a:0,w:14,wCS:5600,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
135. ~r2_hidden(X2,X0) | r1_tarski(X2,X0) | ~v1_ordinal1(X0) [cnf transformation 97] {a:0,w:8,wCS:3200,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
148. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 78] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
155. ~r1_tarski(X0,X1) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 101] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
167. 1 <=> r2_hidden(sK0,sK1) [avatar definition]
169. r2_hidden(sK0,sK1) <- (1) [avatar component clause 167] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
170. 1 [avatar split clause 116,167]
172. 2 <=> v3_ordinal1(sK1) [avatar definition]
174. v3_ordinal1(sK1) <- (2) [avatar component clause 172] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
175. 2 [avatar split clause 115,172]
177. 3 <=> v3_ordinal1(sK0) [avatar definition]
179. v3_ordinal1(sK0) <- (3) [avatar component clause 177] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
180. 3 [avatar split clause 114,177]
303. v1_ordinal1(sK1) <- (2) [resolution 128,174] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
314. 28 <=> v1_ordinal1(sK1) [avatar definition]
316. v1_ordinal1(sK1) <- (28) [avatar component clause 314] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
317. 28 | ~2 [avatar split clause 303,172,314]
383. ~r2_hidden(sK1,sK0) <- (1) [resolution 148,169] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
386. 39 <=> r2_hidden(sK1,sK0) [avatar definition]
389. ~39 | ~1 [avatar split clause 383,167,386]
417. r1_tarski(sK0,sK1) | ~v1_ordinal1(sK1) <- (1) [resolution 135,169] {a:1,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
419. r1_tarski(sK0,sK1) <- (1, 28) [subsumption resolution 417,316] {a:1,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
421. 42 <=> r1_tarski(sK0,sK1) [avatar definition]
423. r1_tarski(sK0,sK1) <- (42) [avatar component clause 421] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
424. 42 | ~1 | ~28 [avatar split clause 419,314,167,421]
462. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (42) [resolution 155,423] {a:2,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
464. r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK0) <- (2, 42) [subsumption resolution 462,174] {a:2,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
465. r1_ordinal1(sK0,sK1) <- (2, 3, 42) [subsumption resolution 464,179] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
467. 45 <=> r1_ordinal1(sK0,sK1) [avatar definition]
469. r1_ordinal1(sK0,sK1) <- (45) [avatar component clause 467] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
470. 45 | ~2 | ~3 | ~42 [avatar split clause 465,421,177,172,467]
474. sK1 != X0 | k1_xboole_0 = sK2(sK0,X0) | ~v3_ordinal1(sK2(sK0,X0)) | ~r1_ordinal1(sK0,X0) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK0) [superposition 117,131] {a:1,w:19,wCS:1900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
479. sK1 != X0 | k1_xboole_0 = sK2(sK0,X0) | ~v3_ordinal1(sK2(sK0,X0)) | ~r1_ordinal1(sK0,X0) | ~v3_ordinal1(X0) <- (3) [subsumption resolution 474,179] {a:1,w:17,wCS:1700,nSel:1,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
486. k1_xboole_0 = sK2(sK0,sK1) | ~v3_ordinal1(sK2(sK0,sK1)) | ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) <- (3) [equality resolution 479] {a:2,w:14,wCS:1400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
487. k1_xboole_0 = sK2(sK0,sK1) | ~v3_ordinal1(sK2(sK0,sK1)) | ~v3_ordinal1(sK1) <- (3, 45) [subsumption resolution 486,469] {a:2,w:11,wCS:1100,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
488. k1_xboole_0 = sK2(sK0,sK1) | ~v3_ordinal1(sK2(sK0,sK1)) <- (2, 3, 45) [subsumption resolution 487,174] {a:2,w:9,wCS:900,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
490. 46 <=> v3_ordinal1(sK2(sK0,sK1)) [avatar definition]
492. ~v3_ordinal1(sK2(sK0,sK1)) <- (~46) [avatar component clause 490] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
494. 47 <=> k1_xboole_0 = sK2(sK0,sK1) [avatar definition]
496. k1_xboole_0 = sK2(sK0,sK1) <- (47) [avatar component clause 494] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
497. ~46 | 47 | ~2 | ~3 | ~45 [avatar split clause 488,467,177,172,494,490]
499. ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~46) [resolution 492,130] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
501. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (45, ~46) [subsumption resolution 499,469] {a:3,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
502. ~v3_ordinal1(sK0) <- (2, 45, ~46) [subsumption resolution 501,174] {a:3,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
503. $false <- (2, 3, 45, ~46) [subsumption resolution 502,179] {a:3,w:0,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
504. ~2 | ~3 | ~45 | 46 [avatar contradiction clause 503]
534. sK1 = k10_ordinal2(sK0,k1_xboole_0) | ~r1_ordinal1(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (47) [superposition 131,496] {a:3,w:12,wCS:1200,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
538. sK1 = k10_ordinal2(sK0,k1_xboole_0) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (45, 47) [subsumption resolution 534,469] {a:3,w:9,wCS:900,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
539. sK1 = k10_ordinal2(sK0,k1_xboole_0) | ~v3_ordinal1(sK0) <- (2, 45, 47) [subsumption resolution 538,174] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
540. sK1 = k10_ordinal2(sK0,k1_xboole_0) <- (2, 3, 45, 47) [subsumption resolution 539,179] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
542. 50 <=> sK1 = k10_ordinal2(sK0,k1_xboole_0) [avatar definition]
544. sK1 = k10_ordinal2(sK0,k1_xboole_0) <- (50) [avatar component clause 542] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
545. 50 | ~2 | ~3 | ~45 | ~47 [avatar split clause 540,494,467,177,172,542]
549. sK0 = sK1 | ~v3_ordinal1(sK0) <- (50) [superposition 127,544] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
555. 51 <=> sK0 = sK1 [avatar definition]
559. sK0 = sK1 <- (3, 50) [subsumption resolution 549,179] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
560. 51 | ~3 | ~50 [avatar split clause 559,542,177,555]
571. sK0 != sK1 | r2_hidden(sK1,sK0) | ~r2_hidden(sK0,sK1) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
572. $false [avatar sat refutation 170,175,180,317,389,424,470,497,504,545,560,571] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:345,thDist:-345}
% SZS output end Proof for ordinal3__t28_ordinal3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 42
% >>> Preprocessing
% Introduced skolems: 11
% Pure predicates: 1
% 
% >>> Saturation
% Initial clauses: 52
% Generated clauses: 236
% Active clauses: 87
% Passive clauses: 120
% Final active clauses: 82
% Final passive clauses: 31
% Main loop iterations started: 89
% 
% >>> Simplifying Inferences
% Duplicate literals: 8
% Fw subsumption resolutions: 16
% Bw demodulations: 4
% 
% >>> Deletion Inferences
% Simple tautologies: 9
% Forward subsumptions: 28
% 
% >>> Generating Inferences
% Binary resolution: 84
% Unit resulting resolution: 1
% Factoring: 1
% Forward superposition: 9
% Backward superposition: 11
% Self superposition: 1
% Equality resolution: 1
% 
% >>> AVATAR
% Split clauses: 3
% Split components: 6
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 62
% SAT solver unit clauses: 16
% SAT solver binary clauses: 32
% 
% Memory used [KB]: 5245
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% other: 0.002 s

