i: [1,20,0,3,1,1,0,0] t16_xtuple_0
a: [3,35,0,3,1,20]
s: 35
i: [1,21,0,5,1,1,0,0] t16_xtuple_0
a: [3,30,0,5,1,21]
s: 30
i: [1,26,0,3,1,0,0,1] reflexivity_r1_tarski
s: 26
i: [1,22,0,6,1,0,0,1] d16_xtuple_0
s: 22
i: [1,23,0,6,1,0,0,2] d14_xtuple_0
s: 23
i: [1,24,0,8,2,0,0,2] t10_xtuple_0
s: 24
d: [2,38,1,7,1,0,44,23,23]
d: [2,39,1,6,1,0,22,38,22]
s: 39
i: [1,25,0,8,2,0,0,2] t8_xtuple_0
s: 25
d: [2,37,1,7,1,0,44,23,22]
s: 37
d: [2,45,1,9,2,0,44,25,22]
s: 45
d: [2,46,1,9,2,0,44,25,23]
s: 46
d: [2,40,2,7,1,0,44,39,22]
s: 40
d: [2,47,1,9,2,0,44,25,22]
s: 47
d: [2,41,2,7,1,0,44,39,23]
s: 41
r: 37
d: [2,70,1,7,1,0,23,37,41]
s: 70
d: [2,48,1,9,2,0,44,25,23]
s: 48
d: [2,44,2,7,1,0,44,22,39]
s: 44
d: [2,52,2,8,1,0,44,23,37]
d: [2,53,2,7,1,0,22,52,22]
s: 53
d: [2,51,2,8,1,0,44,39,37]
s: 51
d: [2,87,2,8,1,0,44,39,70]
s: 87
d: [2,84,2,8,1,0,44,23,70]
d: [2,91,2,8,1,0,22,84,40]
s: 91
d: [2,101,2,8,1,0,22,51,41]
s: 101
d: [2,42,2,9,2,0,44,24,39]
s: 42
d: [2,65,3,8,1,0,44,22,40]
s: 65
d: [2,43,2,9,2,0,44,24,39]
s: 43
d: [2,49,2,10,2,0,44,25,37]
s: 49
d: [2,77,3,8,1,0,44,22,41]
s: 77
d: [2,50,2,10,2,0,44,25,37]
s: 50
d: [2,121,3,9,1,0,44,91,23]
d: [2,136,3,8,1,0,22,121,39]
s: 136
d: [2,54,2,10,2,0,44,45,22]
s: 54
d: [2,55,2,10,2,0,44,45,23]
s: 55
d: [2,173,3,8,2,0,40,54,24]
s: 173
d: [2,200,4,0,0,2,52,35,30,173]
e: 200
d: [2,201,4,3,1,1,40,173,30]
d: [2,203,4,0,0,2,21,201,35]
e: 203
f: 203,20,21
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xtuple_0__t16_xtuple_0
% SZS output start Proof for xtuple_0__t16_xtuple_0
1. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1))) [input t16_xtuple_0]
2. ~! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1))) [negated conjecture 1]
4. ! [X0] : k13_xtuple_0(X0) = k9_xtuple_0(k11_xtuple_0(X0)) [input d16_xtuple_0]
9. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [input t10_xtuple_0]
10. ! [X0] : ! [X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [input t8_xtuple_0]
11. ~! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1))) [flattening 2]
12. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1))) [flattening 9]
13. ! [X0,X1] : (r1_tarski(X0,X1) => r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1))) [flattening 10]
15. ? [X0,X1] : (~r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1)) & r1_tarski(X0,X1)) [ennf transformation 11]
16. ! [X0,X1] : (r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 12]
17. ! [X0,X1] : (r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1)) [ennf transformation 13]
18. ? [X0,X1] : (~r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1)) & r1_tarski(X0,X1)) => (~r1_tarski(k13_xtuple_0(sK0),k13_xtuple_0(sK1)) & r1_tarski(sK0,sK1)) [choice axiom]
19. ~r1_tarski(k13_xtuple_0(sK0),k13_xtuple_0(sK1)) & r1_tarski(sK0,sK1) [skolemisation 15,18]
20. r1_tarski(sK0,sK1) [cnf transformation 19] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
21. ~r1_tarski(k13_xtuple_0(sK0),k13_xtuple_0(sK1)) [cnf transformation 19] {a:0,w:5,wCS:500,goal:1,sine:0,thAx:0,allAx:2,thDist:-2}
22. k13_xtuple_0(X0) = k9_xtuple_0(k11_xtuple_0(X0)) [cnf transformation 4] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
24. r1_tarski(k11_xtuple_0(X0),k11_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 16] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
25. r1_tarski(k9_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(X0,X1) [cnf transformation 17] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
28. 1 <=> r1_tarski(k13_xtuple_0(sK0),k13_xtuple_0(sK1)) [avatar definition]
30. ~r1_tarski(k13_xtuple_0(sK0),k13_xtuple_0(sK1)) <- (~1) [avatar component clause 28] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
31. ~1 [avatar split clause 21,28]
33. 2 <=> r1_tarski(sK0,sK1) [avatar definition]
35. r1_tarski(sK0,sK1) <- (2) [avatar component clause 33] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:2,thDist:-2}
36. 2 [avatar split clause 20,33]
45. r1_tarski(k13_xtuple_0(X0),k9_xtuple_0(X1)) | ~r1_tarski(k11_xtuple_0(X0),X1) [superposition 25,22] {a:1,w:9,wCS:3600,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
54. ~r1_tarski(k11_xtuple_0(X1),k11_xtuple_0(X0)) | r1_tarski(k13_xtuple_0(X1),k13_xtuple_0(X0)) [superposition 45,22] {a:2,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:3,thDist:-3}
173. r1_tarski(k13_xtuple_0(X0),k13_xtuple_0(X1)) | ~r1_tarski(X0,X1) [resolution 54,24] {a:3,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:4,thDist:-4}
201. ~r1_tarski(sK0,sK1) <- (~1) [resolution 173,30] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
203. $false <- (~1, 2) [subsumption resolution 201,35] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
204. 1 | ~2 [avatar contradiction clause 203]
205. $false [avatar sat refutation 31,36,204] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:20,thDist:-20}
% SZS output end Proof for xtuple_0__t16_xtuple_0
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 9
% Introduced skolems: 2
% 
% >>> Saturation
% Initial clauses: 7
% Generated clauses: 175
% Active clauses: 24
% Passive clauses: 113
% Final active clauses: 23
% Final passive clauses: 81
% Main loop iterations started: 32
% 
% >>> Simplifying Inferences
% Fw subsumption resolutions: 1
% Fw demodulations: 43
% Bw demodulations: 1
% 
% >>> Deletion Inferences
% Forward subsumptions: 13
% Fw demodulations to eq. taut.: 9
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 63
% Backward superposition: 50
% Self superposition: 3
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 6
% SAT solver unit clauses: 2
% SAT solver binary clauses: 2
% 
% Memory used [KB]: 9466
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% passive container maintenance: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

