i: [1,24,0,2,1,1,0,0] t60_xxreal_3
a: [3,59,0,2,1,24]
s: 59
i: [1,23,0,2,1,1,0,0] t60_xxreal_3
a: [3,64,0,2,1,23]
s: 64
i: [1,28,0,8,2,1,0,0] t60_xxreal_3
a: [3,39,0,4,1,28]
s: 39
i: [1,27,0,8,2,1,0,0] t60_xxreal_3
a: [3,48,0,4,1,27]
s: 48
i: [1,33,0,7,3,0,0,1] reflexivity_r1_xxreal_0
a: [3,68,0,2,1,33]
s: 68
d: [2,80,1,0,0,2,40,68,64]
e: 80
d: [2,79,1,0,0,2,40,68,59]
e: 79
d: [2,78,1,0,0,2,52,59,68]
e: 78
r: 68
i: [1,35,0,2,1,0,0,1] rc1_xxreal_0
a: [3,76,0,2,1,35]
s: 76
i: [1,29,0,5,2,0,0,1] dt_k2_xxreal_3
s: 29
a: [3,71,0,5,2,33]
s: 71
i: [1,30,0,7,2,0,0,255] involutiveness_k2_xxreal_3
s: 30
i: [1,34,0,10,4,0,0,1] connectedness_r1_xxreal_0
s: 34
i: [1,31,0,12,4,0,0,1] t38_xxreal_3
s: 31
d: [2,85,1,9,3,1,40,34,39]
d: [2,88,1,7,2,2,21,85,59]
a: [3,92,1,3,1,88]
s: 92
d: [2,98,2,0,0,2,52,64,92,29]
e: 98
d: [2,99,2,2,1,1,40,92,29]
d: [2,101,2,0,0,2,21,99,64]
e: 101
r: 92
i: [1,26,0,8,2,1,0,0] t60_xxreal_3
a: [3,53,0,4,1,26]
s: 53
a: [3,91,1,3,1,88]
s: 91
i: [1,32,0,12,4,0,0,1] t38_xxreal_3
s: 32
d: [2,94,1,15,5,0,44,31,30]
d: [2,96,1,12,4,0,21,94,29]
s: 96
d: [2,114,2,0,0,4,52,59,64,39,48,96]
e: 114
d: [2,115,2,8,3,1,40,96,48]
d: [2,124,2,4,2,2,21,115,39]
d: [2,125,2,2,1,3,21,124,64]
d: [2,126,2,0,0,4,21,125,59]
e: 126
r: 39
a: [3,38,0,4,1,28]
s: 38
a: [3,43,0,4,1,28]
s: 43
d: [2,116,2,13,4,1,40,96,71]
d: [2,122,2,10,3,1,18,116]
d: [2,128,2,7,2,1,21,122,29]
s: 128
d: [2,95,1,15,5,0,44,31,30]
d: [2,97,1,12,4,0,21,95,29]
s: 97
d: [2,143,2,0,0,4,52,64,59,43,53,97]
e: 143
d: [2,145,2,8,3,1,40,97,53]
d: [2,155,2,4,2,2,21,145,43]
d: [2,156,2,2,1,3,21,155,59]
d: [2,157,2,0,0,4,21,156,64]
e: 157
d: [2,131,2,6,2,2,21,115,64]
d: [2,132,2,4,1,3,21,131,59]
f: 157,132,101,88,23,24,26,27,28
% Refutation found. Thanks to Tanya!
% SZS status Theorem for xxreal_3__t60_xxreal_3
% SZS output start Proof for xxreal_3__t60_xxreal_3
1. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(k2_xxreal_3(X0),X1) => r1_xxreal_0(k2_xxreal_3(X1),X0)) & (r1_xxreal_0(X0,k2_xxreal_3(X1)) => r1_xxreal_0(X1,k2_xxreal_3(X0)))))) [input t60_xxreal_3]
2. ~! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => ((r1_xxreal_0(k2_xxreal_3(X0),X1) => r1_xxreal_0(k2_xxreal_3(X1),X0)) & (r1_xxreal_0(X0,k2_xxreal_3(X1)) => r1_xxreal_0(X1,k2_xxreal_3(X0)))))) [negated conjecture 1]
3. ! [X0,X1] : ((v1_xxreal_0(X1) & v1_xxreal_0(X0)) => (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1))) [input connectedness_r1_xxreal_0]
4. ! [X0] : (v1_xxreal_0(X0) => v1_xxreal_0(k2_xxreal_3(X0))) [input dt_k2_xxreal_3]
5. ! [X0] : (v1_xxreal_0(X0) => k2_xxreal_3(k2_xxreal_3(X0)) = X0) [input involutiveness_k2_xxreal_3]
8. ! [X0] : (v1_xxreal_0(X0) => ! [X1] : (v1_xxreal_0(X1) => (r1_xxreal_0(X0,X1) <=> r1_xxreal_0(k2_xxreal_3(X1),k2_xxreal_3(X0))))) [input t38_xxreal_3]
9. ? [X0] : (? [X1] : (((~r1_xxreal_0(k2_xxreal_3(X1),X0) & r1_xxreal_0(k2_xxreal_3(X0),X1)) | (~r1_xxreal_0(X1,k2_xxreal_3(X0)) & r1_xxreal_0(X0,k2_xxreal_3(X1)))) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) [ennf transformation 2]
10. ! [X0] : (v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X0)) [ennf transformation 4]
11. ! [X0] : (k2_xxreal_3(k2_xxreal_3(X0)) = X0 | ~v1_xxreal_0(X0)) [ennf transformation 5]
12. ! [X0] : (! [X1] : ((r1_xxreal_0(X0,X1) <=> r1_xxreal_0(k2_xxreal_3(X1),k2_xxreal_3(X0))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [ennf transformation 8]
15. ! [X0,X1] : ((r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1)) | (~v1_xxreal_0(X1) | ~v1_xxreal_0(X0))) [ennf transformation 3]
16. ! [X0,X1] : (r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0)) [flattening 15]
17. ? [X0] : (? [X1] : (((~r1_xxreal_0(k2_xxreal_3(X1),X0) & r1_xxreal_0(k2_xxreal_3(X0),X1)) | (~r1_xxreal_0(X1,k2_xxreal_3(X0)) & r1_xxreal_0(X0,k2_xxreal_3(X1)))) & v1_xxreal_0(X1)) & v1_xxreal_0(X0)) => (? [X1] : (((~r1_xxreal_0(k2_xxreal_3(X1),sK0) & r1_xxreal_0(k2_xxreal_3(sK0),X1)) | (~r1_xxreal_0(X1,k2_xxreal_3(sK0)) & r1_xxreal_0(sK0,k2_xxreal_3(X1)))) & v1_xxreal_0(X1)) & v1_xxreal_0(sK0)) [choice axiom]
18. ? [X1] : (((~r1_xxreal_0(k2_xxreal_3(X1),sK0) & r1_xxreal_0(k2_xxreal_3(sK0),X1)) | (~r1_xxreal_0(X1,k2_xxreal_3(sK0)) & r1_xxreal_0(sK0,k2_xxreal_3(X1)))) & v1_xxreal_0(X1)) => (((~r1_xxreal_0(k2_xxreal_3(sK1),sK0) & r1_xxreal_0(k2_xxreal_3(sK0),sK1)) | (~r1_xxreal_0(sK1,k2_xxreal_3(sK0)) & r1_xxreal_0(sK0,k2_xxreal_3(sK1)))) & v1_xxreal_0(sK1)) [choice axiom]
19. (((~r1_xxreal_0(k2_xxreal_3(sK1),sK0) & r1_xxreal_0(k2_xxreal_3(sK0),sK1)) | (~r1_xxreal_0(sK1,k2_xxreal_3(sK0)) & r1_xxreal_0(sK0,k2_xxreal_3(sK1)))) & v1_xxreal_0(sK1)) & v1_xxreal_0(sK0) [skolemisation 9,18,17]
20. ! [X0] : (! [X1] : (((r1_xxreal_0(X0,X1) | ~r1_xxreal_0(k2_xxreal_3(X1),k2_xxreal_3(X0))) & (r1_xxreal_0(k2_xxreal_3(X1),k2_xxreal_3(X0)) | ~r1_xxreal_0(X0,X1))) | ~v1_xxreal_0(X1)) | ~v1_xxreal_0(X0)) [nnf transformation 12]
23. v1_xxreal_0(sK0) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
24. v1_xxreal_0(sK1) [cnf transformation 19] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
26. r1_xxreal_0(k2_xxreal_3(sK0),sK1) | ~r1_xxreal_0(sK1,k2_xxreal_3(sK0)) [cnf transformation 19] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
27. ~r1_xxreal_0(k2_xxreal_3(sK1),sK0) | r1_xxreal_0(sK0,k2_xxreal_3(sK1)) [cnf transformation 19] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
28. ~r1_xxreal_0(k2_xxreal_3(sK1),sK0) | ~r1_xxreal_0(sK1,k2_xxreal_3(sK0)) [cnf transformation 19] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
29. v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X0) [cnf transformation 10] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
30. k2_xxreal_3(k2_xxreal_3(X0)) = X0 | ~v1_xxreal_0(X0) [cnf transformation 11] {a:0,w:7,wCS:2800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
31. r1_xxreal_0(k2_xxreal_3(X1),k2_xxreal_3(X0)) | ~r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 20] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
34. r1_xxreal_0(X1,X0) | r1_xxreal_0(X0,X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [cnf transformation 16] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
37. 1 <=> r1_xxreal_0(sK1,k2_xxreal_3(sK0)) [avatar definition]
39. ~r1_xxreal_0(sK1,k2_xxreal_3(sK0)) <- (~1) [avatar component clause 37] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
41. 2 <=> r1_xxreal_0(k2_xxreal_3(sK1),sK0) [avatar definition]
43. ~r1_xxreal_0(k2_xxreal_3(sK1),sK0) <- (~2) [avatar component clause 41] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
44. ~1 | ~2 [avatar split clause 28,41,37]
46. 3 <=> r1_xxreal_0(sK0,k2_xxreal_3(sK1)) [avatar definition]
48. r1_xxreal_0(sK0,k2_xxreal_3(sK1)) <- (3) [avatar component clause 46] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
49. 3 | ~2 [avatar split clause 27,41,46]
51. 4 <=> r1_xxreal_0(k2_xxreal_3(sK0),sK1) [avatar definition]
53. r1_xxreal_0(k2_xxreal_3(sK0),sK1) <- (4) [avatar component clause 51] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
54. ~1 | 4 [avatar split clause 26,51,37]
57. 5 <=> v1_xxreal_0(sK1) [avatar definition]
59. v1_xxreal_0(sK1) <- (5) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 5 [avatar split clause 24,57]
62. 6 <=> v1_xxreal_0(sK0) [avatar definition]
64. v1_xxreal_0(sK0) <- (6) [avatar component clause 62] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
65. 6 [avatar split clause 23,62]
85. r1_xxreal_0(k2_xxreal_3(sK0),sK1) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(k2_xxreal_3(sK0)) <- (~1) [resolution 34,39] {a:1,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
88. r1_xxreal_0(k2_xxreal_3(sK0),sK1) | ~v1_xxreal_0(k2_xxreal_3(sK0)) <- (~1, 5) [subsumption resolution 85,59] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
90. 10 <=> v1_xxreal_0(k2_xxreal_3(sK0)) [avatar definition]
92. ~v1_xxreal_0(k2_xxreal_3(sK0)) <- (~10) [avatar component clause 90] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
93. ~10 | 4 | 1 | ~5 [avatar split clause 88,57,37,51,90]
94. r1_xxreal_0(X0,k2_xxreal_3(X1)) | ~r1_xxreal_0(X1,k2_xxreal_3(X0)) | ~v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [superposition 31,30] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
95. r1_xxreal_0(k2_xxreal_3(X1),X0) | ~r1_xxreal_0(k2_xxreal_3(X0),X1) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(k2_xxreal_3(X0)) | ~v1_xxreal_0(X0) [superposition 31,30] {a:1,w:15,wCS:6000,sine:1,thAx:0,allAx:2,thDist:-2}
96. ~r1_xxreal_0(X1,k2_xxreal_3(X0)) | r1_xxreal_0(X0,k2_xxreal_3(X1)) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [subsumption resolution 94,29] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
97. ~r1_xxreal_0(k2_xxreal_3(X0),X1) | r1_xxreal_0(k2_xxreal_3(X1),X0) | ~v1_xxreal_0(X1) | ~v1_xxreal_0(X0) [subsumption resolution 95,29] {a:1,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
99. ~v1_xxreal_0(sK0) <- (~10) [resolution 92,29] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
101. $false <- (6, ~10) [subsumption resolution 99,64] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
102. ~6 | 10 [avatar contradiction clause 101]
115. r1_xxreal_0(sK1,k2_xxreal_3(sK0)) | ~v1_xxreal_0(sK0) | ~v1_xxreal_0(sK1) <- (3) [resolution 96,48] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
131. r1_xxreal_0(sK1,k2_xxreal_3(sK0)) | ~v1_xxreal_0(sK1) <- (3, 6) [subsumption resolution 115,64] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
132. r1_xxreal_0(sK1,k2_xxreal_3(sK0)) <- (3, 5, 6) [subsumption resolution 131,59] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
133. 1 | ~3 | ~5 | ~6 [avatar split clause 132,62,57,46,37]
145. r1_xxreal_0(k2_xxreal_3(sK1),sK0) | ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (4) [resolution 97,53] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
155. ~v1_xxreal_0(sK1) | ~v1_xxreal_0(sK0) <- (~2, 4) [subsumption resolution 145,43] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
156. ~v1_xxreal_0(sK0) <- (~2, 4, 5) [subsumption resolution 155,59] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
157. $false <- (~2, 4, 5, 6) [subsumption resolution 156,64] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
158. 2 | ~4 | ~5 | ~6 [avatar contradiction clause 157]
161. $false [avatar sat refutation 44,49,54,60,65,93,102,133,158] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:114,thDist:-114}
% SZS output end Proof for xxreal_3__t60_xxreal_3
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 7
% Introduced skolems: 3
% 
% >>> Saturation
% Initial clauses: 13
% Generated clauses: 97
% Active clauses: 20
% Passive clauses: 26
% Final active clauses: 17
% Final passive clauses: 6
% Main loop iterations started: 20
% 
% >>> Simplifying Inferences
% Duplicate literals: 12
% Fw subsumption resolutions: 23
% 
% >>> Deletion Inferences
% Simple tautologies: 2
% Forward subsumptions: 9
% 
% >>> Generating Inferences
% Binary resolution: 20
% Unit resulting resolution: 4
% Factoring: 1
% Forward superposition: 8
% Backward superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 21
% SAT solver unit clauses: 3
% SAT solver binary clauses: 10
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% forward subsumption: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

