i: [1,30,0,2,1,1,0,0] t166_relat_1
a: [3,54,0,2,1,30]
s: 54
i: [1,29,0,2,1,1,0,0] t166_relat_1
a: [3,59,0,2,1,29]
s: 59
i: [1,32,0,7,1,1,0,0] t166_relat_1
a: [3,44,0,7,1,32]
s: 44
i: [1,31,0,7,1,1,0,0] t166_relat_1
a: [3,49,0,7,1,31]
s: 49
i: [1,38,0,3,1,0,0,2] reflexivity_r1_tarski
s: 38
i: [1,37,0,4,1,0,0,255] existence_m1_subset_1
s: 37
i: [1,39,0,6,2,0,0,1] dt_k5_relat_1
s: 39
i: [1,35,0,7,2,0,0,255] t3_subset
s: 35
d: [2,61,1,6,2,1,44,39,49]
d: [2,62,1,4,1,2,21,61,54]
a: [3,66,1,4,1,62]
s: 66
i: [1,36,0,7,2,0,0,255] t3_subset
s: 36
i: [1,33,0,8,3,0,0,255] cc2_relat_1
s: 33
d: [2,68,1,5,1,0,40,35,37]
s: 68
i: [1,40,0,11,2,0,0,1] rd5_relat_1
s: 40
d: [2,72,1,11,2,1,44,40,49]
d: [2,75,1,9,1,2,21,72,54]
a: [3,79,1,9,1,75]
s: 79
i: [1,34,0,14,3,0,0,1] t129_relat_1
s: 34
d: [2,83,1,14,3,1,44,34,49]
d: [2,84,1,12,2,2,21,83,54]
s: 84
d: [2,85,2,15,4,2,44,84,34]
d: [2,88,2,12,3,2,18,85]
d: [2,89,2,10,2,3,21,88,59]
s: 89
d: [2,91,3,0,0,4,52,38,44,89]
e: 91
d: [2,92,3,10,2,4,44,44,89]
d: [2,93,3,3,1,4,19,92]
d: [2,95,3,0,0,4,21,93,38]
e: 95
f: 95,29,30,31,32
% Refutation found. Thanks to Tanya!
% SZS status Theorem for relat_1__t166_relat_1
% SZS output start Proof for relat_1__t166_relat_1
1. ! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (k5_relat_1(X1,X0) = k5_relat_1(X2,X0) => k7_relat_1(X1,X0) = k7_relat_1(X2,X0)))) [input t166_relat_1]
2. ~! [X0] : ! [X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (k5_relat_1(X1,X0) = k5_relat_1(X2,X0) => k7_relat_1(X1,X0) = k7_relat_1(X2,X0)))) [negated conjecture 1]
10. ! [X0,X1] : r1_tarski(X0,X0) [input reflexivity_r1_tarski]
11. ! [X0] : (v1_relat_1(X0) => ! [X1] : ! [X2] : (r1_tarski(X1,X2) => k7_relat_1(k5_relat_1(X0,X2),X1) = k7_relat_1(X0,X1))) [input t129_relat_1]
13. ~! [X0,X1] : (v1_relat_1(X1) => ! [X2] : (v1_relat_1(X2) => (k5_relat_1(X1,X0) = k5_relat_1(X2,X0) => k7_relat_1(X1,X0) = k7_relat_1(X2,X0)))) [flattening 2]
14. ! [X0] : (v1_relat_1(X0) => ! [X1,X2] : (r1_tarski(X1,X2) => k7_relat_1(k5_relat_1(X0,X2),X1) = k7_relat_1(X0,X1))) [flattening 11]
16. ! [X0] : r1_tarski(X0,X0) [rectify 10]
17. ? [X0,X1] : (? [X2] : ((k7_relat_1(X1,X0) != k7_relat_1(X2,X0) & k5_relat_1(X1,X0) = k5_relat_1(X2,X0)) & v1_relat_1(X2)) & v1_relat_1(X1)) [ennf transformation 13]
18. ? [X0,X1] : (? [X2] : (k7_relat_1(X1,X0) != k7_relat_1(X2,X0) & k5_relat_1(X1,X0) = k5_relat_1(X2,X0) & v1_relat_1(X2)) & v1_relat_1(X1)) [flattening 17]
20. ! [X0] : (! [X1,X2] : (k7_relat_1(k5_relat_1(X0,X2),X1) = k7_relat_1(X0,X1) | ~r1_tarski(X1,X2)) | ~v1_relat_1(X0)) [ennf transformation 14]
23. ? [X0,X1] : (? [X2] : (k7_relat_1(X1,X0) != k7_relat_1(X2,X0) & k5_relat_1(X1,X0) = k5_relat_1(X2,X0) & v1_relat_1(X2)) & v1_relat_1(X1)) => (? [X2] : (k7_relat_1(sK1,sK0) != k7_relat_1(X2,sK0) & k5_relat_1(sK1,sK0) = k5_relat_1(X2,sK0) & v1_relat_1(X2)) & v1_relat_1(sK1)) [choice axiom]
24. ? [X2] : (k7_relat_1(sK1,sK0) != k7_relat_1(X2,sK0) & k5_relat_1(sK1,sK0) = k5_relat_1(X2,sK0) & v1_relat_1(X2)) => (k7_relat_1(sK1,sK0) != k7_relat_1(sK2,sK0) & k5_relat_1(sK1,sK0) = k5_relat_1(sK2,sK0) & v1_relat_1(sK2)) [choice axiom]
25. (k7_relat_1(sK1,sK0) != k7_relat_1(sK2,sK0) & k5_relat_1(sK1,sK0) = k5_relat_1(sK2,sK0) & v1_relat_1(sK2)) & v1_relat_1(sK1) [skolemisation 18,24,23]
29. v1_relat_1(sK1) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
30. v1_relat_1(sK2) [cnf transformation 25] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
31. k5_relat_1(sK1,sK0) = k5_relat_1(sK2,sK0) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
32. k7_relat_1(sK1,sK0) != k7_relat_1(sK2,sK0) [cnf transformation 25] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
34. k7_relat_1(k5_relat_1(X0,X2),X1) = k7_relat_1(X0,X1) | ~r1_tarski(X1,X2) | ~v1_relat_1(X0) [cnf transformation 20] {a:0,w:14,wCS:5600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
38. r1_tarski(X0,X0) [cnf transformation 16] {a:0,w:3,wCS:1200,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
42. 1 <=> k7_relat_1(sK1,sK0) = k7_relat_1(sK2,sK0) [avatar definition]
44. k7_relat_1(sK1,sK0) != k7_relat_1(sK2,sK0) <- (~1) [avatar component clause 42] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
45. ~1 [avatar split clause 32,42]
47. 2 <=> k5_relat_1(sK1,sK0) = k5_relat_1(sK2,sK0) [avatar definition]
49. k5_relat_1(sK1,sK0) = k5_relat_1(sK2,sK0) <- (2) [avatar component clause 47] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
50. 2 [avatar split clause 31,47]
52. 3 <=> v1_relat_1(sK2) [avatar definition]
54. v1_relat_1(sK2) <- (3) [avatar component clause 52] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
55. 3 [avatar split clause 30,52]
57. 4 <=> v1_relat_1(sK1) [avatar definition]
59. v1_relat_1(sK1) <- (4) [avatar component clause 57] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
60. 4 [avatar split clause 29,57]
83. k7_relat_1(sK2,X0) = k7_relat_1(k5_relat_1(sK1,sK0),X0) | ~r1_tarski(X0,sK0) | ~v1_relat_1(sK2) <- (2) [superposition 34,49] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
84. k7_relat_1(sK2,X0) = k7_relat_1(k5_relat_1(sK1,sK0),X0) | ~r1_tarski(X0,sK0) <- (2, 3) [subsumption resolution 83,54] {a:1,w:12,wCS:1200,nSel:1,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
85. k7_relat_1(sK2,X0) = k7_relat_1(sK1,X0) | ~r1_tarski(X0,sK0) | ~r1_tarski(X0,sK0) | ~v1_relat_1(sK1) <- (2, 3) [superposition 84,34] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
88. k7_relat_1(sK2,X0) = k7_relat_1(sK1,X0) | ~r1_tarski(X0,sK0) | ~v1_relat_1(sK1) <- (2, 3) [duplicate literal removal 85] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
89. k7_relat_1(sK2,X0) = k7_relat_1(sK1,X0) | ~r1_tarski(X0,sK0) <- (2, 3, 4) [subsumption resolution 88,59] {a:2,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
92. k7_relat_1(sK1,sK0) != k7_relat_1(sK1,sK0) | ~r1_tarski(sK0,sK0) <- (~1, 2, 3, 4) [superposition 44,89] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
93. ~r1_tarski(sK0,sK0) <- (~1, 2, 3, 4) [trivial inequality removal 92] {a:3,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
95. $false <- (~1, 2, 3, 4) [subsumption resolution 93,38] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
96. 1 | ~2 | ~3 | ~4 [avatar contradiction clause 95]
97. $false [avatar sat refutation 45,50,55,60,96] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:47,thDist:-47}
% SZS output end Proof for relat_1__t166_relat_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11
% Introduced skolems: 4
% 
% >>> Saturation
% Initial clauses: 12
% Generated clauses: 42
% Active clauses: 17
% Passive clauses: 20
% Final active clauses: 17
% Final passive clauses: 3
% Main loop iterations started: 17
% 
% >>> Simplifying Inferences
% Duplicate literals: 2
% Trivial inequalities: 1
% Fw subsumption resolutions: 6
% 
% >>> Deletion Inferences
% Simple tautologies: 3
% Forward subsumptions: 2
% 
% >>> Generating Inferences
% Binary resolution: 4
% Unit resulting resolution: 1
% Forward superposition: 4
% Backward superposition: 5
% Self superposition: 1
% 
% >>> AVATAR
% Sat splitting refutations: 2
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 10
% SAT solver unit clauses: 4
% 
% Memory used [KB]: 9338
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.001 s ( own 0.0000 s ) 

