i: [1,267,0,2,1,1,0,0] t23_card_1
a: [3,425,0,2,1,267]
s: 425
i: [1,266,0,2,1,1,0,0] t23_card_1
a: [3,430,0,2,1,266]
s: 430
i: [1,269,0,8,2,1,0,0] t23_card_1
a: [3,415,0,3,1,269]
s: 415
a: [3,418,0,5,1,269]
s: 418
i: [1,270,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,435,0,2,1,270]
s: 435
i: [1,271,0,2,1,0,0,255] fc1_xboole_0
a: [3,440,0,2,1,271]
s: 440
i: [1,366,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,449,0,2,1,366]
s: 449
d: [2,666,1,0,0,2,40,449,430]
e: 666
d: [2,665,1,0,0,2,40,449,425]
e: 665
d: [2,664,1,0,0,2,52,425,449]
e: 664
r: 449
i: [1,370,0,2,1,0,0,255] rc2_zfmisc_1
a: [3,457,0,2,1,370]
s: 457
i: [1,371,0,2,1,0,0,2] rc2_xboole_0
a: [3,462,0,2,1,371]
s: 462
i: [1,372,0,2,1,0,0,255] rc4_ordinal1
a: [3,467,0,2,1,372]
s: 467
i: [1,373,0,2,1,0,0,2] rc1_card_1
a: [3,472,0,2,1,373]
s: 472
i: [1,374,0,2,1,0,0,1] rc1_ordinal1
a: [3,477,0,2,1,374]
s: 477
i: [1,375,0,2,1,0,0,2] rc1_xboole_0
a: [3,482,0,2,1,375]
s: 482
i: [1,379,0,2,1,0,0,255] rc2_ordinal1
a: [3,487,0,2,1,379]
s: 487
i: [1,378,0,2,1,0,0,255] rc2_ordinal1
a: [3,492,0,2,1,378]
s: 492
i: [1,377,0,2,1,0,0,255] rc2_ordinal1
a: [3,497,0,2,1,377]
s: 497
i: [1,376,0,2,1,0,0,255] rc2_ordinal1
a: [3,502,0,2,1,376]
s: 502
i: [1,381,0,2,1,0,0,255] rc5_ordinal1
a: [3,507,0,2,1,381]
s: 507
i: [1,380,0,2,1,0,0,255] rc5_ordinal1
a: [3,512,0,2,1,380]
s: 512
i: [1,383,0,2,1,0,0,255] rc7_funct_1
a: [3,517,0,2,1,383]
s: 517
i: [1,382,0,2,1,0,0,255] rc7_funct_1
a: [3,522,0,2,1,382]
s: 522
i: [1,385,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,527,0,2,1,385]
s: 527
i: [1,384,0,2,1,0,0,255] rc1_zfmisc_1
a: [3,532,0,2,1,384]
s: 532
i: [1,387,0,2,1,0,0,255] rc1_finset_1
a: [3,537,0,2,1,387]
s: 537
i: [1,386,0,2,1,0,0,255] rc1_finset_1
a: [3,542,0,2,1,386]
s: 542
i: [1,390,0,2,1,0,0,255] rc8_finset_1
a: [3,547,0,2,1,390]
s: 547
i: [1,389,0,2,1,0,0,255] rc8_finset_1
a: [3,552,0,2,1,389]
s: 552
i: [1,388,0,2,1,0,0,255] rc8_finset_1
a: [3,557,0,2,1,388]
s: 557
i: [1,393,0,2,1,0,0,255] rc9_finset_1
a: [3,562,0,2,1,393]
s: 562
i: [1,392,0,2,1,0,0,255] rc9_finset_1
a: [3,567,0,2,1,392]
s: 567
i: [1,391,0,2,1,0,0,255] rc9_finset_1
a: [3,572,0,2,1,391]
s: 572
i: [1,397,0,2,1,0,0,255] rc3_finset_1
a: [3,577,0,2,1,397]
s: 577
i: [1,396,0,2,1,0,0,255] rc3_finset_1
a: [3,582,0,2,1,396]
s: 582
i: [1,395,0,2,1,0,0,255] rc3_finset_1
a: [3,587,0,2,1,395]
s: 587
i: [1,394,0,2,1,0,0,255] rc3_finset_1
a: [3,592,0,2,1,394]
s: 592
i: [1,399,0,2,1,0,0,255] rc1_funct_1
a: [3,597,0,2,1,399]
s: 597
i: [1,398,0,2,1,0,0,255] rc1_funct_1
a: [3,602,0,2,1,398]
s: 602
i: [1,402,0,2,1,0,0,255] rc5_funct_1
a: [3,607,0,2,1,402]
s: 607
i: [1,401,0,2,1,0,0,255] rc5_funct_1
a: [3,612,0,2,1,401]
s: 612
i: [1,400,0,2,1,0,0,255] rc5_funct_1
a: [3,617,0,2,1,400]
s: 617
i: [1,404,0,2,1,0,0,255] rc2_funct_1
a: [3,622,0,2,1,404]
s: 622
i: [1,403,0,2,1,0,0,255] rc2_funct_1
a: [3,627,0,2,1,403]
s: 627
i: [1,407,0,2,1,0,0,255] rc1_funcop_1
a: [3,632,0,2,1,407]
s: 632
i: [1,406,0,2,1,0,0,255] rc1_funcop_1
a: [3,637,0,2,1,406]
s: 637
i: [1,405,0,2,1,0,0,255] rc1_funcop_1
a: [3,642,0,2,1,405]
s: 642
i: [1,411,0,2,1,0,0,255] rc2_funcop_1
a: [3,647,0,2,1,411]
s: 647
i: [1,410,0,2,1,0,0,255] rc2_funcop_1
a: [3,652,0,2,1,410]
s: 652
i: [1,409,0,2,1,0,0,255] rc2_funcop_1
a: [3,657,0,2,1,409]
s: 657
i: [1,408,0,2,1,0,0,255] rc2_funcop_1
a: [3,662,0,2,1,408]
s: 662
i: [1,273,0,3,1,0,0,3] fc1_subset_1
s: 273
i: [1,359,0,3,1,0,0,3] rc2_subset_1
s: 359
i: [1,362,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 362
i: [1,363,0,3,1,0,0,2] reflexivity_r1_tarski
s: 363
i: [1,272,0,3,1,0,0,255] d2_xboole_0
a: [3,445,0,3,1,272]
s: 445
r: 435
i: [1,274,0,4,2,0,0,255] cc2_zfmisc_1
s: 274
i: [1,280,0,4,2,0,0,255] cc5_finset_1
s: 280
d: [2,671,1,2,1,1,40,274,440]
a: [3,677,1,2,1,671]
s: 677
i: [1,295,0,4,2,0,0,255] cc6_ordinal1
s: 295
d: [2,673,1,2,1,1,40,274,482]
a: [3,682,1,2,1,673]
s: 682
i: [1,296,0,4,2,0,0,2] cc1_card_1
s: 296
i: [1,302,0,4,2,0,0,255] cc4_finset_1
s: 302
d: [2,684,1,2,1,1,40,280,527]
a: [3,688,1,2,1,684]
s: 688
i: [1,305,0,4,2,0,0,255] cc1_ordinal1
s: 305
d: [2,720,1,2,1,1,40,305,430]
a: [3,727,1,2,1,720]
s: 727
i: [1,306,0,4,2,0,0,255] cc1_ordinal1
s: 306
i: [1,314,0,4,2,0,0,255] cc9_ordinal1
s: 314
d: [2,721,1,2,1,1,40,305,425]
a: [3,732,1,2,1,721]
s: 732
i: [1,315,0,4,2,0,0,2] cc2_card_1
s: 315
d: [2,739,1,2,1,1,40,306,430]
a: [3,746,1,2,1,739]
s: 746
i: [1,316,0,4,2,0,0,255] cc7_ordinal1
s: 316
i: [1,317,0,4,2,0,0,255] cc7_funct_1
s: 317
d: [2,740,1,2,1,1,40,306,425]
a: [3,751,1,2,1,740]
s: 751
i: [1,318,0,4,2,0,0,255] cc6_finset_1
s: 318
d: [2,696,1,2,1,1,40,295,467]
a: [3,701,1,2,1,696]
s: 701
i: [1,319,0,4,2,0,0,255] cc1_zfmisc_1
s: 319
i: [1,320,0,4,2,0,0,1] cc3_ordinal1
s: 320
d: [2,697,1,2,1,1,40,295,507]
a: [3,706,1,2,1,697]
s: 706
i: [1,321,0,4,2,0,0,255] cc1_finset_1
s: 321
d: [2,714,1,2,1,1,40,296,472]
a: [3,718,1,2,1,714]
s: 718
i: [1,322,0,4,2,0,0,255] cc1_funct_1
s: 322
i: [1,357,0,4,1,0,0,2] existence_m1_subset_1
s: 357
d: [2,722,1,2,1,1,40,305,477]
a: [3,737,1,2,1,722]
s: 737
i: [1,361,0,4,1,0,0,255] rc3_subset_1
s: 361
d: [2,741,1,2,1,1,40,306,477]
a: [3,756,1,2,1,741]
s: 756
i: [1,276,0,5,2,0,0,255] rc6_subset_1
s: 276
i: [1,278,0,5,2,0,0,255] rc7_finset_1
s: 278
d: [2,758,1,2,1,1,40,314,440]
a: [3,764,1,2,1,758]
s: 764
i: [1,279,0,5,2,0,0,255] rc7_finset_1
s: 279
d: [2,760,1,2,1,1,40,314,482]
a: [3,769,1,2,1,760]
s: 769
i: [1,284,0,5,2,0,0,3] rc1_subset_1
s: 284
i: [1,286,0,5,2,0,0,255] rc5_subset_1
s: 286
d: [2,771,1,2,1,1,40,315,440]
a: [3,777,1,2,1,771]
s: 777
i: [1,287,0,5,2,0,0,255] rc5_subset_1
s: 287
d: [2,773,1,2,1,1,40,315,482]
a: [3,782,1,2,1,773]
s: 782
i: [1,289,0,5,2,0,0,255] rc2_finset_1
s: 289
i: [1,290,0,5,2,0,0,255] rc2_finset_1
s: 290
d: [2,784,1,2,1,1,40,316,440]
a: [3,790,1,2,1,784]
s: 790
i: [1,304,0,5,2,0,0,1] fc3_card_1
s: 304
d: [2,786,1,2,1,1,40,316,482]
a: [3,795,1,2,1,786]
s: 795
i: [1,311,0,5,2,0,0,255] fc31_finset_1
s: 311
i: [1,312,0,5,2,0,0,255] fc17_finset_1
s: 312
d: [2,797,1,2,1,1,40,317,440]
a: [3,803,1,2,1,797]
s: 803
i: [1,323,0,5,2,0,0,255] t6_boole
s: 323
d: [2,799,1,2,1,1,40,317,482]
a: [3,808,1,2,1,799]
s: 808
i: [1,354,0,5,2,0,0,2] t7_boole
s: 354
i: [1,358,0,5,1,0,0,3] rc2_subset_1
s: 358
d: [2,810,1,2,1,1,40,318,440]
a: [3,816,1,2,1,810]
s: 816
i: [1,360,0,5,1,0,0,255] rc3_subset_1
s: 360
d: [2,812,1,2,1,1,40,318,482]
a: [3,821,1,2,1,812]
s: 821
a: [3,452,0,5,2,366]
s: 452
i: [1,292,0,6,2,0,0,255] rc4_subset_1
s: 292
d: [2,835,1,2,1,1,40,320,440]
a: [3,841,1,2,1,835]
s: 841
i: [1,326,0,6,3,0,0,255] cc2_ordinal1
s: 326
d: [2,837,1,2,1,1,40,320,482]
a: [3,846,1,2,1,837]
s: 846
i: [1,346,0,6,2,0,0,2] d2_ordinal1
s: 346
i: [1,347,0,6,2,0,0,2] d2_ordinal1
s: 347
d: [2,877,1,2,1,1,40,322,440]
a: [3,883,1,2,1,877]
s: 883
i: [1,348,0,6,2,0,0,2] t1_subset
s: 348
d: [2,879,1,2,1,1,40,322,482]
a: [3,888,1,2,1,879]
s: 888
i: [1,364,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 364
i: [1,365,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 365
d: [2,690,2,2,1,1,40,677,280]
a: [3,694,2,2,1,690]
s: 694
i: [1,275,0,7,2,0,0,255] rc6_subset_1
s: 275
d: [2,708,2,2,1,1,40,682,280]
a: [3,712,2,2,1,708]
s: 712
i: [1,277,0,7,2,0,0,255] rc7_finset_1
s: 277
i: [1,283,0,7,2,0,0,3] rc1_subset_1
s: 283
d: [2,823,2,2,1,1,40,701,306]
a: [3,828,2,2,1,823]
s: 828
i: [1,285,0,7,2,0,0,255] rc5_subset_1
s: 285
d: [2,824,2,2,1,1,40,701,305]
a: [3,833,2,2,1,824]
s: 833
i: [1,288,0,7,2,0,0,255] rc2_finset_1
s: 288
i: [1,291,0,7,2,0,0,255] rc4_subset_1
s: 291
d: [2,848,2,2,1,1,40,706,306]
a: [3,853,2,2,1,848]
s: 853
i: [1,297,0,7,3,0,0,255] cc8_funct_1
s: 297
d: [2,849,2,2,1,1,40,706,305]
a: [3,858,2,2,1,849]
s: 858
i: [1,298,0,7,3,0,0,255] cc8_funct_1
s: 298
i: [1,300,0,7,3,0,0,255] cc7_finset_1
s: 300
d: [2,865,2,2,1,1,40,718,306]
a: [3,870,2,2,1,865]
s: 870
i: [1,310,0,7,3,0,0,1] cc5_ordinal1
s: 310
d: [2,866,2,2,1,1,40,718,305]
a: [3,875,2,2,1,866]
s: 875
i: [1,351,0,7,2,0,0,2] t3_subset
s: 351
i: [1,352,0,7,2,0,0,2] t3_subset
s: 352
d: [2,908,2,2,1,1,40,841,306]
a: [3,914,2,2,1,908]
s: 914
i: [1,353,0,7,3,0,0,2] t8_boole
s: 353
d: [2,909,2,2,1,1,40,841,305]
a: [3,919,2,2,1,909]
s: 919
i: [1,293,0,8,3,0,0,255] cc10_ordinal1
s: 293
i: [1,299,0,8,3,0,0,255] cc9_funct_1
s: 299
d: [2,925,2,2,1,1,40,846,306]
a: [3,930,2,2,1,925]
s: 930
i: [1,301,0,8,3,0,0,255] cc8_finset_1
s: 301
d: [2,926,2,2,1,1,40,846,305]
a: [3,935,2,2,1,926]
s: 935
i: [1,303,0,8,3,0,0,255] cc5_subset_1
s: 303
i: [1,313,0,8,3,0,0,255] cc2_finset_1
s: 313
d: [2,672,1,3,1,0,40,274,359]
s: 672
i: [1,324,0,8,3,0,0,3] cc1_subset_1
s: 324
d: [2,759,1,3,1,0,40,314,359]
s: 759
i: [1,329,0,8,4,0,0,255] cc6_funct_1
s: 329
i: [1,334,0,8,4,0,0,255] cc4_funct_1
s: 334
d: [2,1060,1,6,3,1,40,329,567]
d: [2,1067,1,4,2,2,21,1060,562]
a: [3,1071,1,2,1,1067]
s: 1071
d: [2,1124,0,6,3,0,21,334,322]
s: 1124
d: [2,1061,1,6,3,1,40,329,587]
d: [2,1077,1,4,2,2,21,1061,582]
a: [3,1081,1,2,1,1077]
s: 1081
i: [1,336,0,8,4,0,0,255] cc5_funct_1
s: 336
i: [1,343,0,8,4,0,0,255] cc3_funcop_1
s: 343
d: [2,1062,1,6,3,1,40,329,602]
d: [2,1087,1,4,2,2,21,1062,597]
a: [3,1091,1,2,1,1087]
s: 1091
d: [2,1153,0,6,3,0,21,343,322]
s: 1153
d: [2,1063,1,6,3,1,40,329,617]
d: [2,1097,1,4,2,2,21,1063,612]
d: [2,1098,1,2,1,3,21,1097,607]
a: [3,1102,1,2,1,1098]
s: 1102
i: [1,345,0,8,3,0,0,2] d2_ordinal1
s: 345
i: [1,349,0,8,3,0,0,2] t2_subset
s: 349
d: [2,1064,1,6,3,1,40,329,627]
d: [2,1104,1,4,2,2,21,1064,622]
a: [3,1108,1,2,1,1104]
s: 1108
i: [1,325,0,9,3,0,0,255] cc4_subset_1
s: 325
d: [2,1065,1,6,3,1,40,329,642]
d: [2,1114,1,4,2,2,21,1065,637]
a: [3,1118,1,2,1,1114]
s: 1118
i: [1,350,0,9,3,0,0,255] d8_xboole_0
s: 350
i: [1,356,0,9,3,0,0,2] t5_subset
s: 356
d: [2,1127,1,4,2,1,40,1124,602]
a: [3,1135,1,2,1,1127]
s: 1135
i: [1,294,0,10,4,0,0,255] t11_ordinal1
s: 294
d: [2,1128,1,4,2,1,40,1124,617]
d: [2,1137,1,2,1,2,21,1128,607]
a: [3,1141,1,2,1,1137]
s: 1141
i: [1,344,0,10,4,0,0,255] cc3_funct_1
s: 344
i: [1,355,0,10,3,0,0,2] t4_subset
s: 355
d: [2,1129,1,4,2,1,40,1124,627]
a: [3,1146,1,2,1,1129]
s: 1146
i: [1,367,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 367
r: 415
r: 418
a: [3,414,0,3,1,269]
s: 414
a: [3,419,0,5,1,269]
s: 419
i: [1,368,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 368
r: 414
r: 419
s: 415
s: 418
d: [2,1237,1,7,3,1,40,367,415]
d: [2,1240,1,5,2,2,21,1237,430]
d: [2,1241,1,3,1,3,21,1240,425]
a: [3,1245,1,3,1,1241]
s: 1245
r: 415
r: 418
r: 1245
s: 414
s: 419
a: [3,1244,1,3,1,1241]
s: 1244
i: [1,369,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 369
d: [2,1247,1,11,3,1,40,419,367]
a: [3,1251,1,3,1,1247]
s: 1251
i: [1,281,0,11,4,0,0,255] cc3_subset_1
s: 281
i: [1,282,0,11,4,0,0,255] cc2_subset_1
s: 282
d: [2,1263,1,7,3,1,40,368,414]
d: [2,1266,1,5,2,2,21,1263,425]
d: [2,1267,1,3,1,3,21,1266,430]
a: [3,1271,1,3,1,1267]
s: 1271
r: 414
r: 419
r: 1244
r: 1271
s: 415
s: 1245
s: 418
i: [1,307,0,12,4,0,0,1] t22_card_1
s: 307
a: [3,1270,1,3,1,1267]
s: 1270
i: [1,308,0,12,4,0,0,1] t21_card_1
s: 308
d: [2,1283,2,7,3,1,40,1245,368]
d: [2,1284,2,5,2,2,21,1283,430]
d: [2,1285,2,3,1,3,21,1284,425]
a: [3,1289,2,3,1,1285]
s: 1289
r: 415
r: 418
r: 1245
r: 1270
r: 1289
s: 414
s: 419
s: 1244
i: [1,309,0,12,4,0,0,1] t21_card_1
s: 309
s: 1271
d: [2,1130,1,4,2,1,40,1124,642]
a: [3,1151,1,2,1,1130]
s: 1151
d: [2,772,1,3,1,0,40,315,359]
s: 772
a: [3,1288,2,3,1,1285]
s: 1288
d: [2,785,1,3,1,0,40,316,359]
s: 785
d: [2,1319,2,6,2,1,40,1271,350]
a: [3,1323,2,3,1,1319]
s: 1323
r: 414
r: 419
r: 1244
r: 1271
r: 1288
r: 1323
s: 415
s: 418
s: 1245
s: 1270
s: 1289
d: [2,798,1,3,1,0,40,317,359]
s: 798
d: [2,811,1,3,1,0,40,318,359]
s: 811
a: [3,1322,2,3,1,1319]
s: 1322
d: [2,836,1,3,1,0,40,320,359]
s: 836
d: [2,1345,3,6,2,1,40,1289,350]
a: [3,1349,3,3,1,1345]
s: 1349
d: [2,861,1,3,1,0,40,321,359]
s: 861
d: [2,878,1,3,1,0,40,322,359]
s: 878
d: [2,1389,4,7,3,1,40,1349,294]
d: [2,1391,4,5,2,2,21,1389,430]
d: [2,1392,4,3,1,3,21,1391,732]
a: [3,1396,4,3,1,1392]
s: 1396
d: [2,937,1,4,2,0,40,346,354]
s: 937
d: [2,1401,5,2,1,1,40,1396,354]
a: [3,1411,5,2,1,1401]
s: 1411
d: [2,988,1,4,1,0,40,351,358]
s: 988
d: [2,989,1,4,1,0,40,351,360]
s: 989
d: [2,1364,3,7,3,1,40,1323,294]
d: [2,1366,3,5,2,2,21,1364,425]
d: [2,1367,3,3,1,3,21,1366,727]
a: [3,1370,3,3,1,1367]
s: 1370
d: [2,892,1,5,2,0,40,287,280]
s: 892
d: [2,1400,5,3,1,1,40,1396,348]
a: [3,1406,5,3,1,1400]
s: 1406
d: [2,897,1,5,2,0,40,304,296]
s: 897
d: [2,1449,2,0,0,2,52,425,1251,897]
e: 1449
d: [2,1450,2,2,1,1,40,897,1251]
d: [2,1454,2,0,0,2,21,1450,425]
e: 1454
r: 1251
a: [3,1250,1,3,1,1247]
s: 1250
d: [2,1423,6,4,2,1,40,1406,300]
a: [3,1429,6,2,1,1423]
s: 1429
a: [3,1255,1,3,1,1247]
s: 1255
d: [2,1469,2,0,0,2,52,430,1255,897]
e: 1469
d: [2,1470,2,2,1,1,40,1255,897]
d: [2,1472,2,0,0,2,21,1470,430]
e: 1472
r: 1255
d: [2,1424,6,4,2,1,40,1406,298]
a: [3,1438,6,2,1,1424]
s: 1438
a: [3,1254,1,3,1,1247]
s: 1254
d: [2,1277,1,11,3,1,40,418,368]
a: [3,1281,1,5,1,1277]
s: 1281
d: [2,1457,2,3,1,1,40,1250,306]
a: [3,1462,2,3,1,1457]
s: 1462
d: [2,950,1,5,2,0,40,297,357]
s: 950
d: [2,1458,2,3,1,1,40,1250,305]
a: [3,1467,2,3,1,1458]
s: 1467
d: [2,959,1,5,2,0,40,298,357]
s: 959
d: [2,968,1,5,2,0,40,300,357]
s: 968
d: [2,1474,2,3,1,1,40,1254,306]
a: [3,1479,2,3,1,1474]
s: 1479
d: [2,978,1,5,2,0,40,310,357]
s: 978
d: [2,1475,2,3,1,1,40,1254,305]
a: [3,1484,2,3,1,1475]
s: 1484
d: [2,987,1,5,1,0,40,351,357]
s: 987
d: [2,1006,1,5,2,0,40,293,358]
s: 1006
d: [2,1487,2,10,2,1,40,1281,350]
a: [3,1491,2,5,1,1487]
s: 1491
d: [2,1007,1,5,2,0,40,293,360]
s: 1007
d: [2,1506,3,11,3,1,40,1491,294]
d: [2,1508,3,8,2,2,21,1506,1250]
d: [2,1509,3,5,1,3,21,1508,1484]
a: [3,1513,3,5,1,1509]
s: 1513
d: [2,1520,4,0,0,4,52,430,425,1370,1513,309]
e: 1520
d: [2,1521,4,7,3,1,40,1513,309]
d: [2,1527,4,4,2,2,21,1521,1370]
d: [2,1528,4,2,1,3,21,1527,425]
d: [2,1529,4,0,0,4,21,1528,430]
e: 1529
r: 1491
r: 1513
d: [2,1016,1,5,2,0,40,299,358]
s: 1016
d: [2,1017,1,5,2,0,40,299,360]
s: 1017
a: [3,1490,2,5,1,1487]
s: 1490
d: [2,1027,1,5,2,0,40,301,358]
s: 1027
a: [3,1495,2,5,1,1487]
s: 1495
r: 418
r: 1250
r: 1281
r: 1462
r: 1467
r: 1490
d: [2,1531,0,5,1,2,23,418,1495]
a: [3,1548,0,5,1,1531]
s: 1548
d: [2,1028,1,5,2,0,40,301,360]
s: 1028
d: [2,1537,3,5,2,1,44,304,1495]
d: [2,1550,3,3,1,2,21,1537,425]
a: [3,1554,3,3,1,1550]
s: 1554
d: [2,1037,1,5,2,0,40,303,358]
s: 1037
d: [2,1507,3,5,1,1,40,1491,364]
a: [3,1518,3,5,1,1507]
s: 1518
d: [2,1038,1,5,2,0,40,303,360]
s: 1038
d: [2,1047,1,5,2,0,40,313,358]
s: 1047
d: [2,1538,3,12,4,1,44,307,1495]
d: [2,1556,3,10,3,2,21,1538,425]
s: 1556
r: 415
r: 1245
r: 1270
r: 1289
r: 1322
r: 1349
r: 1370
r: 1396
r: 1556
r: 1495
r: 1406
s: 414
s: 1244
s: 419
s: 1250
s: 1271
a: [3,1259,1,5,1,1247]
s: 1259
s: 1288
d: [2,1048,1,5,2,0,40,313,360]
s: 1048
d: [2,1059,1,5,2,0,40,324,360]
s: 1059
s: 1323
d: [2,1185,1,8,2,0,40,349,358]
d: [2,1194,1,5,1,0,21,1185,273]
s: 1194
a: [3,1326,2,3,1,1319]
s: 1326
d: [2,1186,1,8,2,0,40,349,360]
d: [2,1195,1,5,1,0,21,1186,273]
s: 1195
d: [2,1309,1,9,3,0,40,281,360]
d: [2,1317,1,5,2,0,21,1309,361]
s: 1317
s: 1462
d: [2,939,1,6,2,0,40,348,346]
s: 939
s: 1467
d: [2,940,1,6,2,0,40,365,346]
s: 940
d: [2,951,1,6,2,0,40,297,358]
s: 951
a: [3,1348,3,3,1,1345]
s: 1348
d: [2,952,1,6,2,0,40,297,360]
s: 952
a: [3,1371,3,3,1,1367]
s: 1371
d: [2,960,1,6,2,0,40,298,358]
s: 960
d: [2,961,1,6,2,0,40,298,360]
s: 961
d: [2,1643,4,2,1,1,40,1371,354]
a: [3,1653,4,2,1,1643]
s: 1653
d: [2,969,1,6,2,0,40,300,358]
s: 969
a: [3,1395,4,3,1,1392]
s: 1395
d: [2,970,1,6,2,0,40,300,360]
s: 970
d: [2,979,1,6,2,0,40,310,358]
s: 979
d: [2,1642,4,3,1,1,40,1371,348]
a: [3,1648,4,3,1,1642]
s: 1648
d: [2,980,1,6,2,0,40,310,360]
s: 980
d: [2,1658,5,4,2,1,40,1648,300]
a: [3,1664,5,2,1,1658]
s: 1664
d: [2,981,1,6,2,0,40,351,275]
s: 981
d: [2,982,1,6,2,0,40,351,277]
s: 982
d: [2,1659,5,4,2,1,40,1648,298]
a: [3,1673,5,2,1,1659]
s: 1673
d: [2,983,1,6,2,0,40,351,283]
s: 983
a: [3,1405,5,3,1,1400]
s: 1405
d: [2,984,1,6,2,0,40,351,285]
s: 984
d: [2,985,1,6,2,0,40,351,288]
s: 985
s: 1491
d: [2,986,1,6,2,0,40,351,291]
s: 986
a: [3,1494,2,5,1,1487]
s: 1494
d: [2,1005,1,6,2,0,40,293,357]
s: 1005
d: [2,1015,1,6,2,0,40,299,357]
s: 1015
d: [2,1589,2,11,3,1,40,1259,368]
d: [2,1590,2,8,2,2,21,1589,1254]
d: [2,1591,2,5,1,3,21,1590,1250]
a: [3,1595,2,5,1,1591]
s: 1595
r: 1491
a: [3,1280,1,5,1,1277]
s: 1280
s: 1490
d: [2,1026,1,6,2,0,40,301,357]
s: 1026
d: [2,1036,1,6,2,0,40,303,357]
s: 1036
s: 1513
d: [2,1707,4,0,0,3,52,1467,1280,1513,345]
e: 1707
d: [2,1709,4,8,2,1,40,1513,345]
d: [2,1714,4,3,1,2,21,1709,1280]
d: [2,1715,4,0,0,3,21,1714,1467]
e: 1715
r: 1513
d: [2,1046,1,6,2,0,40,313,357]
s: 1046
a: [3,1512,3,5,1,1509]
s: 1512
d: [2,1717,4,0,0,4,52,430,425,1371,1512,308]
e: 1717
d: [2,1718,4,7,3,1,40,1512,308]
d: [2,1720,4,4,2,2,21,1718,1371]
d: [2,1721,4,2,1,3,21,1720,425]
d: [2,1722,4,0,0,4,21,1721,430]
e: 1722
i: [1,268,0,8,2,1,0,0] t23_card_1
d: [2,1375,1,7,3,1,40,415,367]
d: [2,1376,1,5,2,2,21,1375,430]
d: [2,1377,1,3,1,3,21,1376,425]
d: [2,1380,2,7,3,1,40,1245,368]
d: [2,1381,2,5,2,2,21,1380,430]
d: [2,1382,2,3,1,3,21,1381,425]
d: [2,1385,3,6,2,1,40,1289,350]
d: [2,1379,1,11,3,1,40,418,368]
d: [2,1568,4,5,2,2,47,1556]
d: [2,1569,4,3,1,3,21,1568,430]
d: [2,1574,1,7,3,1,40,414,368]
d: [2,1575,1,5,2,2,21,1574,425]
d: [2,1576,1,3,1,3,21,1575,430]
d: [2,1579,1,11,3,1,40,419,367]
d: [2,1580,1,8,2,2,21,1579,1254]
d: [2,1583,2,3,1,1,40,1250,305]
d: [2,1587,2,6,2,1,40,1271,350]
d: [2,1607,3,7,3,1,40,1323,294]
d: [2,1609,3,5,2,2,21,1607,425]
d: [2,1610,3,3,1,3,21,1609,727]
d: [2,1641,4,3,1,1,40,1371,365]
d: [2,1694,3,11,3,1,40,1491,294]
d: [2,1696,3,8,2,2,21,1694,1250]
d: [2,1697,3,5,1,3,21,1696,1484]
d: [2,1695,3,5,1,1,40,1491,364]
d: [2,1703,3,10,2,1,40,1595,350]
d: [2,1704,3,5,1,2,21,1703,1494]
f: 1722,1715,1704,1695,1697,1641,1610,1591,1587,1583,1580,1576,1569,1529,1487,1475,1472,1379,1454,1392,1385,1382,1377,721,720,266,267,268,269
% Refutation found. Thanks to Tanya!
% SZS status Theorem for card_1__t23_card_1
% SZS output start Proof for card_1__t23_card_1
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r1_ordinal1(X0,X1) <=> r1_ordinal1(k3_card_1(X0),k3_card_1(X1))))) [input t23_card_1]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r1_ordinal1(X0,X1) <=> r1_ordinal1(k3_card_1(X0),k3_card_1(X1))))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
4. ! [X0,X1] : (r2_xboole_0(X0,X1) => ~r2_xboole_0(X1,X0)) [input antisymmetry_r2_xboole_0]
6. ! [X0] : (v1_card_1(X0) => v3_ordinal1(X0)) [input cc1_card_1]
9. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
42. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
43. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r2_hidden(X1,X0) => r1_tarski(X1,X0))) [input d2_ordinal1]
45. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
56. ! [X0] : (v3_ordinal1(X0) => v1_card_1(k3_card_1(X0))) [input fc3_card_1]
85. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
88. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
90. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,X1) <=> r2_hidden(k3_card_1(X0),k3_card_1(X1))))) [input t21_card_1]
91. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (k3_card_1(X0) = k3_card_1(X1) => X0 = X1))) [input t22_card_1]
101. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 45]
109. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 101]
116. ? [X0] : (? [X1] : ((r1_ordinal1(X0,X1) <~> r1_ordinal1(k3_card_1(X0),k3_card_1(X1))) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
130. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 88]
131. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 130]
133. ! [X0] : (v3_ordinal1(X0) | ~v1_card_1(X0)) [ennf transformation 6]
140. ! [X0] : (v1_card_1(k3_card_1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 56]
141. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 9]
142. ! [X0] : (! [X1] : ((X0 = X1 | k3_card_1(X0) != k3_card_1(X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 91]
143. ! [X0] : (! [X1] : (X0 = X1 | k3_card_1(X0) != k3_card_1(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 142]
144. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) <=> r2_hidden(k3_card_1(X0),k3_card_1(X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 90]
179. ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0))) [ennf transformation 43]
183. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 109]
184. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 183]
190. ! [X0,X1] : (~r2_xboole_0(X1,X0) | ~r2_xboole_0(X0,X1)) [ennf transformation 4]
191. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
194. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 42]
195. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 194]
196. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 85]
197. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 196]
198. ? [X0] : (? [X1] : (((~r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | r1_ordinal1(X0,X1))) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [nnf transformation 116]
199. ? [X0] : (? [X1] : ((~r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | r1_ordinal1(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 198]
200. ? [X0] : (? [X1] : ((~r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | ~r1_ordinal1(X0,X1)) & (r1_ordinal1(k3_card_1(X0),k3_card_1(X1)) | r1_ordinal1(X0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : ((~r1_ordinal1(k3_card_1(sK0),k3_card_1(X1)) | ~r1_ordinal1(sK0,X1)) & (r1_ordinal1(k3_card_1(sK0),k3_card_1(X1)) | r1_ordinal1(sK0,X1)) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
201. ? [X1] : ((~r1_ordinal1(k3_card_1(sK0),k3_card_1(X1)) | ~r1_ordinal1(sK0,X1)) & (r1_ordinal1(k3_card_1(sK0),k3_card_1(X1)) | r1_ordinal1(sK0,X1)) & v3_ordinal1(X1)) => ((~r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | ~r1_ordinal1(sK0,sK1)) & (r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | r1_ordinal1(sK0,sK1)) & v3_ordinal1(sK1)) [choice axiom]
202. ((~r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | ~r1_ordinal1(sK0,sK1)) & (r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | r1_ordinal1(sK0,sK1)) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 199,201,200]
215. ! [X0] : (! [X1] : (((r2_hidden(X0,X1) | ~r2_hidden(k3_card_1(X0),k3_card_1(X1))) & (r2_hidden(k3_card_1(X0),k3_card_1(X1)) | ~r2_hidden(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [nnf transformation 144]
216. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)) | ~v1_ordinal1(X0))) [nnf transformation 179]
217. ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [rectify 216]
218. ! [X0] : (? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0)) => (~r1_tarski(sK8(X0),X0) & r2_hidden(sK8(X0),X0))) [choice axiom]
219. ! [X0] : ((v1_ordinal1(X0) | (~r1_tarski(sK8(X0),X0) & r2_hidden(sK8(X0),X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0))) [skolemisation 217,218]
227. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 197]
266. v3_ordinal1(sK0) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
267. v3_ordinal1(sK1) [cnf transformation 202] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
268. r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | r1_ordinal1(sK0,sK1) [cnf transformation 202] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
269. ~r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | ~r1_ordinal1(sK0,sK1) [cnf transformation 202] {a:0,w:8,wCS:800,goal:1,sine:0,thAx:0,allAx:3,thDist:-3}
294. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 131] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
296. ~v1_card_1(X0) | v3_ordinal1(X0) [cnf transformation 133] {a:0,w:4,wCS:1600,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
304. v1_card_1(k3_card_1(X0)) | ~v3_ordinal1(X0) [cnf transformation 140] {a:0,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
305. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 141] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
307. k3_card_1(X0) != k3_card_1(X1) | X0 = X1 | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 143] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
308. r2_hidden(k3_card_1(X0),k3_card_1(X1)) | ~r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 215] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
309. ~r2_hidden(k3_card_1(X0),k3_card_1(X1)) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 215] {a:0,w:12,wCS:4800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
345. ~r2_hidden(X2,X0) | r1_tarski(X2,X0) | ~v1_ordinal1(X0) [cnf transformation 219] {a:0,w:8,wCS:3200,nSel:1,sine:2,thAx:0,allAx:2,thDist:-2}
350. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 184] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
364. ~r2_xboole_0(X1,X0) | ~r2_xboole_0(X0,X1) [cnf transformation 190] {a:0,w:6,wCS:2400,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
365. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 191] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
367. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 195] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
368. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 227] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
413. 1 <=> r1_ordinal1(sK0,sK1) [avatar definition]
414. r1_ordinal1(sK0,sK1) <- (1) [avatar component clause 413] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
415. ~r1_ordinal1(sK0,sK1) <- (~1) [avatar component clause 413] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
417. 2 <=> r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) [avatar definition]
418. r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) <- (2) [avatar component clause 417] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
419. ~r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) <- (~2) [avatar component clause 417] {a:0,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
420. ~1 | ~2 [avatar split clause 269,417,413]
421. 1 | 2 [avatar split clause 268,417,413]
423. 3 <=> v3_ordinal1(sK1) [avatar definition]
425. v3_ordinal1(sK1) <- (3) [avatar component clause 423] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
426. 3 [avatar split clause 267,423]
428. 4 <=> v3_ordinal1(sK0) [avatar definition]
430. v3_ordinal1(sK0) <- (4) [avatar component clause 428] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:3,thDist:-3}
431. 4 [avatar split clause 266,428]
720. v1_ordinal1(sK0) <- (4) [resolution 305,430] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
721. v1_ordinal1(sK1) <- (3) [resolution 305,425] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
725. 60 <=> v1_ordinal1(sK0) [avatar definition]
727. v1_ordinal1(sK0) <- (60) [avatar component clause 725] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
728. 60 | ~4 [avatar split clause 720,428,725]
730. 61 <=> v1_ordinal1(sK1) [avatar definition]
732. v1_ordinal1(sK1) <- (61) [avatar component clause 730] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
733. 61 | ~3 [avatar split clause 721,423,730]
897. v3_ordinal1(k3_card_1(X0)) | ~v3_ordinal1(X0) [resolution 304,296] {a:1,w:5,wCS:2000,nSel:1,sine:1,thAx:0,allAx:2,thDist:-2}
1243. 108 <=> r1_ordinal1(sK1,sK0) [avatar definition]
1245. r1_ordinal1(sK1,sK0) <- (108) [avatar component clause 1243] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1249. 109 <=> v3_ordinal1(k3_card_1(sK1)) [avatar definition]
1250. v3_ordinal1(k3_card_1(sK1)) <- (109) [avatar component clause 1249] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1251. ~v3_ordinal1(k3_card_1(sK1)) <- (~109) [avatar component clause 1249] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1253. 110 <=> v3_ordinal1(k3_card_1(sK0)) [avatar definition]
1254. v3_ordinal1(k3_card_1(sK0)) <- (110) [avatar component clause 1253] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1255. ~v3_ordinal1(k3_card_1(sK0)) <- (~110) [avatar component clause 1253] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1257. 111 <=> r1_ordinal1(k3_card_1(sK1),k3_card_1(sK0)) [avatar definition]
1259. r1_ordinal1(k3_card_1(sK1),k3_card_1(sK0)) <- (111) [avatar component clause 1257] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1269. 112 <=> r1_tarski(sK0,sK1) [avatar definition]
1271. r1_tarski(sK0,sK1) <- (112) [avatar component clause 1269] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1279. 113 <=> r1_tarski(k3_card_1(sK0),k3_card_1(sK1)) [avatar definition]
1280. ~r1_tarski(k3_card_1(sK0),k3_card_1(sK1)) <- (~113) [avatar component clause 1279] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1281. r1_tarski(k3_card_1(sK0),k3_card_1(sK1)) <- (113) [avatar component clause 1279] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1287. 114 <=> r1_tarski(sK1,sK0) [avatar definition]
1289. r1_tarski(sK1,sK0) <- (114) [avatar component clause 1287] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1321. 115 <=> r2_xboole_0(sK0,sK1) [avatar definition]
1323. r2_xboole_0(sK0,sK1) <- (115) [avatar component clause 1321] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1325. 116 <=> sK0 = sK1 [avatar definition]
1347. 117 <=> r2_xboole_0(sK1,sK0) [avatar definition]
1349. r2_xboole_0(sK1,sK0) <- (117) [avatar component clause 1347] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1369. 118 <=> r2_hidden(sK0,sK1) [avatar definition]
1370. ~r2_hidden(sK0,sK1) <- (~118) [avatar component clause 1369] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1371. r2_hidden(sK0,sK1) <- (118) [avatar component clause 1369] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1374. sK0 != sK1 | r1_ordinal1(sK0,sK1) | ~r1_ordinal1(sK1,sK0) [theory tautology sat conflict] {a:0,w:9,wCS:3600,sine:255,thAx:0,allAx:1,thDist:-1}
1375. r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (~1) [resolution 415,367] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1376. r1_ordinal1(sK1,sK0) | ~v3_ordinal1(sK1) <- (~1, 4) [subsumption resolution 1375,430] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1377. r1_ordinal1(sK1,sK0) <- (~1, 3, 4) [subsumption resolution 1376,425] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1378. 108 | 1 | ~3 | ~4 [avatar split clause 1377,428,423,413,1243]
1379. r1_tarski(k3_card_1(sK0),k3_card_1(sK1)) | ~v3_ordinal1(k3_card_1(sK1)) | ~v3_ordinal1(k3_card_1(sK0)) <- (2) [resolution 418,368] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1380. r1_tarski(sK1,sK0) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK1) <- (108) [resolution 1245,368] {a:2,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1381. r1_tarski(sK1,sK0) | ~v3_ordinal1(sK1) <- (4, 108) [subsumption resolution 1380,430] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1382. r1_tarski(sK1,sK0) <- (3, 4, 108) [subsumption resolution 1381,425] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1383. 114 | ~3 | ~4 | ~108 [avatar split clause 1382,1243,428,423,1287]
1385. sK0 = sK1 | r2_xboole_0(sK1,sK0) <- (114) [resolution 1289,350] {a:3,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1386. 117 | 116 | ~114 [avatar split clause 1385,1287,1325,1347]
1389. r2_hidden(sK1,sK0) | ~v3_ordinal1(sK0) | ~v1_ordinal1(sK1) <- (117) [resolution 1349,294] {a:4,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1391. r2_hidden(sK1,sK0) | ~v1_ordinal1(sK1) <- (4, 117) [subsumption resolution 1389,430] {a:4,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1392. r2_hidden(sK1,sK0) <- (4, 61, 117) [subsumption resolution 1391,732] {a:4,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1394. 119 <=> r2_hidden(sK1,sK0) [avatar definition]
1397. 119 | ~4 | ~61 | ~117 [avatar split clause 1392,1347,730,428,1394]
1450. ~v3_ordinal1(sK1) <- (~109) [resolution 897,1251] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1454. $false <- (3, ~109) [subsumption resolution 1450,425] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1455. ~3 | 109 [avatar contradiction clause 1454]
1456. ~110 | ~109 | 113 | ~2 [avatar split clause 1379,417,1279,1249,1253]
1465. 128 <=> v1_ordinal1(k3_card_1(sK1)) [avatar definition]
1467. v1_ordinal1(k3_card_1(sK1)) <- (128) [avatar component clause 1465] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1470. ~v3_ordinal1(sK0) <- (~110) [resolution 1255,897] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1472. $false <- (4, ~110) [subsumption resolution 1470,430] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1473. ~4 | 110 [avatar contradiction clause 1472]
1475. v1_ordinal1(k3_card_1(sK0)) <- (110) [resolution 1254,305] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1482. 130 <=> v1_ordinal1(k3_card_1(sK0)) [avatar definition]
1484. v1_ordinal1(k3_card_1(sK0)) <- (130) [avatar component clause 1482] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1485. 130 | ~110 [avatar split clause 1475,1253,1482]
1487. k3_card_1(sK0) = k3_card_1(sK1) | r2_xboole_0(k3_card_1(sK0),k3_card_1(sK1)) <- (113) [resolution 1281,350] {a:2,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1489. 131 <=> r2_xboole_0(k3_card_1(sK0),k3_card_1(sK1)) [avatar definition]
1491. r2_xboole_0(k3_card_1(sK0),k3_card_1(sK1)) <- (131) [avatar component clause 1489] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1493. 132 <=> k3_card_1(sK0) = k3_card_1(sK1) [avatar definition]
1494. k3_card_1(sK0) != k3_card_1(sK1) <- (~132) [avatar component clause 1493] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1495. k3_card_1(sK0) = k3_card_1(sK1) <- (132) [avatar component clause 1493] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1496. 131 | 132 | ~113 [avatar split clause 1487,1279,1493,1489]
1511. 133 <=> r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) [avatar definition]
1512. ~r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) <- (~133) [avatar component clause 1511] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1513. r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) <- (133) [avatar component clause 1511] {a:3,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1516. 134 <=> r2_xboole_0(k3_card_1(sK1),k3_card_1(sK0)) [avatar definition]
1521. r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (133) [resolution 1513,309] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1527. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~118, 133) [subsumption resolution 1521,1370] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1528. ~v3_ordinal1(sK0) <- (3, ~118, 133) [subsumption resolution 1527,425] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1529. $false <- (3, 4, ~118, 133) [subsumption resolution 1528,430] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1530. ~3 | ~4 | 118 | ~133 [avatar contradiction clause 1529]
1538. k3_card_1(X0) != k3_card_1(sK0) | sK1 = X0 | ~v3_ordinal1(X0) | ~v3_ordinal1(sK1) <- (132) [superposition 307,1495] {a:3,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1556. k3_card_1(X0) != k3_card_1(sK0) | sK1 = X0 | ~v3_ordinal1(X0) <- (3, 132) [subsumption resolution 1538,425] {a:3,w:10,wCS:1000,nSel:1,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1568. sK0 = sK1 | ~v3_ordinal1(sK0) <- (3, 132) [equality resolution 1556] {a:4,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1569. sK0 = sK1 <- (3, 4, 132) [subsumption resolution 1568,430] {a:4,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1570. 116 | ~3 | ~4 | ~132 [avatar split clause 1569,1493,428,423,1325]
1571. sK0 != sK1 | r1_ordinal1(k3_card_1(sK0),k3_card_1(sK1)) | ~r1_ordinal1(k3_card_1(sK1),k3_card_1(sK0)) [theory tautology sat conflict] {a:0,w:13,wCS:5200,sine:255,thAx:0,allAx:1,thDist:-1}
1574. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (1) [resolution 414,368] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1575. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK0) <- (1, 3) [subsumption resolution 1574,425] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1576. r1_tarski(sK0,sK1) <- (1, 3, 4) [subsumption resolution 1575,430] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1577. 112 | ~1 | ~3 | ~4 [avatar split clause 1576,428,423,413,1269]
1579. r1_ordinal1(k3_card_1(sK1),k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK1)) <- (~2) [resolution 419,367] {a:1,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1580. r1_ordinal1(k3_card_1(sK1),k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK1)) <- (~2, 110) [subsumption resolution 1579,1254] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:4,thDist:-4}
1581. ~109 | 111 | 2 | ~110 [avatar split clause 1580,1253,417,1257,1249]
1583. v1_ordinal1(k3_card_1(sK1)) <- (109) [resolution 1250,305] {a:2,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1585. 128 | ~109 [avatar split clause 1583,1249,1465]
1587. sK0 = sK1 | r2_xboole_0(sK0,sK1) <- (112) [resolution 1271,350] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1588. 115 | 116 | ~112 [avatar split clause 1587,1269,1325,1321]
1589. r1_tarski(k3_card_1(sK1),k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK1)) <- (111) [resolution 1259,368] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1590. r1_tarski(k3_card_1(sK1),k3_card_1(sK0)) | ~v3_ordinal1(k3_card_1(sK1)) <- (110, 111) [subsumption resolution 1589,1254] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1591. r1_tarski(k3_card_1(sK1),k3_card_1(sK0)) <- (109, 110, 111) [subsumption resolution 1590,1250] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1593. 137 <=> r1_tarski(k3_card_1(sK1),k3_card_1(sK0)) [avatar definition]
1595. r1_tarski(k3_card_1(sK1),k3_card_1(sK0)) <- (137) [avatar component clause 1593] {a:2,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1596. 137 | ~109 | ~110 | ~111 [avatar split clause 1591,1257,1253,1249,1593]
1607. r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v1_ordinal1(sK0) <- (115) [resolution 1323,294] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1609. r2_hidden(sK0,sK1) | ~v1_ordinal1(sK0) <- (3, 115) [subsumption resolution 1607,425] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1610. r2_hidden(sK0,sK1) <- (3, 60, 115) [subsumption resolution 1609,727] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1611. 118 | ~3 | ~60 | ~115 [avatar split clause 1610,1321,725,423,1369]
1641. ~r2_hidden(sK1,sK0) <- (118) [resolution 1371,365] {a:4,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:7,thDist:-7}
1644. ~119 | ~118 [avatar split clause 1641,1369,1394]
1694. r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) | ~v3_ordinal1(k3_card_1(sK1)) | ~v1_ordinal1(k3_card_1(sK0)) <- (131) [resolution 1491,294] {a:3,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1695. ~r2_xboole_0(k3_card_1(sK1),k3_card_1(sK0)) <- (131) [resolution 1491,364] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1696. r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) | ~v1_ordinal1(k3_card_1(sK0)) <- (109, 131) [subsumption resolution 1694,1250] {a:3,w:8,wCS:800,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1697. r2_hidden(k3_card_1(sK0),k3_card_1(sK1)) <- (109, 130, 131) [subsumption resolution 1696,1484] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1698. 133 | ~109 | ~130 | ~131 [avatar split clause 1697,1489,1482,1249,1511]
1699. ~134 | ~131 [avatar split clause 1695,1489,1516]
1703. k3_card_1(sK0) = k3_card_1(sK1) | r2_xboole_0(k3_card_1(sK1),k3_card_1(sK0)) <- (137) [resolution 1595,350] {a:3,w:10,wCS:1000,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1704. r2_xboole_0(k3_card_1(sK1),k3_card_1(sK0)) <- (~132, 137) [subsumption resolution 1703,1494] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1705. 134 | 132 | ~137 [avatar split clause 1704,1593,1493,1516]
1709. r1_tarski(k3_card_1(sK0),k3_card_1(sK1)) | ~v1_ordinal1(k3_card_1(sK1)) <- (133) [resolution 1513,345] {a:4,w:8,wCS:800,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
1714. ~v1_ordinal1(k3_card_1(sK1)) <- (~113, 133) [subsumption resolution 1709,1280] {a:4,w:3,wCS:300,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
1715. $false <- (~113, 128, 133) [subsumption resolution 1714,1467] {a:4,w:0,goal:1,sine:2,thAx:0,allAx:8,thDist:-8}
1716. 113 | ~128 | ~133 [avatar contradiction clause 1715]
1718. ~r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (~133) [resolution 1512,308] {a:4,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1720. ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (118, ~133) [subsumption resolution 1718,1371] {a:4,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1721. ~v3_ordinal1(sK0) <- (3, 118, ~133) [subsumption resolution 1720,425] {a:4,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1722. $false <- (3, 4, 118, ~133) [subsumption resolution 1721,430] {a:4,w:0,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
1723. ~3 | ~4 | ~118 | 133 [avatar contradiction clause 1722]
1724. $false [avatar sat refutation 420,421,426,431,728,733,1374,1378,1383,1386,1397,1455,1456,1473,1485,1496,1530,1570,1571,1577,1581,1585,1588,1596,1611,1644,1698,1699,1705,1716,1723] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:1071,thDist:-1071}
% SZS output end Proof for card_1__t23_card_1
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 97
% >>> Preprocessing
% Introduced skolems: 31
% Pure predicates: 4
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 146
% Generated clauses: 886
% Active clauses: 299
% Passive clauses: 489
% Final active clauses: 252
% Final passive clauses: 169
% Main loop iterations started: 314
% 
% >>> Simplifying Inferences
% Duplicate literals: 14
% Fw subsumption resolutions: 79
% Fw demodulations: 1
% Bw demodulations: 7
% 
% >>> Deletion Inferences
% Simple tautologies: 32
% Equational tautologies: 2
% Forward subsumptions: 104
% 
% >>> Generating Inferences
% Binary resolution: 432
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 6
% Backward superposition: 14
% Self superposition: 1
% Equality resolution: 2
% 
% >>> AVATAR
% Split clauses: 31
% Split components: 66
% Sat splitting refutations: 13
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 179
% SAT solver unit clauses: 47
% SAT solver binary clauses: 70
% 
% Memory used [KB]: 6012
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
% Time measurement results:
% property evaluation: 0.001 s
% backward superposition index maintenance: 0.001 s
% resolution: 0.001 s
% passive container maintenance: 0.004 s
% other: 0.009 s ( own 0.002 s ) 

