i: [1,114,0,2,1,1,0,0] l1_afvect01
a: [3,193,0,2,1,114]
s: 193
i: [1,113,0,2,1,1,0,0] l1_afvect01
a: [3,198,0,2,1,113]
s: 198
i: [1,112,0,2,1,1,0,0] l1_afvect01
a: [3,203,0,2,1,112]
s: 203
i: [1,119,0,3,1,1,0,0] l1_afvect01
a: [3,168,0,3,1,119]
s: 168
i: [1,117,0,4,1,1,0,0] l1_afvect01
a: [3,178,0,4,1,117]
s: 178
i: [1,116,0,4,1,1,0,0] l1_afvect01
a: [3,183,0,4,1,116]
s: 183
i: [1,115,0,4,1,1,0,0] l1_afvect01
a: [3,188,0,4,1,115]
s: 188
i: [1,120,0,6,1,1,0,0] l1_afvect01
a: [3,163,0,6,1,120]
s: 163
i: [1,118,0,6,1,1,0,0] l1_afvect01
a: [3,173,0,6,1,118]
s: 173
i: [1,121,0,2,1,0,0,255] spc1_boole
a: [3,208,0,2,1,121]
s: 208
i: [1,122,0,2,1,0,0,2] fc1_xboole_0
a: [3,213,0,2,1,122]
s: 213
i: [1,156,0,2,1,0,0,2] rc2_xboole_0
a: [3,218,0,2,1,156]
s: 218
i: [1,157,0,2,1,0,0,1] existence_l1_analoaf
a: [3,223,0,2,1,157]
s: 223
i: [1,158,0,2,1,0,0,2] existence_l1_struct_0
a: [3,228,0,2,1,158]
s: 228
i: [1,159,0,2,1,0,0,2] rc1_xboole_0
a: [3,233,0,2,1,159]
s: 233
i: [1,123,0,4,2,0,0,1] dt_l1_analoaf
s: 123
d: [2,235,1,2,1,1,40,123,193]
a: [3,240,1,2,1,235]
s: 240
i: [1,154,0,4,1,0,0,1] existence_m1_subset_1
s: 154
d: [2,236,1,2,1,1,40,123,223]
a: [3,245,1,2,1,236]
s: 245
i: [1,138,0,5,2,0,0,1] t6_boole
s: 138
i: [1,153,0,5,2,0,0,255] t7_boole
s: 153
i: [1,124,0,6,3,0,0,255] cc7_struct_0
s: 124
i: [1,125,0,6,3,0,0,255] cc5_struct_0
s: 125
d: [2,249,1,4,2,1,40,124,228]
a: [3,255,1,2,1,249]
s: 255
i: [1,127,0,6,3,0,0,1] cc17_struct_0
s: 127
i: [1,128,0,6,3,0,0,1] cc2_struct_0
s: 128
d: [2,251,2,4,2,1,40,124,245]
a: [3,264,2,2,1,251]
s: 264
i: [1,129,0,6,3,0,0,1] cc1_struct_0
s: 129
i: [1,132,0,6,3,0,0,255] cc4_struct_0
s: 132
i: [1,133,0,6,3,0,0,255] cc6_struct_0
s: 133
i: [1,150,0,6,2,0,0,255] t1_subset
s: 150
i: [1,155,0,6,2,0,0,255] antisymmetry_r2_hidden
s: 155
i: [1,130,0,7,3,0,0,2] cc8_struct_0
s: 130
i: [1,134,0,7,3,0,0,255] cc11_struct_0
s: 134
i: [1,135,0,7,3,0,0,255] cc11_struct_0
s: 135
i: [1,136,0,7,3,0,0,2] cc9_struct_0
s: 136
i: [1,139,0,7,3,0,0,255] fc9_struct_0
s: 139
i: [1,143,0,7,3,0,0,1] fc2_struct_0
s: 143
i: [1,146,0,7,3,0,0,255] fc6_struct_0
s: 146
i: [1,147,0,7,3,0,0,255] fc8_struct_0
s: 147
i: [1,148,0,7,3,0,0,1] fc1_struct_0
s: 148
i: [1,149,0,7,3,0,0,255] fc7_struct_0
s: 149
i: [1,152,0,7,3,0,0,1] t8_boole
s: 152
d: [2,247,1,7,3,0,44,138,138]
s: 247
i: [1,151,0,8,3,0,0,255] t2_subset
s: 151
i: [1,137,0,9,4,0,0,255] cc10_struct_0
s: 137
d: [2,280,1,7,2,1,40,151,178]
a: [3,291,1,3,1,280]
s: 291
r: 291
i: [1,144,0,27,8,0,0,1] t4_afvect0
s: 144
a: [3,290,1,3,1,280]
s: 290
i: [1,141,0,32,8,0,0,1] d4_diraf
s: 141
i: [1,142,0,32,8,0,0,1] d4_diraf
s: 142
d: [2,307,2,4,2,1,40,291,143]
d: [2,308,2,2,1,2,21,307,240]
a: [3,311,2,2,1,308]
s: 311
i: [1,145,0,34,9,0,0,1] t7_afvect0
s: 145
a: [3,287,1,4,1,280]
s: 287
i: [1,140,0,38,9,0,0,1] d4_diraf
s: 140
r: 290
r: 311
r: 287
s: 291
a: [3,312,2,2,1,308]
s: 312
d: [2,343,3,0,0,3,52,240,203,312,127]
e: 343
d: [2,344,3,4,2,1,40,312,127]
d: [2,347,3,2,1,2,21,344,203]
d: [2,348,3,0,0,3,21,347,240]
e: 348
r: 291
r: 312
s: 290
s: 311
s: 287
d: [2,281,1,7,2,1,40,151,183]
a: [3,296,1,4,1,281]
s: 296
d: [2,282,1,7,2,1,40,151,188]
a: [3,301,1,4,1,282]
s: 301
d: [2,325,1,32,8,1,40,140,173]
d: [2,326,1,28,7,1,18,325]
d: [2,327,1,22,6,2,21,326,163]
d: [2,328,1,18,5,3,21,327,178]
d: [2,329,1,14,4,4,21,328,183]
d: [2,330,1,10,3,5,21,329,188]
d: [2,331,1,8,2,6,21,330,193]
d: [2,332,1,6,1,7,21,331,311]
a: [3,336,1,6,1,332]
s: 336
d: [2,317,2,4,1,1,40,287,155]
a: [3,323,2,4,1,317]
s: 323
d: [2,283,1,6,2,0,40,151,154]
s: 283
d: [2,270,1,9,4,0,44,130,138]
s: 270
d: [2,358,2,4,1,1,40,296,155]
a: [3,364,2,4,1,358]
s: 364
d: [2,272,1,9,4,0,44,136,138]
s: 272
d: [2,366,2,4,1,1,40,301,155]
a: [3,372,2,4,1,366]
s: 372
d: [2,374,2,28,8,1,40,336,145]
d: [2,379,2,24,7,1,18,374]
d: [2,380,2,20,6,2,21,379,183]
d: [2,381,2,16,5,3,21,380,178]
d: [2,382,2,12,4,4,21,381,188]
d: [2,383,2,10,3,5,21,382,193]
d: [2,384,2,8,2,6,21,383,198]
d: [2,385,2,6,1,7,21,384,203]
a: [3,389,2,6,1,385]
s: 389
d: [2,416,3,0,0,8,52,203,198,193,168,183,188,178,389,144]
e: 416
d: [2,417,3,21,7,1,40,389,144]
d: [2,425,3,18,6,2,21,417,168]
d: [2,426,3,14,5,3,21,425,178]
d: [2,427,3,10,4,4,21,426,188]
d: [2,428,3,6,3,5,21,427,183]
d: [2,429,3,4,2,6,21,428,193]
d: [2,430,3,2,1,7,21,429,198]
d: [2,431,3,0,0,8,21,430,203]
e: 431
f: 431,385,348,331,235,112,113,114,115,116,117,118,119,120
% Refutation found. Thanks to Tanya!
% SZS status Theorem for afvect01__l1_afvect01
% SZS output start Proof for afvect01__l1_afvect01
1. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X2,X3) => (r2_analoaf(X0,X1,X2,X2,X3) | X1 = X3)))))) [input l1_afvect01]
2. ~! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X2,X3) => (r2_analoaf(X0,X1,X2,X2,X3) | X1 = X3)))))) [negated conjecture 1]
6. ! [X0] : (l1_struct_0(X0) => (~v7_struct_0(X0) => ~v2_struct_0(X0))) [input cc17_struct_0]
15. ! [X0] : ((l1_analoaf(X0) & ~v2_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)))))))) [input d4_diraf]
17. ! [X0] : (l1_analoaf(X0) => l1_struct_0(X0)) [input dt_l1_analoaf]
36. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X1,X3) => X2 = X3))))) [input t4_afvect0]
38. ! [X0] : ((l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => ! [X3] : (m1_subset_1(X3,u1_struct_0(X0)) => ! [X4] : (m1_subset_1(X4,u1_struct_0(X0)) => (r2_analoaf(X0,X1,X2,X3,X4) => r2_analoaf(X0,X2,X1,X4,X3))))))) [input t7_afvect0]
45. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((~r2_analoaf(X0,X1,X2,X2,X3) & X1 != X3) & r2_diraf(X0,X1,X2,X2,X3)) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & (l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0))) [ennf transformation 2]
46. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_analoaf(X0,X1,X2,X2,X3) & X1 != X3 & r2_diraf(X0,X1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) [flattening 45]
47. ! [X0] : (l1_struct_0(X0) | ~l1_analoaf(X0)) [ennf transformation 17]
52. ! [X0] : ((~v2_struct_0(X0) | v7_struct_0(X0)) | ~l1_struct_0(X0)) [ennf transformation 6]
53. ! [X0] : (~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0)) [flattening 52]
73. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | v2_struct_0(X0))) [ennf transformation 15]
74. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_diraf(X0,X1,X2,X3,X4) <=> (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 73]
77. ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((X2 = X3 | ~r2_analoaf(X0,X1,X2,X1,X3)) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 36]
78. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (X2 = X3 | ~r2_analoaf(X0,X1,X2,X1,X3) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 77]
79. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : ((r2_analoaf(X0,X2,X1,X4,X3) | ~r2_analoaf(X0,X1,X2,X3,X4)) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0))) [ennf transformation 38]
80. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r2_analoaf(X0,X2,X1,X4,X3) | ~r2_analoaf(X0,X1,X2,X3,X4) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0)) [flattening 79]
95. ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~r2_analoaf(X0,X1,X2,X2,X3) & X1 != X3 & r2_diraf(X0,X1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(X0))) & m1_subset_1(X2,u1_struct_0(X0))) & m1_subset_1(X1,u1_struct_0(X0))) & l1_analoaf(X0) & v1_afvect0(X0) & ~v7_struct_0(X0)) => (? [X1] : (? [X2] : (? [X3] : (~r2_analoaf(sK0,X1,X2,X2,X3) & X1 != X3 & r2_diraf(sK0,X1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0)) [choice axiom]
96. ? [X1] : (? [X2] : (? [X3] : (~r2_analoaf(sK0,X1,X2,X2,X3) & X1 != X3 & r2_diraf(sK0,X1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(X1,u1_struct_0(sK0))) => (? [X2] : (? [X3] : (~r2_analoaf(sK0,sK1,X2,X2,X3) & sK1 != X3 & r2_diraf(sK0,sK1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) [choice axiom]
97. ? [X2] : (? [X3] : (~r2_analoaf(sK0,sK1,X2,X2,X3) & sK1 != X3 & r2_diraf(sK0,sK1,X2,X2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(X2,u1_struct_0(sK0))) => (? [X3] : (~r2_analoaf(sK0,sK1,sK2,sK2,X3) & sK1 != X3 & r2_diraf(sK0,sK1,sK2,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) [choice axiom]
98. ? [X3] : (~r2_analoaf(sK0,sK1,sK2,sK2,X3) & sK1 != X3 & r2_diraf(sK0,sK1,sK2,sK2,X3) & m1_subset_1(X3,u1_struct_0(sK0))) => (~r2_analoaf(sK0,sK1,sK2,sK2,sK3) & sK1 != sK3 & r2_diraf(sK0,sK1,sK2,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) [choice axiom]
99. (((~r2_analoaf(sK0,sK1,sK2,sK2,sK3) & sK1 != sK3 & r2_diraf(sK0,sK1,sK2,sK2,sK3) & m1_subset_1(sK3,u1_struct_0(sK0))) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_subset_1(sK1,u1_struct_0(sK0))) & l1_analoaf(sK0) & v1_afvect0(sK0) & ~v7_struct_0(sK0) [skolemisation 46,98,97,96,95]
100. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & ((r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4)) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [nnf transformation 74]
101. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (((r2_diraf(X0,X1,X2,X3,X4) | (~r2_analoaf(X0,X1,X2,X4,X3) & ~r2_analoaf(X0,X1,X2,X3,X4))) & (r2_analoaf(X0,X1,X2,X4,X3) | r2_analoaf(X0,X1,X2,X3,X4) | ~r2_diraf(X0,X1,X2,X3,X4))) | ~m1_subset_1(X4,u1_struct_0(X0))) | ~m1_subset_1(X3,u1_struct_0(X0))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_analoaf(X0) | v2_struct_0(X0)) [flattening 100]
112. ~v7_struct_0(sK0) [cnf transformation 99] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
113. v1_afvect0(sK0) [cnf transformation 99] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
114. l1_analoaf(sK0) [cnf transformation 99] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
115. m1_subset_1(sK1,u1_struct_0(sK0)) [cnf transformation 99] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
116. m1_subset_1(sK2,u1_struct_0(sK0)) [cnf transformation 99] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
117. m1_subset_1(sK3,u1_struct_0(sK0)) [cnf transformation 99] {a:0,w:4,wCS:400,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
118. r2_diraf(sK0,sK1,sK2,sK2,sK3) [cnf transformation 99] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
119. sK1 != sK3 [cnf transformation 99] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
120. ~r2_analoaf(sK0,sK1,sK2,sK2,sK3) [cnf transformation 99] {a:0,w:6,wCS:600,goal:1,sine:0,thAx:0,allAx:5,thDist:-5}
123. ~l1_analoaf(X0) | l1_struct_0(X0) [cnf transformation 47] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
127. ~v2_struct_0(X0) | v7_struct_0(X0) | ~l1_struct_0(X0) [cnf transformation 53] {a:0,w:6,wCS:2400,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
140. ~r2_diraf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X1,X2,X4,X3) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | v2_struct_0(X0) [cnf transformation 101] {a:0,w:38,wCS:15200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
144. ~r2_analoaf(X0,X1,X2,X1,X3) | X2 = X3 | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 78] {a:0,w:27,wCS:10800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
145. ~r2_analoaf(X0,X1,X2,X3,X4) | r2_analoaf(X0,X2,X1,X4,X3) | ~m1_subset_1(X4,u1_struct_0(X0)) | ~m1_subset_1(X3,u1_struct_0(X0)) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_analoaf(X0) | ~v1_afvect0(X0) | v7_struct_0(X0) [cnf transformation 80] {a:0,w:34,wCS:13600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
161. 1 <=> r2_analoaf(sK0,sK1,sK2,sK2,sK3) [avatar definition]
163. ~r2_analoaf(sK0,sK1,sK2,sK2,sK3) <- (~1) [avatar component clause 161] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
164. ~1 [avatar split clause 120,161]
166. 2 <=> sK1 = sK3 [avatar definition]
168. sK1 != sK3 <- (~2) [avatar component clause 166] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
169. ~2 [avatar split clause 119,166]
171. 3 <=> r2_diraf(sK0,sK1,sK2,sK2,sK3) [avatar definition]
173. r2_diraf(sK0,sK1,sK2,sK2,sK3) <- (3) [avatar component clause 171] {a:0,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
174. 3 [avatar split clause 118,171]
176. 4 <=> m1_subset_1(sK3,u1_struct_0(sK0)) [avatar definition]
178. m1_subset_1(sK3,u1_struct_0(sK0)) <- (4) [avatar component clause 176] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
179. 4 [avatar split clause 117,176]
181. 5 <=> m1_subset_1(sK2,u1_struct_0(sK0)) [avatar definition]
183. m1_subset_1(sK2,u1_struct_0(sK0)) <- (5) [avatar component clause 181] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
184. 5 [avatar split clause 116,181]
186. 6 <=> m1_subset_1(sK1,u1_struct_0(sK0)) [avatar definition]
188. m1_subset_1(sK1,u1_struct_0(sK0)) <- (6) [avatar component clause 186] {a:0,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
189. 6 [avatar split clause 115,186]
191. 7 <=> l1_analoaf(sK0) [avatar definition]
193. l1_analoaf(sK0) <- (7) [avatar component clause 191] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
194. 7 [avatar split clause 114,191]
196. 8 <=> v1_afvect0(sK0) [avatar definition]
198. v1_afvect0(sK0) <- (8) [avatar component clause 196] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
199. 8 [avatar split clause 113,196]
201. 9 <=> v7_struct_0(sK0) [avatar definition]
203. ~v7_struct_0(sK0) <- (~9) [avatar component clause 201] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
204. ~9 [avatar split clause 112,201]
235. l1_struct_0(sK0) <- (7) [resolution 123,193] {a:1,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
238. 16 <=> l1_struct_0(sK0) [avatar definition]
240. l1_struct_0(sK0) <- (16) [avatar component clause 238] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
241. 16 | ~7 [avatar split clause 235,191,238]
310. 26 <=> v2_struct_0(sK0) [avatar definition]
312. v2_struct_0(sK0) <- (26) [avatar component clause 310] {a:2,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
325. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (3) [resolution 140,173] {a:1,w:32,wCS:3200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
326. r2_analoaf(sK0,sK1,sK2,sK2,sK3) | r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (3) [duplicate literal removal 325] {a:1,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
327. r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3) [subsumption resolution 326,163] {a:1,w:22,wCS:2200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
328. r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4) [subsumption resolution 327,178] {a:1,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
329. r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 5) [subsumption resolution 328,183] {a:1,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
330. r2_analoaf(sK0,sK1,sK2,sK3,sK2) | ~l1_analoaf(sK0) | v2_struct_0(sK0) <- (~1, 3, 4, 5, 6) [subsumption resolution 329,188] {a:1,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
331. r2_analoaf(sK0,sK1,sK2,sK3,sK2) | v2_struct_0(sK0) <- (~1, 3, 4, 5, 6, 7) [subsumption resolution 330,193] {a:1,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
334. 28 <=> r2_analoaf(sK0,sK1,sK2,sK3,sK2) [avatar definition]
336. r2_analoaf(sK0,sK1,sK2,sK3,sK2) <- (28) [avatar component clause 334] {a:1,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
339. 26 | 28 | 1 | ~3 | ~4 | ~5 | ~6 | ~7 [avatar split clause 331,191,186,181,176,171,161,334,310]
344. v7_struct_0(sK0) | ~l1_struct_0(sK0) <- (26) [resolution 312,127] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
347. ~l1_struct_0(sK0) <- (~9, 26) [subsumption resolution 344,203] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
348. $false <- (~9, 16, 26) [subsumption resolution 347,240] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
349. 9 | ~16 | ~26 [avatar contradiction clause 348]
374. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (28) [resolution 336,145] {a:2,w:28,wCS:2800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
379. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (28) [duplicate literal removal 374] {a:2,w:24,wCS:2400,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
380. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (5, 28) [subsumption resolution 379,183] {a:2,w:20,wCS:2000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
381. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (4, 5, 28) [subsumption resolution 380,178] {a:2,w:16,wCS:1600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
382. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (4, 5, 6, 28) [subsumption resolution 381,188] {a:2,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
383. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (4, 5, 6, 7, 28) [subsumption resolution 382,193] {a:2,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
384. r2_analoaf(sK0,sK2,sK1,sK2,sK3) | v7_struct_0(sK0) <- (4, 5, 6, 7, 8, 28) [subsumption resolution 383,198] {a:2,w:8,wCS:800,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
385. r2_analoaf(sK0,sK2,sK1,sK2,sK3) <- (4, 5, 6, 7, 8, ~9, 28) [subsumption resolution 384,203] {a:2,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:7,thDist:-7}
387. 31 <=> r2_analoaf(sK0,sK2,sK1,sK2,sK3) [avatar definition]
389. r2_analoaf(sK0,sK2,sK1,sK2,sK3) <- (31) [avatar component clause 387] {a:2,w:6,wCS:600,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
390. 31 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~28 [avatar split clause 385,334,201,196,191,186,181,176,387]
417. sK1 = sK3 | ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (31) [resolution 389,144] {a:3,w:21,wCS:2100,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
425. ~m1_subset_1(sK3,u1_struct_0(sK0)) | ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, 31) [subsumption resolution 417,168] {a:3,w:18,wCS:1800,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
426. ~m1_subset_1(sK1,u1_struct_0(sK0)) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, 4, 31) [subsumption resolution 425,178] {a:3,w:14,wCS:1400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
427. ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, 4, 6, 31) [subsumption resolution 426,188] {a:3,w:10,wCS:1000,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
428. ~l1_analoaf(sK0) | ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, 4, 5, 6, 31) [subsumption resolution 427,183] {a:3,w:6,wCS:600,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
429. ~v1_afvect0(sK0) | v7_struct_0(sK0) <- (~2, 4, 5, 6, 7, 31) [subsumption resolution 428,193] {a:3,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
430. v7_struct_0(sK0) <- (~2, 4, 5, 6, 7, 8, 31) [subsumption resolution 429,198] {a:3,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
431. $false <- (~2, 4, 5, 6, 7, 8, ~9, 31) [subsumption resolution 430,203] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:8,thDist:-8}
432. 2 | ~4 | ~5 | ~6 | ~7 | ~8 | 9 | ~31 [avatar contradiction clause 431]
443. $false [avatar sat refutation 164,169,174,179,184,189,194,199,204,241,339,349,390,432] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:343,thDist:-343}
% SZS output end Proof for afvect01__l1_afvect01
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 39
% Introduced skolems: 9
% 
% >>> Saturation
% Initial clauses: 48
% Generated clauses: 203
% Active clauses: 65
% Passive clauses: 73
% Final active clauses: 59
% Final passive clauses: 2
% Main loop iterations started: 70
% 
% >>> Simplifying Inferences
% Duplicate literals: 22
% Fw subsumption resolutions: 41
% 
% >>> Deletion Inferences
% Simple tautologies: 13
% Forward subsumptions: 20
% 
% >>> Generating Inferences
% Binary resolution: 50
% Unit resulting resolution: 2
% Forward superposition: 2
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 6
% Split components: 12
% Sat splitting refutations: 4
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 41
% SAT solver unit clauses: 15
% SAT solver binary clauses: 7
% 
% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% Time measurement results:
% superposition: 0.001 s
% other: 0.002 s ( own 0.001 s ) 

