i: [1,209,0,2,1,1,0,0] t41_ordinal2
a: [3,331,0,2,1,209]
s: 331
i: [1,208,0,2,1,1,0,0] t41_ordinal2
a: [3,336,0,2,1,208]
s: 336
i: [1,207,0,2,1,1,0,0] t41_ordinal2
a: [3,341,0,2,1,207]
s: 341
i: [1,210,0,3,1,1,0,0] t41_ordinal2
a: [3,326,0,3,1,210]
s: 326
i: [1,211,0,7,1,1,0,0] t41_ordinal2
a: [3,321,0,7,1,211]
s: 321
i: [1,212,0,2,1,0,0,255] dt_o_0_0_xboole_0
a: [3,346,0,2,1,212]
s: 346
i: [1,213,0,2,1,0,0,2] fc1_xboole_0
a: [3,351,0,2,1,213]
s: 351
i: [1,292,0,2,1,0,0,2] rc2_xboole_0
a: [3,368,0,2,1,292]
s: 368
i: [1,293,0,2,1,0,0,255] rc4_ordinal1
a: [3,373,0,2,1,293]
s: 373
i: [1,294,0,2,1,0,0,1] rc1_ordinal1
a: [3,378,0,2,1,294]
s: 378
i: [1,295,0,2,1,0,0,2] rc1_xboole_0
a: [3,383,0,2,1,295]
s: 383
i: [1,299,0,2,1,0,0,255] rc2_ordinal1
a: [3,388,0,2,1,299]
s: 388
i: [1,298,0,2,1,0,0,255] rc2_ordinal1
a: [3,393,0,2,1,298]
s: 393
i: [1,297,0,2,1,0,0,255] rc2_ordinal1
a: [3,398,0,2,1,297]
s: 398
i: [1,296,0,2,1,0,0,255] rc2_ordinal1
a: [3,403,0,2,1,296]
s: 403
i: [1,301,0,2,1,0,0,255] rc5_ordinal1
a: [3,408,0,2,1,301]
s: 408
i: [1,300,0,2,1,0,0,255] rc5_ordinal1
a: [3,413,0,2,1,300]
s: 413
i: [1,303,0,2,1,0,0,255] rc7_funct_1
a: [3,418,0,2,1,303]
s: 418
i: [1,302,0,2,1,0,0,255] rc7_funct_1
a: [3,423,0,2,1,302]
s: 423
i: [1,305,0,2,1,0,0,255] rc1_funct_1
a: [3,428,0,2,1,305]
s: 428
i: [1,304,0,2,1,0,0,255] rc1_funct_1
a: [3,433,0,2,1,304]
s: 433
i: [1,308,0,2,1,0,0,255] rc5_funct_1
a: [3,438,0,2,1,308]
s: 438
i: [1,307,0,2,1,0,0,255] rc5_funct_1
a: [3,443,0,2,1,307]
s: 443
i: [1,306,0,2,1,0,0,255] rc5_funct_1
a: [3,448,0,2,1,306]
s: 448
i: [1,310,0,2,1,0,0,255] rc2_funct_1
a: [3,453,0,2,1,310]
s: 453
i: [1,309,0,2,1,0,0,255] rc2_funct_1
a: [3,458,0,2,1,309]
s: 458
i: [1,313,0,2,1,0,0,255] rc1_funcop_1
a: [3,463,0,2,1,313]
s: 463
i: [1,312,0,2,1,0,0,255] rc1_funcop_1
a: [3,468,0,2,1,312]
s: 468
i: [1,311,0,2,1,0,0,255] rc1_funcop_1
a: [3,473,0,2,1,311]
s: 473
i: [1,317,0,2,1,0,0,255] rc2_funcop_1
a: [3,478,0,2,1,317]
s: 478
i: [1,316,0,2,1,0,0,255] rc2_funcop_1
a: [3,483,0,2,1,316]
s: 483
i: [1,315,0,2,1,0,0,255] rc2_funcop_1
a: [3,488,0,2,1,315]
s: 488
i: [1,314,0,2,1,0,0,255] rc2_funcop_1
a: [3,493,0,2,1,314]
s: 493
i: [1,215,0,3,1,0,0,3] fc1_subset_1
s: 215
i: [1,280,0,3,1,0,0,3] rc2_subset_1
s: 280
i: [1,283,0,3,1,0,0,255] irreflexivity_r2_xboole_0
s: 283
i: [1,284,0,3,1,0,0,2] reflexivity_r1_tarski
s: 284
i: [1,214,0,3,1,0,0,255] d2_xboole_0
a: [3,356,0,3,1,214]
s: 356
r: 346
i: [1,229,0,4,2,0,0,255] cc6_ordinal1
s: 229
i: [1,235,0,4,2,0,0,255] cc1_ordinal1
s: 235
i: [1,236,0,4,2,0,0,255] cc1_ordinal1
s: 236
d: [2,508,1,2,1,1,40,235,341]
a: [3,516,1,2,1,508]
s: 516
i: [1,240,0,4,2,0,0,255] cc9_ordinal1
s: 240
d: [2,509,1,2,1,1,40,235,336]
a: [3,521,1,2,1,509]
s: 521
i: [1,241,0,4,2,0,0,255] cc7_ordinal1
s: 241
i: [1,242,0,4,2,0,0,255] cc7_funct_1
s: 242
d: [2,510,1,2,1,1,40,235,331]
a: [3,526,1,2,1,510]
s: 526
i: [1,243,0,4,2,0,0,1] cc3_ordinal1
s: 243
d: [2,533,1,2,1,1,40,236,341]
a: [3,541,1,2,1,533]
s: 541
i: [1,244,0,4,2,0,0,255] cc1_funct_1
s: 244
i: [1,278,0,4,1,0,0,2] existence_m1_subset_1
s: 278
d: [2,534,1,2,1,1,40,236,336]
a: [3,546,1,2,1,534]
s: 546
i: [1,282,0,4,1,0,0,255] rc3_subset_1
s: 282
d: [2,535,1,2,1,1,40,236,331]
a: [3,551,1,2,1,535]
s: 551
i: [1,217,0,5,2,0,0,255] rc6_subset_1
s: 217
i: [1,221,0,5,2,0,0,3] rc1_subset_1
s: 221
d: [2,496,1,2,1,1,40,229,373]
a: [3,501,1,2,1,496]
s: 501
i: [1,223,0,5,2,0,0,255] rc5_subset_1
s: 223
d: [2,497,1,2,1,1,40,229,408]
a: [3,506,1,2,1,497]
s: 506
i: [1,224,0,5,2,0,0,255] rc5_subset_1
s: 224
i: [1,245,0,5,2,0,0,2] t6_boole
s: 245
d: [2,511,1,2,1,1,40,235,378]
a: [3,531,1,2,1,511]
s: 531
i: [1,275,0,5,2,0,0,2] t7_boole
s: 275
d: [2,536,1,2,1,1,40,236,378]
a: [3,556,1,2,1,536]
s: 556
i: [1,279,0,5,1,0,0,3] rc2_subset_1
s: 279
i: [1,281,0,5,1,0,0,255] rc3_subset_1
s: 281
d: [2,558,1,2,1,1,40,240,351]
a: [3,564,1,2,1,558]
s: 564
i: [1,287,0,7,3,0,0,1] reflexivity_r1_ordinal1
a: [3,363,0,5,2,287]
s: 363
d: [2,560,1,2,1,1,40,240,383]
a: [3,569,1,2,1,560]
s: 569
i: [1,226,0,6,2,0,0,255] rc4_subset_1
s: 226
i: [1,248,0,6,3,0,0,255] cc2_ordinal1
s: 248
d: [2,571,1,2,1,1,40,241,351]
a: [3,577,1,2,1,571]
s: 577
i: [1,267,0,6,2,0,0,2] d2_ordinal1
s: 267
d: [2,573,1,2,1,1,40,241,383]
a: [3,582,1,2,1,573]
s: 582
i: [1,268,0,6,2,0,0,2] d2_ordinal1
s: 268
i: [1,269,0,6,2,0,0,2] t1_subset
s: 269
d: [2,584,1,2,1,1,40,242,351]
a: [3,590,1,2,1,584]
s: 590
i: [1,285,0,6,2,0,0,255] antisymmetry_r2_xboole_0
s: 285
d: [2,586,1,2,1,1,40,242,383]
a: [3,595,1,2,1,586]
s: 595
i: [1,286,0,6,2,0,0,2] antisymmetry_r2_hidden
s: 286
i: [1,216,0,7,2,0,0,255] rc6_subset_1
s: 216
d: [2,597,1,2,1,1,40,243,351]
a: [3,603,1,2,1,597]
s: 603
i: [1,220,0,7,2,0,0,3] rc1_subset_1
s: 220
d: [2,599,1,2,1,1,40,243,383]
a: [3,608,1,2,1,599]
s: 608
i: [1,222,0,7,2,0,0,255] rc5_subset_1
s: 222
i: [1,225,0,7,2,0,0,255] rc4_subset_1
s: 225
d: [2,610,1,2,1,1,40,244,351]
a: [3,616,1,2,1,610]
s: 616
i: [1,230,0,7,3,0,0,255] cc8_funct_1
s: 230
d: [2,612,1,2,1,1,40,244,383]
a: [3,621,1,2,1,612]
s: 621
i: [1,231,0,7,3,0,0,255] cc8_funct_1
s: 231
i: [1,234,0,7,2,0,0,1] t38_ordinal2
s: 234
d: [2,623,2,2,1,1,40,501,236]
a: [3,628,2,2,1,623]
s: 628
i: [1,238,0,7,3,0,0,1] cc2_ordinal2
s: 238
d: [2,624,2,2,1,1,40,501,235]
a: [3,633,2,2,1,624]
s: 633
i: [1,239,0,7,3,0,0,1] cc5_ordinal1
s: 239
i: [1,272,0,7,2,0,0,2] t3_subset
s: 272
d: [2,635,2,2,1,1,40,506,236]
a: [3,640,2,2,1,635]
s: 640
i: [1,273,0,7,2,0,0,2] t3_subset
s: 273
d: [2,636,2,2,1,1,40,506,235]
a: [3,645,2,2,1,636]
s: 645
i: [1,274,0,7,3,0,0,2] t8_boole
s: 274
i: [1,227,0,8,3,0,0,255] cc10_ordinal1
s: 227
d: [2,664,2,2,1,1,40,603,236]
a: [3,670,2,2,1,664]
s: 670
i: [1,232,0,8,3,0,0,255] cc9_funct_1
s: 232
d: [2,665,2,2,1,1,40,603,235]
a: [3,675,2,2,1,665]
s: 675
i: [1,233,0,8,3,0,0,255] cc5_subset_1
s: 233
i: [1,246,0,8,3,0,0,3] cc1_subset_1
s: 246
d: [2,677,2,2,1,1,40,608,236]
a: [3,682,2,2,1,677]
s: 682
i: [1,251,0,8,4,0,0,255] cc6_funct_1
s: 251
d: [2,678,2,2,1,1,40,608,235]
a: [3,687,2,2,1,678]
s: 687
i: [1,256,0,8,4,0,0,255] cc4_funct_1
s: 256
d: [2,762,0,6,3,0,21,256,244]
s: 762
d: [2,763,1,4,2,1,40,762,433]
a: [3,771,1,2,1,763]
s: 771
i: [1,257,0,8,4,0,0,255] cc5_funct_1
s: 257
d: [2,764,1,4,2,1,40,762,448]
d: [2,777,1,2,1,2,21,764,438]
a: [3,781,1,2,1,777]
s: 781
i: [1,264,0,8,4,0,0,255] cc3_funcop_1
s: 264
d: [2,801,0,6,3,0,21,264,244]
s: 801
d: [2,765,1,4,2,1,40,762,458]
a: [3,786,1,2,1,765]
s: 786
i: [1,266,0,8,3,0,0,2] d2_ordinal1
s: 266
d: [2,766,1,4,2,1,40,762,473]
a: [3,795,1,2,1,766]
s: 795
i: [1,270,0,8,3,0,0,2] t2_subset
s: 270
i: [1,288,0,8,3,0,0,1] dt_k11_ordinal2
s: 288
d: [2,559,1,3,1,0,40,240,280]
s: 559
i: [1,247,0,9,3,0,0,255] cc4_subset_1
s: 247
d: [2,572,1,3,1,0,40,241,280]
s: 572
i: [1,271,0,9,3,0,0,255] d8_xboole_0
s: 271
i: [1,277,0,9,3,0,0,2] t5_subset
s: 277
d: [2,585,1,3,1,0,40,242,280]
s: 585
i: [1,228,0,10,4,0,0,255] t11_ordinal1
s: 228
d: [2,598,1,3,1,0,40,243,280]
s: 598
i: [1,265,0,10,4,0,0,255] cc3_funct_1
s: 265
i: [1,276,0,10,3,0,0,2] t4_subset
s: 276
d: [2,611,1,3,1,0,40,244,280]
s: 611
i: [1,289,0,10,4,0,0,1] connectedness_r1_ordinal1
s: 289
d: [2,868,1,15,3,1,40,289,321]
a: [3,874,1,4,1,868]
s: 874
d: [2,884,2,0,0,3,52,336,331,874,288]
e: 884
d: [2,885,2,4,2,1,40,874,288]
d: [2,887,2,2,1,2,21,885,331]
d: [2,888,2,0,0,3,21,887,336]
e: 888
r: 874
i: [1,290,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 290
i: [1,291,0,10,4,0,0,1] redefinition_r1_ordinal1
s: 291
d: [2,890,1,7,3,1,40,290,326]
d: [2,895,1,5,2,2,21,890,336]
d: [2,896,1,3,1,3,21,895,341]
a: [3,900,1,3,1,896]
s: 900
i: [1,218,0,11,4,0,0,255] cc3_subset_1
s: 218
d: [2,905,2,6,2,1,40,900,271]
a: [3,909,2,3,1,905]
s: 909
i: [1,219,0,11,4,0,0,255] cc2_subset_1
s: 219
i: [1,237,0,19,6,0,0,1] t40_ordinal2
s: 237
a: [3,912,2,3,1,905]
s: 912
a: [3,873,1,4,1,868]
s: 873
d: [2,928,3,7,3,1,40,909,228]
d: [2,930,3,5,2,2,21,928,336]
d: [2,931,3,3,1,3,21,930,516]
a: [3,935,3,3,1,931]
s: 935
a: [3,882,1,7,1,868]
s: 882
d: [2,658,1,4,2,0,40,267,275]
s: 658
d: [2,965,4,2,1,1,40,935,275]
a: [3,979,4,2,1,965]
s: 979
d: [2,718,1,4,1,0,40,272,279]
s: 718
d: [2,929,3,3,1,1,40,909,285]
a: [3,940,3,3,1,929]
s: 940
d: [2,719,1,4,1,0,40,272,281]
s: 719
d: [2,694,1,5,2,0,40,230,278]
s: 694
d: [2,963,4,3,1,1,40,935,286]
a: [3,969,4,3,1,963]
s: 969
d: [2,701,1,5,2,0,40,231,278]
s: 701
d: [2,964,4,3,1,1,40,935,269]
a: [3,974,4,3,1,964]
s: 974
d: [2,710,1,5,2,0,40,238,278]
s: 710
d: [2,717,1,5,1,0,40,272,278]
s: 717
d: [2,1000,5,4,2,1,40,974,231]
a: [3,1005,5,2,1,1000]
s: 1005
d: [2,731,1,5,2,0,40,227,279]
s: 731
d: [2,950,2,4,1,1,40,873,236]
a: [3,955,2,4,1,950]
s: 955
d: [2,732,1,5,2,0,40,227,281]
s: 732
d: [2,741,1,5,2,0,40,232,279]
s: 741
d: [2,951,2,4,1,1,40,873,235]
a: [3,960,2,4,1,951]
s: 960
d: [2,742,1,5,2,0,40,232,281]
s: 742
d: [2,981,2,15,3,1,40,882,290]
d: [2,982,2,11,2,2,21,981,873]
a: [3,986,2,7,1,982]
s: 986
d: [2,750,1,5,2,0,40,233,279]
s: 750
d: [2,751,1,5,2,0,40,233,281]
s: 751
d: [2,1022,3,14,2,1,40,986,271]
a: [3,1026,3,7,1,1022]
s: 1026
d: [2,759,1,5,2,0,40,246,281]
s: 759
d: [2,1036,4,15,3,1,40,1026,228]
d: [2,1038,4,11,2,2,21,1036,960]
a: [3,1042,4,7,1,1038]
s: 1042
d: [2,819,1,8,2,0,40,270,279]
d: [2,826,1,5,1,0,21,819,215]
s: 826
d: [2,820,1,8,2,0,40,270,281]
d: [2,827,1,5,1,0,21,820,215]
s: 827
d: [2,1058,5,4,1,1,40,1042,275]
a: [3,1072,5,4,1,1058]
s: 1072
d: [2,922,1,9,3,0,40,218,281]
d: [2,927,1,5,2,0,21,922,282]
s: 927
d: [2,1037,4,7,1,1,40,1026,285]
a: [3,1047,4,7,1,1037]
s: 1047
d: [2,661,1,6,2,0,40,269,267]
s: 661
d: [2,663,1,6,2,0,40,286,267]
s: 663
d: [2,1056,5,7,1,1,40,1042,286]
a: [3,1062,5,7,1,1056]
s: 1062
d: [2,695,1,6,2,0,40,230,279]
s: 695
d: [2,1097,6,12,5,1,40,1062,237]
d: [2,1098,6,9,4,2,21,1097,935]
d: [2,1099,6,7,3,3,21,1098,336]
d: [2,1100,6,5,2,4,21,1099,341]
d: [2,1101,6,3,1,5,21,1100,331]
a: [3,1105,6,3,1,1101]
s: 1105
r: 321
r: 331
r: 526
r: 551
r: 873
r: 882
r: 955
r: 960
r: 986
r: 1026
r: 1042
r: 1047
r: 1062
r: 1072
d: [2,1107,0,7,1,2,23,321,1105]
a: [3,1124,0,7,1,1107]
s: 1124
d: [2,1111,1,4,1,2,23,873,1105]
a: [3,1129,1,4,1,1111]
s: 1129
d: [2,1113,2,4,1,2,23,955,1105]
a: [3,1139,2,4,1,1113]
s: 1139
d: [2,1177,1,7,2,1,44,1124,234]
d: [2,1185,1,5,1,2,21,1177,341]
a: [3,1189,1,5,1,1185]
s: 1189
d: [2,1208,2,5,2,1,44,1189,234]
d: [2,1216,2,3,1,2,21,1208,336]
a: [3,1220,2,3,1,1216]
s: 1220
d: [2,1222,3,0,0,3,52,603,1220,363]
e: 1222
d: [2,1223,3,0,0,2,52,603,603,284,1220,291]
e: 1223
d: [2,1224,3,0,0,2,52,603,603,1220,1220,289]
e: 1224
d: [2,1225,3,0,0,2,52,603,603,1220,1220,289]
e: 1225
d: [2,1226,3,2,1,2,40,1220,363]
d: [2,1234,3,0,0,3,21,1226,603]
e: 1234
r: 1105
r: 1124
r: 1189
r: 1220
s: 331
s: 321
s: 526
a: [3,1123,0,7,1,1107]
s: 1123
s: 551
s: 873
a: [3,878,1,4,1,868]
s: 878
d: [2,1261,2,0,0,3,52,341,331,878,288]
e: 1261
d: [2,1262,2,4,2,1,40,878,288]
d: [2,1264,2,2,1,2,21,1262,331]
d: [2,1265,2,0,0,3,21,1264,341]
e: 1265
d: [2,1240,1,15,3,1,40,321,289]
d: [2,1258,2,4,1,1,40,873,235]
f: 1265,1258,1240,1225,1216,1185,1107,1101,1056,1038,1022,982,931,905,896,888,597,508,213,207,208,209,210,211
% Refutation found. Thanks to Tanya!
% SZS status Theorem for ordinal2__t41_ordinal2
% SZS output start Proof for ordinal2__t41_ordinal2
1. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(X0,X1) => r1_ordinal1(k11_ordinal2(X0,X2),k11_ordinal2(X1,X2)))))) [input t41_ordinal2]
2. ~! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r1_ordinal1(X0,X1) => r1_ordinal1(k11_ordinal2(X0,X2),k11_ordinal2(X1,X2)))))) [negated conjecture 1]
3. ! [X0,X1] : (r2_hidden(X0,X1) => ~r2_hidden(X1,X0)) [input antisymmetry_r2_hidden]
7. ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0))) [input cc1_ordinal1]
16. ! [X0] : (v1_xboole_0(X0) => v3_ordinal1(X0)) [input cc3_ordinal1]
31. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1))) [input connectedness_r1_ordinal1]
34. ! [X0] : ! [X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [input d8_xboole_0]
35. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k11_ordinal2(X0,X1))) [input dt_k11_ordinal2]
42. v1_xboole_0(k1_xboole_0) [input fc1_xboole_0]
62. ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1))) [input redefinition_r1_ordinal1]
65. ! [X0] : (v1_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_xboole_0(X0,X1) => r2_hidden(X0,X1)))) [input t11_ordinal1]
68. ! [X0] : (v3_ordinal1(X0) => k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0)) [input t38_ordinal2]
70. ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(X1,X2) => (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0))))) [input t40_ordinal2]
78. ! [X0,X1] : (r2_xboole_0(X0,X1) <=> (X0 != X1 & r1_tarski(X0,X1))) [flattening 34]
86. ! [X0,X1] : ((X0 != X1 & r1_tarski(X0,X1)) => r2_xboole_0(X0,X1)) [unused predicate definition removal 78]
91. ? [X0] : (? [X1] : (? [X2] : ((~r1_ordinal1(k11_ordinal2(X0,X2),k11_ordinal2(X1,X2)) & r1_ordinal1(X0,X1)) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [ennf transformation 2]
92. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(k11_ordinal2(X0,X2),k11_ordinal2(X1,X2)) & r1_ordinal1(X0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) [flattening 91]
102. ! [X0] : (! [X1] : ((r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1)) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [ennf transformation 65]
103. ! [X0] : (! [X1] : (r2_hidden(X0,X1) | ~r2_xboole_0(X0,X1) | ~v3_ordinal1(X1)) | ~v1_ordinal1(X0)) [flattening 102]
108. ! [X0] : (k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0)) [ennf transformation 68]
109. ! [X0] : ((v2_ordinal1(X0) & v1_ordinal1(X0)) | ~v3_ordinal1(X0)) [ennf transformation 7]
110. ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0) | ~r2_hidden(X1,X2)) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [ennf transformation 70]
111. ! [X0] : (! [X1] : (! [X2] : (r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0)) [flattening 110]
117. ! [X0] : (v3_ordinal1(X0) | ~v1_xboole_0(X0)) [ennf transformation 16]
142. ! [X0,X1] : (r2_xboole_0(X0,X1) | (X0 = X1 | ~r1_tarski(X0,X1))) [ennf transformation 86]
143. ! [X0,X1] : (r2_xboole_0(X0,X1) | X0 = X1 | ~r1_tarski(X0,X1)) [flattening 142]
150. ! [X0,X1] : (~r2_hidden(X1,X0) | ~r2_hidden(X0,X1)) [ennf transformation 3]
153. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 35]
154. ! [X0,X1] : (v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 153]
155. ! [X0,X1] : ((r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 31]
156. ! [X0,X1] : (r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 155]
157. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0))) [ennf transformation 62]
158. ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [flattening 157]
159. ? [X0] : (? [X1] : (? [X2] : (~r1_ordinal1(k11_ordinal2(X0,X2),k11_ordinal2(X1,X2)) & r1_ordinal1(X0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(X0)) => (? [X1] : (? [X2] : (~r1_ordinal1(k11_ordinal2(sK0,X2),k11_ordinal2(X1,X2)) & r1_ordinal1(sK0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) & v3_ordinal1(sK0)) [choice axiom]
160. ? [X1] : (? [X2] : (~r1_ordinal1(k11_ordinal2(sK0,X2),k11_ordinal2(X1,X2)) & r1_ordinal1(sK0,X1) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (~r1_ordinal1(k11_ordinal2(sK0,X2),k11_ordinal2(sK1,X2)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(X2)) & v3_ordinal1(sK1)) [choice axiom]
161. ? [X2] : (~r1_ordinal1(k11_ordinal2(sK0,X2),k11_ordinal2(sK1,X2)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(X2)) => (~r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(sK2)) [choice axiom]
162. ((~r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) & r1_ordinal1(sK0,sK1) & v3_ordinal1(sK2)) & v3_ordinal1(sK1)) & v3_ordinal1(sK0) [skolemisation 92,161,160,159]
182. ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) [nnf transformation 158]
207. v3_ordinal1(sK0) [cnf transformation 162] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
208. v3_ordinal1(sK1) [cnf transformation 162] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
209. v3_ordinal1(sK2) [cnf transformation 162] {a:0,w:2,wCS:200,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
210. r1_ordinal1(sK0,sK1) [cnf transformation 162] {a:0,w:3,wCS:300,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
211. ~r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) [cnf transformation 162] {a:0,w:7,wCS:700,goal:1,sine:0,thAx:0,allAx:4,thDist:-4}
213. v1_xboole_0(k1_xboole_0) [cnf transformation 42] {a:0,w:2,wCS:800,sine:2,thAx:0,allAx:1,thDist:-1}
228. ~r2_xboole_0(X0,X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v1_ordinal1(X0) [cnf transformation 103] {a:0,w:10,wCS:4000,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
234. k1_xboole_0 = k11_ordinal2(X0,k1_xboole_0) | ~v3_ordinal1(X0) [cnf transformation 108] {a:0,w:7,wCS:2800,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
235. ~v3_ordinal1(X0) | v1_ordinal1(X0) [cnf transformation 109] {a:0,w:4,wCS:1600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
237. r2_hidden(k11_ordinal2(X1,X0),k11_ordinal2(X2,X0)) | k1_xboole_0 = X0 | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 111] {a:0,w:19,wCS:7600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
243. ~v1_xboole_0(X0) | v3_ordinal1(X0) [cnf transformation 117] {a:0,w:4,wCS:1600,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
271. ~r1_tarski(X0,X1) | X0 = X1 | r2_xboole_0(X0,X1) [cnf transformation 143] {a:0,w:9,wCS:3600,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
286. ~r2_hidden(X1,X0) | ~r2_hidden(X0,X1) [cnf transformation 150] {a:0,w:6,wCS:2400,nSel:1,sine:2,thAx:0,allAx:1,thDist:-1}
288. v3_ordinal1(k11_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 154] {a:0,w:8,wCS:3200,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
289. r1_ordinal1(X1,X0) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 156] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
290. ~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) [cnf transformation 182] {a:0,w:10,wCS:4000,nSel:1,sine:1,thAx:0,allAx:1,thDist:-1}
319. 1 <=> r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) [avatar definition]
321. ~r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) <- (~1) [avatar component clause 319] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
322. ~1 [avatar split clause 211,319]
324. 2 <=> r1_ordinal1(sK0,sK1) [avatar definition]
326. r1_ordinal1(sK0,sK1) <- (2) [avatar component clause 324] {a:0,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
327. 2 [avatar split clause 210,324]
329. 3 <=> v3_ordinal1(sK2) [avatar definition]
331. v3_ordinal1(sK2) <- (3) [avatar component clause 329] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
332. 3 [avatar split clause 209,329]
334. 4 <=> v3_ordinal1(sK1) [avatar definition]
336. v3_ordinal1(sK1) <- (4) [avatar component clause 334] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
337. 4 [avatar split clause 208,334]
339. 5 <=> v3_ordinal1(sK0) [avatar definition]
341. v3_ordinal1(sK0) <- (5) [avatar component clause 339] {a:0,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
342. 5 [avatar split clause 207,339]
349. 7 <=> v1_xboole_0(k1_xboole_0) [avatar definition]
351. v1_xboole_0(k1_xboole_0) <- (7) [avatar component clause 349] {a:0,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:1,thDist:-1}
352. 7 [avatar split clause 213,349]
508. v1_ordinal1(sK0) <- (5) [resolution 235,341] {a:1,w:2,wCS:200,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
514. 39 <=> v1_ordinal1(sK0) [avatar definition]
516. v1_ordinal1(sK0) <- (39) [avatar component clause 514] {a:1,w:2,wCS:200,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
517. 39 | ~5 [avatar split clause 508,339,514]
597. v3_ordinal1(k1_xboole_0) <- (7) [resolution 243,351] {a:1,w:2,wCS:800,sine:1,thAx:0,allAx:2,thDist:-2}
601. 53 <=> v3_ordinal1(k1_xboole_0) [avatar definition]
603. v3_ordinal1(k1_xboole_0) <- (53) [avatar component clause 601] {a:1,w:2,wCS:800,nSel:1,sine:255,thAx:0,allAx:2,thDist:-2}
604. 53 | ~7 [avatar split clause 597,349,601]
872. 73 <=> v3_ordinal1(k11_ordinal2(sK1,sK2)) [avatar definition]
873. v3_ordinal1(k11_ordinal2(sK1,sK2)) <- (73) [avatar component clause 872] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
874. ~v3_ordinal1(k11_ordinal2(sK1,sK2)) <- (~73) [avatar component clause 872] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
876. 74 <=> v3_ordinal1(k11_ordinal2(sK0,sK2)) [avatar definition]
878. ~v3_ordinal1(k11_ordinal2(sK0,sK2)) <- (~74) [avatar component clause 876] {a:1,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
880. 75 <=> r1_ordinal1(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [avatar definition]
882. r1_ordinal1(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) <- (75) [avatar component clause 880] {a:1,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
885. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK1) <- (~73) [resolution 874,288] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
887. ~v3_ordinal1(sK1) <- (3, ~73) [subsumption resolution 885,331] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
888. $false <- (3, 4, ~73) [subsumption resolution 887,336] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
889. ~3 | ~4 | 73 [avatar contradiction clause 888]
890. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) <- (2) [resolution 290,326] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
895. r1_tarski(sK0,sK1) | ~v3_ordinal1(sK0) <- (2, 4) [subsumption resolution 890,336] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
896. r1_tarski(sK0,sK1) <- (2, 4, 5) [subsumption resolution 895,341] {a:1,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
898. 76 <=> r1_tarski(sK0,sK1) [avatar definition]
900. r1_tarski(sK0,sK1) <- (76) [avatar component clause 898] {a:1,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
901. 76 | ~2 | ~4 | ~5 [avatar split clause 896,339,334,324,898]
905. sK0 = sK1 | r2_xboole_0(sK0,sK1) <- (76) [resolution 900,271] {a:2,w:6,wCS:600,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
907. 77 <=> r2_xboole_0(sK0,sK1) [avatar definition]
909. r2_xboole_0(sK0,sK1) <- (77) [avatar component clause 907] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
911. 78 <=> sK0 = sK1 [avatar definition]
914. 77 | 78 | ~76 [avatar split clause 905,898,911,907]
928. r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v1_ordinal1(sK0) <- (77) [resolution 909,228] {a:3,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
930. r2_hidden(sK0,sK1) | ~v1_ordinal1(sK0) <- (4, 77) [subsumption resolution 928,336] {a:3,w:5,wCS:500,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
931. r2_hidden(sK0,sK1) <- (4, 39, 77) [subsumption resolution 930,516] {a:3,w:3,wCS:300,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
933. 79 <=> r2_hidden(sK0,sK1) [avatar definition]
935. r2_hidden(sK0,sK1) <- (79) [avatar component clause 933] {a:3,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
936. 79 | ~4 | ~39 | ~77 [avatar split clause 931,907,514,334,933]
943. sK0 != sK1 | r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) | ~r1_ordinal1(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [theory tautology sat conflict] {a:0,w:17,wCS:6800,sine:255,thAx:0,allAx:1,thDist:-1}
958. 82 <=> v1_ordinal1(k11_ordinal2(sK1,sK2)) [avatar definition]
960. v1_ordinal1(k11_ordinal2(sK1,sK2)) <- (82) [avatar component clause 958] {a:2,w:4,wCS:400,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
981. r1_tarski(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK1,sK2)) <- (75) [resolution 882,290] {a:2,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
982. r1_tarski(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK0,sK2)) <- (73, 75) [subsumption resolution 981,873] {a:2,w:11,wCS:1100,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
984. 86 <=> r1_tarski(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [avatar definition]
986. r1_tarski(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) <- (86) [avatar component clause 984] {a:2,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
987. ~74 | 86 | ~73 | ~75 [avatar split clause 982,880,872,984,876]
1022. k11_ordinal2(sK0,sK2) = k11_ordinal2(sK1,sK2) | r2_xboole_0(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) <- (86) [resolution 986,271] {a:3,w:14,wCS:1400,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1024. 90 <=> r2_xboole_0(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [avatar definition]
1026. r2_xboole_0(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) <- (90) [avatar component clause 1024] {a:3,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:7,thDist:-7}
1028. 91 <=> k11_ordinal2(sK0,sK2) = k11_ordinal2(sK1,sK2) [avatar definition]
1031. 90 | 91 | ~86 [avatar split clause 1022,984,1028,1024]
1032. k11_ordinal2(sK0,sK2) != k11_ordinal2(sK1,sK2) | r1_ordinal1(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) | ~r1_ordinal1(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [theory tautology sat conflict] {a:0,w:21,wCS:8400,sine:255,thAx:0,allAx:1,thDist:-1}
1036. r2_hidden(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK0,sK2)) | ~v1_ordinal1(k11_ordinal2(sK1,sK2)) <- (90) [resolution 1026,228] {a:4,w:15,wCS:1500,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1038. r2_hidden(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK0,sK2)) <- (82, 90) [subsumption resolution 1036,960] {a:4,w:11,wCS:1100,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1040. 92 <=> r2_hidden(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) [avatar definition]
1042. r2_hidden(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) <- (92) [avatar component clause 1040] {a:4,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:8,thDist:-8}
1043. ~74 | 92 | ~82 | ~90 [avatar split clause 1038,1024,958,1040,876]
1056. ~r2_hidden(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) <- (92) [resolution 1042,286] {a:5,w:7,wCS:700,goal:1,sine:2,thAx:0,allAx:9,thDist:-9}
1060. 94 <=> r2_hidden(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) [avatar definition]
1062. ~r2_hidden(k11_ordinal2(sK0,sK2),k11_ordinal2(sK1,sK2)) <- (~94) [avatar component clause 1060] {a:5,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:9,thDist:-9}
1063. ~94 | ~92 [avatar split clause 1056,1040,1060]
1097. k1_xboole_0 = sK2 | ~r2_hidden(sK0,sK1) | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (~94) [resolution 1062,237] {a:6,w:12,wCS:1200,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1098. k1_xboole_0 = sK2 | ~v3_ordinal1(sK1) | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (79, ~94) [subsumption resolution 1097,935] {a:6,w:9,wCS:900,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1099. k1_xboole_0 = sK2 | ~v3_ordinal1(sK0) | ~v3_ordinal1(sK2) <- (4, 79, ~94) [subsumption resolution 1098,336] {a:6,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1100. k1_xboole_0 = sK2 | ~v3_ordinal1(sK2) <- (4, 5, 79, ~94) [subsumption resolution 1099,341] {a:6,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1101. k1_xboole_0 = sK2 <- (3, 4, 5, 79, ~94) [subsumption resolution 1100,331] {a:6,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:10,thDist:-10}
1103. 97 <=> k1_xboole_0 = sK2 [avatar definition]
1105. k1_xboole_0 = sK2 <- (97) [avatar component clause 1103] {a:6,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:10,thDist:-10}
1106. 97 | ~3 | ~4 | ~5 | ~79 | 94 [avatar split clause 1101,1060,933,339,334,329,1103]
1107. ~r1_ordinal1(k11_ordinal2(sK0,k1_xboole_0),k11_ordinal2(sK1,k1_xboole_0)) <- (~1, 97) [backward demodulation 321,1105] {a:0,w:7,wCS:700,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1122. 98 <=> r1_ordinal1(k11_ordinal2(sK0,k1_xboole_0),k11_ordinal2(sK1,k1_xboole_0)) [avatar definition]
1124. ~r1_ordinal1(k11_ordinal2(sK0,k1_xboole_0),k11_ordinal2(sK1,k1_xboole_0)) <- (~98) [avatar component clause 1122] {a:0,w:7,wCS:700,nSel:1,goal:1,sine:255,thAx:0,allAx:4,thDist:-4}
1125. ~98 | 1 | ~97 [avatar split clause 1107,1103,319,1122]
1177. ~r1_ordinal1(k1_xboole_0,k11_ordinal2(sK1,k1_xboole_0)) | ~v3_ordinal1(sK0) <- (~98) [superposition 1124,234] {a:1,w:7,wCS:700,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1185. ~r1_ordinal1(k1_xboole_0,k11_ordinal2(sK1,k1_xboole_0)) <- (5, ~98) [subsumption resolution 1177,341] {a:1,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1187. 110 <=> r1_ordinal1(k1_xboole_0,k11_ordinal2(sK1,k1_xboole_0)) [avatar definition]
1189. ~r1_ordinal1(k1_xboole_0,k11_ordinal2(sK1,k1_xboole_0)) <- (~110) [avatar component clause 1187] {a:1,w:5,wCS:500,nSel:1,goal:1,sine:255,thAx:0,allAx:5,thDist:-5}
1190. ~110 | ~5 | 98 [avatar split clause 1185,1122,339,1187]
1208. ~r1_ordinal1(k1_xboole_0,k1_xboole_0) | ~v3_ordinal1(sK1) <- (~110) [superposition 1189,234] {a:2,w:5,wCS:500,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1216. ~r1_ordinal1(k1_xboole_0,k1_xboole_0) <- (4, ~110) [subsumption resolution 1208,336] {a:2,w:3,wCS:300,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1218. 113 <=> r1_ordinal1(k1_xboole_0,k1_xboole_0) [avatar definition]
1220. ~r1_ordinal1(k1_xboole_0,k1_xboole_0) <- (~113) [avatar component clause 1218] {a:2,w:3,wCS:300,nSel:1,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1221. ~113 | ~4 | 110 [avatar split clause 1216,1187,334,1218]
1225. $false <- (53, ~113) [unit resulting resolution 603,603,1220,1220,289] {a:3,w:0,goal:1,sine:1,thAx:0,allAx:17,thDist:-17}
1233. ~53 | 113 [avatar contradiction clause 1225]
1240. r1_ordinal1(k11_ordinal2(sK1,sK2),k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK0,sK2)) | ~v3_ordinal1(k11_ordinal2(sK1,sK2)) <- (~1) [resolution 321,289] {a:1,w:15,wCS:1500,goal:1,sine:1,thAx:0,allAx:5,thDist:-5}
1241. ~73 | ~74 | 75 | 1 [avatar split clause 1240,319,880,876,872]
1258. v1_ordinal1(k11_ordinal2(sK1,sK2)) <- (73) [resolution 873,235] {a:2,w:4,wCS:400,goal:1,sine:255,thAx:0,allAx:6,thDist:-6}
1260. 82 | ~73 [avatar split clause 1258,872,958]
1262. ~v3_ordinal1(sK2) | ~v3_ordinal1(sK0) <- (~74) [resolution 878,288] {a:2,w:4,wCS:400,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1264. ~v3_ordinal1(sK0) <- (3, ~74) [subsumption resolution 1262,331] {a:2,w:2,wCS:200,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1265. $false <- (3, 5, ~74) [subsumption resolution 1264,341] {a:2,w:0,goal:1,sine:1,thAx:0,allAx:6,thDist:-6}
1266. ~3 | ~5 | 74 [avatar contradiction clause 1265]
1267. $false [avatar sat refutation 322,327,332,337,342,352,517,604,889,901,914,936,943,987,1031,1032,1043,1063,1106,1125,1190,1221,1233,1241,1260,1266] {a:0,w:0,goal:1,sine:0,thAx:0,allAx:849,thDist:-849}
% SZS output end Proof for ordinal2__t41_ordinal2
% ------------------------------
% Version: Vampire 4.5.1 (commit 226f88162 on 2021-09-12 20:13:44 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 74
% >>> Preprocessing
% Introduced skolems: 23
% Pure predicates: 3
% Unused predicate definitions: 1
% 
% >>> Saturation
% Initial clauses: 111
% Generated clauses: 601
% Active clauses: 184
% Passive clauses: 349
% Final active clauses: 164
% Final passive clauses: 154
% Main loop iterations started: 191
% 
% >>> Simplifying Inferences
% Duplicate literals: 11
% Fw subsumption resolutions: 37
% Bw demodulations: 15
% 
% >>> Deletion Inferences
% Simple tautologies: 27
% Equational tautologies: 1
% Forward subsumptions: 59
% 
% >>> Generating Inferences
% Binary resolution: 273
% Unit resulting resolution: 6
% Factoring: 1
% Forward superposition: 22
% Backward superposition: 1
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 15
% Split components: 33
% Sat splitting refutations: 9
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 130
% SAT solver unit clauses: 34
% SAT solver binary clauses: 50
% 
% Memory used [KB]: 5756
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.001 s
% SAT solver time: 0.001 s
% non unit clause index maintenance: 0.001 s
% superposition: 0.001 s
% other: 0.006 s ( own 0.002 s ) 

